<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 396598  -->
  <question type="stack">
    <name>
      <text>Conversion Base 16 (Hexadécimal) vers Base 10</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<div style="background:#4c1d95;border-left:5px solid #7c3aed;border-radius:0 8px 8px 0;padding:10px 16px;margin-bottom:12px;display:flex;align-items:center;gap:10px;flex-wrap:wrap;">
        <strong style="font-weight:800;color:#fff;font-size:.95rem;">Numération — Conversion Décimale</strong>
        <span style="background:#7c3aed;color:#fff;padding:2px 9px;border-radius:20px;font-size:.78rem;font-weight:700;">/ 1 pt</span>
        <span style="background:#ffffff;color:#7c3aed;border:1px solid #7c3aed;padding:2px 9px;border-radius:20px;font-size:.75rem;font-weight:600;">🔢 Base 10</span>
      </div>
      
      <p><strong>Consigne :</strong> Convertissez le nombre hexadécimal suivant en base 10 (décimal) :</p>
      <p style="text-align: center; font-size: 1.5em; font-weight: bold; margin: 15px 0;">
        \( {@hex_str@} _{16} \) = [[input:ans1]] \( _{10} \)
      </p>
      <div style="margin-bottom: 15px;">[[validation:ans1]]</div>
      
      <div style="background:#f8f9fa;border:1px solid #dee2e6;padding:15px;border-radius:8px;margin:15px 0;">
        <strong>Conseils pour la saisie :</strong>
        <ul style="margin:8px 0 0 20px;line-height:1.7">
          <li>Saisissez uniquement le nombre entier résultat (sans unité).</li>
          <li>Rappel : A=10, B=11, C=12, D=13, E=14, F=15.</li>
        </ul>
      </div>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<div style="margin-top:20px; padding:15px; background:#f8fafc; border:1px solid #cbd5e1; border-radius:8px; box-shadow:0 2px 4px rgba(0,0,0,0.05);">
        <div style="font-weight:bold; color:#4c1d95; margin-bottom:10px; display:flex; align-items:center; gap:8px;">
          <span style="font-size:1.2rem;">🔑</span> Réponse attendue
        </div>
        <div style="margin-bottom:8px;font-size:.9rem;border-bottom:1px dashed #e2e8f0;padding-bottom:6px;">
          <span style="font-weight:bold;color:#4c1d95;">Conversion :</span>
          <p>Pour convertir de l'hexadécimal vers le décimal, on multiplie chaque chiffre par la puissance de 16 correspondant à son rang (en partant de la droite : \(16^0, 16^1, 16^2...\) ), puis on additionne le tout.</p>
          <p>{@hex_str@}<sub>16</sub> = {@ta_dec@}<sub>10</sub></p>
        </div>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.3333333</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text></text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* 1. Liste des symboles hexadécimaux */
sym_hex: ["0","1","2","3","4","5","6","7","8","9","A","B","C","D","E","F"];

/* 2. Génère une liste de 2 ou 3 chiffres hexadécimaux aléatoires (valeurs de 0 à 15) */
len_hex: 2 + rand(2);
hex_list: makelist(rand(16), i, 1, len_hex);

/* 3. Convertit la liste d'entiers en chaîne pour l'affichage (rapide via map) */
hex_str: simplode(map(lambda([x], sym_hex[x+1]), hex_list));

/* 4. Calcule la vraie réponse décimale (Méthode de Horner : ultra rapide) */
ta_dec: block([res: 0],
    for d in hex_list do (
        res: res * 16 + d
    ),
    res
);

/* 5. Erreur classique : l'élève a multiplié par 10 au lieu de 16 */
ta_err_10: block([res: 0],
    for d in hex_list do (
        res: res * 10 + d
    ),
    res
);

/* 6. Erreur classique : l'élève a lu les chiffres de droite à gauche */
ta_err_rev: block([res: 0],
    for d in reverse(hex_list) do (
        res: res * 16 + d
    ),
    res
);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="html">
      <text>Hexa: {@hex_str@} | Décimal: {@ta_dec@}</text>
    </questionnote>
    <questiondescription format="html">
      <text>Généré avec format strict STACK</text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta_dec</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>2</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Réponse exacte</description>
        <answertest>EqualComAss</answertest>
        <sans>ans1</sans>
        <tans>ta_dec</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-0-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left:4px solid #15803d;padding:8px 12px;background:#f0fdf4;border-radius:4px;margin-bottom:10px;">✅ <strong>Correct !</strong> La conversion est bonne.</div>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-0-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Multiplication par 10 au lieu de 16</description>
        <answertest>EqualComAss</answertest>
        <sans>ans1</sans>
        <tans>ta_err_10</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>-</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left:4px solid #dc2626;padding:8px 12px;background:#fef2f2;border-radius:4px;margin-bottom:10px;">❌ <strong>Erreur de base.</strong> Vous avez calculé comme si le nombre était en base 10 (en multipliant par 10, 100, 1000...). C'est un nombre en <strong>base 16</strong>, il faut utiliser des puissances de 16 !</div>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Poids des chiffres inversés</description>
        <answertest>EqualComAss</answertest>
        <sans>ans1</sans>
        <tans>ta_err_rev</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>-</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left:4px solid #dc2626;padding:8px 12px;background:#fef2f2;border-radius:4px;margin-bottom:10px;">❌ <strong>Erreur de sens de lecture.</strong> Vous avez attribué le poids le plus fort au chiffre de droite (ou inversement). Le chiffre le plus à droite vaut toujours \(16^0\), le suivant vers la gauche \(16^1\), etc.</div>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<div style="border-left:4px solid #dc2626;padding:8px 12px;background:#fef2f2;border-radius:4px;margin-bottom:10px;">❌ <strong>Incorrect.</strong> Votre résultat est faux. Vérifiez la valeur des lettres (A=10, B=11...) et vos calculs de puissances de 16.</div>]]></text>
        </falsefeedback>
      </node>
    </prt>
  </question>

</quiz>