<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 49  -->
  <question type="stack">
    <name>
      <text>marche mais pas de correction (copie)</text>
    </name>
    <questiontext format="plain_text">
      <text><![CDATA[<!-- ENONCE-START --><div style="margin-bottom:14px;">hghhghg</div><!-- ENONCE-END -->
[[jsxgraph width="100%" height="600px" input-ref-ans1_html="refAns1" input-ref-ans1_flag="refAns2"]]

// --- UTILITAIRE ---
function setRef(ref, value) {
  var el = document.getElementById(ref);
  if (el) { el.value = value; el.dispatchEvent(new Event('change')); }
}

// --- CONFIGURATION ---
var box = document.getElementById('jxgbox');
box.style.cssText = 'border:none; background:#f9f9f9; padding:15px; border-radius:8px; font-family:sans-serif; box-sizing:border-box; width:100%; overflow:hidden;';
box.innerHTML = '';

var txtLatex = document.createElement('textarea');
txtLatex.rows = 3;
txtLatex.style.cssText = 'width:100%; padding:8px; font-family:monospace; border:1px solid #ccc; border-radius:4px; box-sizing:border-box; font-size:1rem;';

// --- PANNEAUX ---
var paneMath  = document.createElement('div');
paneMath.style.cssText  = 'display:flex; flex-wrap:wrap; gap:5px; align-items:center; min-height:100px; width:100%;';
var panePhys  = document.createElement('div');
panePhys.style.cssText  = 'display:none; flex-wrap:wrap; gap:5px; align-items:center; min-height:100px; width:100%;';
var paneGreek = document.createElement('div');
paneGreek.style.cssText = 'display:none; flex-wrap:wrap; gap:3px; align-items:center; min-height:100px; width:100%;';

// --- BOUTONS SYMBOLES ---
function makeSymBtn(label, code, pane) {
  var b = document.createElement('button');
  b.style.cssText = 'padding:4px 8px; cursor:pointer; background:#fff; border:1px solid #ccc; border-radius:3px; font-size:0.85rem; width:50px; height:40px; flex-shrink:0; display:flex; align-items:center; justify-content:center; overflow:hidden; user-select:none;';
  b.onmouseover = function() { this.style.borderColor='#3498db'; this.style.background='#eaf2f8'; };
  b.onmouseout  = function() { this.style.borderColor='#ccc';    this.style.background='#fff'; };
  b.onclick = function() { insertAtCursor(code); txtLatex.focus(); updatePreview(); };
  var displayCode = code;
  if (code === '_{}')             displayCode = 'a_n';
  else if (code === '^{n}')       displayCode = 'x^n';
  else if (code === '\\widehat{...}') displayCode = '\\widehat{A}';
  else if (code === '\ \\sqrt[a]{b} ') displayCode = '\\sqrt[a]{b}';
  var img = document.createElement('img');
  img.src = 'https://latex.codecogs.com/svg.image?' + encodeURIComponent(displayCode);
  img.style.cssText = 'max-width:100%; max-height:20px; display:block;';
  b.appendChild(img);
  pane.appendChild(b);
}

// Panneau Math
makeSymBtn('a/b', '\\frac{a}{b}', paneMath);
makeSymBtn('a.b', '\\cdot', paneMath);
makeSymBtn('a×b', '\\times', paneMath);
makeSymBtn('a^n', '^{n}', paneMath);
makeSymBtn('a_n', '_{}', paneMath);
makeSymBtn('√a', '\\sqrt{a}', paneMath);
makeSymBtn('b√a', '\ \\sqrt[a]{b} ', paneMath);
var s1=document.createElement('div'); s1.style.cssText='width:100%;border-top:1px solid #eee;margin:5px 0;font-size:.8em;color:#aaa;'; s1.innerText='Opérations'; paneMath.appendChild(s1);
makeSymBtn('≠', '\\neq', paneMath);
makeSymBtn('≤', '\\le', paneMath);
makeSymBtn('≥', '\\ge', paneMath);
makeSymBtn('∈', '\\in', paneMath);
makeSymBtn('∉', '\\notin', paneMath);
var s2=document.createElement('div'); s2.style.cssText='width:100%;border-top:1px solid #eee;margin:5px 0;font-size:.8em;color:#aaa;'; s2.innerText='Ensembles'; paneMath.appendChild(s2);
makeSymBtn('N','\\mathbb{N}',paneMath); makeSymBtn('Z','\\mathbb{Z}',paneMath); makeSymBtn('Q','\\mathbb{Q}',paneMath); makeSymBtn('R','\\mathbb{R}',paneMath); makeSymBtn('C','\\mathbb{C}',paneMath);
var s3=document.createElement('div'); s3.style.cssText='width:100%;border-top:1px solid #eee;margin:5px 0;font-size:.8em;color:#aaa;'; s3.innerText='Fonctions & Vecteurs'; paneMath.appendChild(s3);
makeSymBtn('|a|','|a|',paneMath); makeSymBtn('n!','n!',paneMath); makeSymBtn('∑','\\sum a',paneMath); makeSymBtn('∑lim','\\sum_{n=1}^{10} n^2',paneMath); makeSymBtn('∫','\\int_{a}^{b} x dx',paneMath); makeSymBtn('a⃗','\\vec{a}',paneMath); makeSymBtn('ȧ','\\dot{a}',paneMath); makeSymBtn('ä','\\ddot{a}',paneMath); makeSymBtn('â','\\widehat{...}',paneMath); makeSymBtn('∞','\\infty',paneMath);

// Panneau Physique
makeSymBtn('ā','\\overline{ab}',panePhys); makeSymBtn('a⃗','\\overrightarrow{ab}',panePhys); makeSymBtn('→','\\rightarrow',panePhys); makeSymBtn('←','\\leftarrow',panePhys); makeSymBtn('⇄','\\rightleftharpoons',panePhys);
var s4=document.createElement('div'); s4.style.cssText='width:100%;border-top:1px solid #eee;margin:5px 0;font-size:.8em;color:#aaa;'; s4.innerText='Chimie & Logique'; panePhys.appendChild(s4);
makeSymBtn('H₂SO₄','H_2SO_4',panePhys); makeSymBtn('²³⁸U','{}_{92}^{238}U',panePhys); makeSymBtn('SO₄²⁻','SO_4^{2-}',panePhys); makeSymBtn('∧','\\land',panePhys); makeSymBtn('a⃗','\\vec{a}',panePhys); makeSymBtn('ȧ','\\dot{a}',panePhys); makeSymBtn('ä','\\ddot{a}',panePhys);

// Panneau Grec
var gW1=document.createElement('div'); gW1.style.cssText='display:flex;flex-wrap:wrap;gap:2px;border-bottom:1px dashed #ccc;padding-bottom:5px;margin-bottom:5px;width:100%;';
var gW2=document.createElement('div'); gW2.style.cssText='display:flex;flex-wrap:wrap;gap:2px;width:100%;';
['α','β','γ','δ','ε','θ','λ','μ','π','ρ','σ','τ','φ','χ','ψ','ω'].forEach(function(c,i){
  var codes=['\\alpha','\\beta','\\gamma','\\delta','\\epsilon','\\theta','\\lambda','\\mu','\\pi','\\rho','\\sigma','\\tau','\\phi','\\chi','\\psi','\\omega'];
  makeSymBtn(c,codes[i],gW1);
});
['Γ','Δ','Θ','Λ','Π','Σ','Φ','Ψ','Ω'].forEach(function(c,i){
  var codes=['\\Gamma','\\Delta','\\Theta','\\Lambda','\\Pi','\\Sigma','\\Phi','\\Psi','\\Omega'];
  makeSymBtn(c,codes[i],gW2);
});
paneGreek.appendChild(gW1); paneGreek.appendChild(document.createElement('div')); paneGreek.appendChild(gW2);

// --- MODALE ---
var modal = document.createElement('div');
modal.style.cssText = 'display:none; position:fixed; top:0; left:0; width:100%; height:100%; background:rgba(0,0,0,0.5); z-index:9999; align-items:center; justify-content:center; padding:20px; box-sizing:border-box;';
var modalContent = document.createElement('div');
modalContent.style.cssText = 'background:white; padding:0; border-radius:8px; width:95%; max-width:600px; box-shadow:0 4px 15px rgba(0,0,0,0.2); max-height:90vh; display:flex; flex-direction:column; box-sizing:border-box;';
var mHead = document.createElement('h3');
mHead.innerText = 'Éditeur LaTeX';
mHead.style.cssText = 'margin:0; padding:15px; border-bottom:1px solid #eee; background:#f9f9f9; font-size:1.1rem; color:#333; text-align:center;';
var mBody = document.createElement('div');
mBody.style.cssText = 'flex:1; overflow-y:auto; padding:10px; overflow-x:hidden; min-height:200px;';
var mFoot = document.createElement('div');
mFoot.style.cssText = 'padding:15px; background:#f9f9f9; border-top:1px solid #ccc; border-radius:0 0 8px 8px; display:flex; flex-direction:column; gap:10px;';

// Onglets
var tabBar = document.createElement('div');
tabBar.style.cssText = 'display:flex; border-bottom:1px solid #ccc; margin-bottom:10px;';
var tabPanes = { 'Math': paneMath, 'Physique': panePhys, 'Lettres (Grec)': paneGreek };
['Math','Physique','Lettres (Grec)'].forEach(function(name, i) {
  var tb = document.createElement('button');
  tb.innerText = name;
  tb.style.cssText = 'flex:1; padding:8px; cursor:pointer; font-weight:bold; outline:none; border:none; border-bottom:2px solid ' + (i===0?'#34495e':'transparent') + '; background:' + (i===0?'#34495e':'none') + '; color:' + (i===0?'white':'#666') + ';';
  tb.onclick = function() {
    tabBar.querySelectorAll('button').forEach(function(b){ b.style.background='none'; b.style.color='#666'; b.style.borderBottomColor='transparent'; });
    tb.style.background='#34495e'; tb.style.color='white'; tb.style.borderBottomColor='#34495e';
    Object.keys(tabPanes).forEach(function(k){ tabPanes[k].style.display='none'; });
    tabPanes[name].style.display='flex';
  };
  tabBar.appendChild(tb);
});
mBody.appendChild(tabBar);
mBody.appendChild(paneMath); mBody.appendChild(panePhys); mBody.appendChild(paneGreek);

// Prévisualisation
var prevDiv = document.createElement('div');
prevDiv.style.cssText = 'border:1px dashed #ccc; padding:5px; min-height:20px; text-align:center; background:#fafafa; overflow-x:auto;';
var prevImg = document.createElement('img');
prevImg.style.cssText = 'max-width:100%; max-height:30px;';
prevDiv.appendChild(prevImg);
prevDiv.appendChild(document.createTextNode(' Prévisualisation'));
mFoot.appendChild(prevDiv);
mFoot.appendChild(txtLatex);

// Boutons footer
var btnRow = document.createElement('div');
btnRow.style.cssText = 'display:flex; gap:10px;';
var btnCancel = document.createElement('button');
btnCancel.innerText = 'Annuler';
btnCancel.style.cssText = 'flex:1; padding:10px; cursor:pointer; background:#95a5a6; color:white; border:none; border-radius:4px; font-weight:bold;';
btnCancel.onclick = function() { modal.style.display='none'; };
var btnIns = document.createElement('button');
btnIns.innerText = 'Insérer la formule';
btnIns.style.cssText = 'flex:2; padding:10px; cursor:pointer; background:#27ae60; color:white; border:none; border-radius:4px; font-weight:bold;';
btnIns.onclick = function() {
  var txt = txtLatex.value;
  if (txt.trim() !== '') {
    var imgEl = document.createElement('img');
    imgEl.src = 'https://latex.codecogs.com/svg.image?' + encodeURIComponent('\\displaystyle ' + txt);
    imgEl.alt = txt;
    imgEl.style.cssText = 'vertical-align:middle; margin:0 3px; max-height:1.8em; max-width:100%;';
    editor.focus();
    var sel = window.getSelection();
    if (sel.rangeCount > 0) {
      var range = sel.getRangeAt(0);
      range.deleteContents();
      range.insertNode(imgEl);
      range.setStartAfter(imgEl);
      range.collapse(true);
      sel.removeAllRanges(); sel.addRange(range);
    }
    updateData();
  }
  modal.style.display = 'none';
};
btnRow.appendChild(btnCancel); btnRow.appendChild(btnIns);
mFoot.appendChild(btnRow);

modalContent.appendChild(mHead); modalContent.appendChild(mBody); modalContent.appendChild(mFoot);
modal.appendChild(modalContent);
box.appendChild(modal);

// --- BARRE D'OUTILS ---
var toolbar = document.createElement('div');
toolbar.style.cssText = 'display:flex; gap:8px; margin-bottom:10px; flex-wrap:wrap; align-items:center;';
function makeToolBtn(html, cmd, color) {
  var b = document.createElement('button');
  b.type = 'button'; b.innerHTML = html;
  b.style.cssText = 'padding:6px 12px; cursor:pointer; background:' + (color||'#34495e') + '; color:white; border:none; border-radius:4px; font-weight:bold; font-size:.9rem;';
  b.onclick = function() { editor.focus(); document.execCommand(cmd, false, null); updateData(); };
  return b;
}
toolbar.appendChild(makeToolBtn('<b>Gras</b>', 'bold'));
toolbar.appendChild(makeToolBtn('<i>Italique</i>', 'italic'));
toolbar.appendChild(makeToolBtn('<u>Souligné</u>', 'underline'));
var btnMath = document.createElement('button');
btnMath.type = 'button'; btnMath.innerHTML = '∑ Formule (LaTeX)';
btnMath.style.cssText = 'padding:6px 12px; cursor:pointer; background:#8e44ad; color:white; border:none; border-radius:4px; font-weight:bold; font-size:.9rem; margin-left:8px;';
btnMath.onclick = function() { modal.style.display='flex'; txtLatex.value=''; prevImg.src=''; txtLatex.focus(); updatePreview(); };
toolbar.appendChild(btnMath);

// --- ÉDITEUR ---
var editor = document.createElement('div');
editor.contentEditable = true;
editor.innerHTML = 'Tapez votre réponse ici...';
editor.style.cssText = 'width:100%; min-height:300px; border:2px solid #34495e; padding:10px; background:white; border-radius:4px; outline:none; line-height:1.6; font-size:1.1rem; box-sizing:border-box; overflow-wrap:break-word;';
editor.addEventListener('focus', function() { if (editor.innerHTML === 'Tapez votre réponse ici...') editor.innerHTML = ''; }, {once:true});

box.appendChild(toolbar);
box.appendChild(editor);

// --- LOGIQUE ---
function insertAtCursor(text) {
  var s=txtLatex.selectionStart, e=txtLatex.selectionEnd, v=txtLatex.value;
  txtLatex.value = v.substring(0,s) + text + v.substring(e);
  var np = s + text.length;
  if (text==='\\frac{a}{b}') np=s+6;
  else if (text==='\\sqrt{a}') np=s+6;
  else if (text==='^{n}') np=s+2;
  else if (text==='_{}') np=s+2;
  txtLatex.selectionStart = txtLatex.selectionEnd = np;
}

function updatePreview() {
  var txt = txtLatex.value;
  if (txt.trim()==='') { prevImg.style.display='none'; return; }
  prevImg.style.display='inline-block';
  prevImg.src = 'https://latex.codecogs.com/svg.image?' + encodeURIComponent('\\displaystyle '+txt);
}
txtLatex.addEventListener('keyup', updatePreview);

function updateData() {
  var html = editor.innerHTML;
  var plain = editor.innerText || '';
  setRef(refAns1, html);
  setRef(refAns2, (plain.trim()!=='' && plain.trim()!=='Tapez votre réponse ici...') ? '0' : '');
}
editor.addEventListener('input', updateData);
editor.addEventListener('keyup', updateData);
editor.addEventListener('mouseup', updateData);

// Restauration si déjà saisi
var savedVal = document.getElementById(refAns1);
if (savedVal && savedVal.value && savedVal.value.trim()!=='') {
  editor.innerHTML = savedVal.value;
  updateData();
}

[[/jsxgraph]]
<p style="margin-top:10px;padding:10px 14px;background:#fef3c7;border-left:4px solid #f59e0b;border-radius:4px;font-size:.9rem;color:#78350f;">👨‍🏫 Votre réponse sera lue et corrigée par votre professeur.</p>
<div style="display:none;">
  [[input:ans1_html]] [[validation:ans1_html]]
  [[input:ans1_flag]] [[validation:ans1_flag]]
</div>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>4</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>/* Q1 : Composition libre (correction manuelle) */
ta1:1;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<ol>
  <li>[[feedback:prt1]]</li>
</ol>]]></text>
    </specificfeedback>
    <questionnote format="html">
      <text>Composition Q1</text>
    </questionnote>
    <questiondescription format="html">
      <text>Généré avec l'aide de Héstack V1.0</text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1_flag</name>
      <type>algebraic</type>
      <tans>1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans1_html</name>
      <type>notes</type>
      <tans>1</tans>
      <boxsize>2000</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options>manualgraded:true</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>4.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Toujours vrai — correction manuelle par le professeur</description>
        <answertest>AlgEquiv</answertest>
        <sans>1</sans>
        <tans>1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-1-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p style="padding:10px;background:#f0fdf4;border-left:4px solid #22c55e;border-radius:4px;color:#166534;">✅ Réponse enregistrée. Votre professeur la corrigera et attribuera la note.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>PRT-1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
    <tags>
      <tag><text>stack</text>
</tag>
      <tag><text>qtype:redaction</text>
</tag>
      <tag><text>allemand</text>
</tag>
      <tag><text>2nde</text>
</tag>
      <tag><text>chapitre:axes</text>
</tag>
      <tag><text>chapitre:representationdesoi</text>
</tag>
      <tag><text>bloom:memoriser</text>
</tag>
      <tag><text>difficulte:simple</text>
</tag>
    </tags>
  </question>

</quiz>