<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 396667  -->
  <question type="stack">
    <name>
      <text>string Réponse textuelle (Liste de synonymes)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<div style="background:#dc2626;border-left:5px solid #991b1b;border-radius:0 8px 8px 0;padding:10px 16px;margin-bottom:12px;display:flex;align-items:center;gap:10px;flex-wrap:wrap;">
        <strong style="font-weight:800;color:#fff;font-size:.95rem;">Q1 — Réponse textuelle</strong>
        <span style="background:#991b1b;color:#fff;padding:2px 9px;border-radius:20px;font-size:.78rem;font-weight:700;">/ 1 pt</span>
        <span style="background:#ffffff;color:#991b1b;border:1px solid #991b1b;padding:2px 9px;border-radius:20px;font-size:.75rem;font-weight:600;">🔤 Réponse texte</span>
      </div>
      <!-- ENONCE-START -->Quelle est l'unité de force ?<!-- ENONCE-END -->
      <p>[[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<div style="margin-top:20px; padding:15px; background:#f8fafc; border:1px solid #cbd5e1; border-radius:8px; box-shadow:0 2px 4px rgba(0,0,0,0.05);">
        <div style="font-weight:bold; color:#0f766e; margin-bottom:10px; display:flex; align-items:center; gap:8px;">
          <span style="font-size:1.2rem;">🔑</span> Réponses attendues
        </div>
        <div style="margin-bottom:8px;font-size:.9rem;border-bottom:1px dashed #e2e8f0;padding-bottom:6px;">
          <span style="font-weight:bold;color:#0f766e;">Q1 Texte :</span>
          <code style="background:#f1f5f9;padding:2px 6px;border-radius:4px;color:#d946ef;">Le Newton</code> (ou Newton, newton, N...)
        </div>
      </div>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2024070500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* Q1 : String (1pt) */

/* Définition des listes de synonymes acceptés */
liste_parfaite : ["Le Newton", "le newton", "Le newton", "le Newton"];
liste_sans_article : ["Newton", "newton"];
liste_orthographe : ["Neuton", "Le Neuton", "le neuton", "Newtons", "le newtons", "Le newtons", "newtons"];
liste_symbole : ["N"];

/* La réponse modèle par défaut (pour l'interface) */
ta1: "Le Newton";]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="html">
      <text>Réponse : {@ans1@}</text>
    </questionnote>
    <questiondescription format="html">
      <text>Généré avec HéStack V2 | hestack-types:string</text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>E</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>string</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>2</feedbackstyle>
      <feedbackvariables>
        <text>/* On vérifie si la réponse de l'élève est dans les différentes listes */
est_parfait : if member(ans1, liste_parfaite) then true else false;
est_sans_article : if member(ans1, liste_sans_article) then true else false;
a_faute : if member(ans1, liste_orthographe) then true else false;
est_symbole : if member(ans1, liste_symbole) then true else false;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Vérifie si la réponse est parfaite (avec article)</description>
        <answertest>AlgEquiv</answertest>
        <sans>est_parfait</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-0-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left:4px solid #15803d;padding:10px 14px;background:#f0fdf4;border-radius:4px;margin:4px 0;"><p>✅ <strong>Excellent !</strong> La réponse est parfaitement formulée avec l'article.</p></div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-0-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Vérifie si c'est juste sans l'article</description>
        <answertest>AlgEquiv</answertest>
        <sans>est_sans_article</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left:4px solid #15803d;padding:10px 14px;background:#f0fdf4;border-radius:4px;margin:4px 0;"><p>✅ <strong>Bonne réponse.</strong> On préfère généralement écrire "Le Newton" avec l'article pour nommer une unité, mais votre réponse est acceptée.</p></div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Piège : Faute de frappe (ex: Neuton)</description>
        <answertest>AlgEquiv</answertest>
        <sans>a_faute</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left:4px solid #ca8a04;padding:10px 14px;background:#fefce8;border-radius:4px;margin:4px 0;"><p>🔶 <strong>Presque ça !</strong> Attention à l'orthographe, il s'écrit avec un "w" : <strong>N<u>e</u>wton</strong>. (Moitié des points accordée).</p></div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Piège : A donné le symbole (N) au lieu du nom</description>
        <answertest>AlgEquiv</answertest>
        <sans>est_symbole</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left:4px solid #ca8a04;padding:10px 14px;background:#fefce8;border-radius:4px;margin:4px 0;"><p>🔶 <strong>Attention au sens de la question.</strong> "N" est le symbole de l'unité. La question demandait le <strong>nom</strong> de l'unité (Newton). (Moitié des points accordée).</p></div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>4</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description>Faux</description>
        <answertest>AlgEquiv</answertest>
        <sans>true</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left:4px solid #dc2626;padding:10px 14px;background:#fef2f2;border-radius:4px;margin:4px 0;"><p>❌ <strong>Réponse incorrecte.</strong> L'unité légale de la force est le <strong>Newton</strong> (symbole N).</p></div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
  </question>

</quiz>