<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 396293  -->
  <question type="stack">
    <name>
      <text>Suites : Limite d'une suite géométrique (|q| inf 1)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<div style="background: #475569; border-left: 5px solid #334155; border-radius: 0 8px 8px 0; padding: 10px 16px; margin-bottom: 12px; display: flex; align-items: center; gap: 10px; flex-wrap: wrap;"><strong style="font-weight: 800; color: #fff; font-size: .95rem;">Suites numériques — Limite géométrique</strong> <span style="background: #334155; color: #fff; padding: 2px 9px; border-radius: 20px; font-size: .78rem; font-weight: bold;">/ 1 pt</span></div>
<p>Soit \((U_n)\) la suite géométrique définie par :</p>
<div style="display: flex; gap: 30px; justify-content: center; margin: 15px 0; flex-wrap: wrap;">
<div>\(U_0 = {@U0@}\)</div>
<div>\(q = {@q@}\)</div>
</div>
<p><strong>1.</strong> Déterminer la limite de la suite \((U_n)\) lorsque \(n\) tend vers \(+\infty\).</p>
<div style="margin: 10px 0;">\(\lim\limits_{n \to +\infty} U_n = \) [[input:ans_lim]] [[validation:ans_lim]]</div>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<div style="margin-top: 20px; padding: 15px; background: #f8fafc; border: 1px solid #cbd5e1; border-radius: 8px;">
<div style="font-weight: bold; margin-bottom: 10px;">🔑 Correction détaillée</div>
<div style="font-size: .9rem;">
Le terme général de la suite est \(U_n = U_0 \times q^n\).<br /><br />
Ici, la raison est \(q = {@q@}\). On a :<br />
\[
|q| = \left|{@q@}\right| = {@abs_q@} < 1
\]<br />
<strong>Théorème :</strong> Si \(|q| < 1\), alors \(\lim\limits_{n \to +\infty} q^n = 0\).<br /><br />
<strong>Conclusion :</strong><br />
Puisque \(U_n = {@U0@} \times q^n\) et que \(q^n\) tend vers 0, le produit tend vers 0 :<br />
\[
\lim\limits_{n \to +\infty} U_n = {@U0@} \times 0 = 0
\]
</div>
</div>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2024092500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* Redéfinition locale du tirage aléatoire inclusif (Pôle 2.9) */
ri(a,b) := a + rand(b - a + 1)$ 
/* Génération des paramètres */
U0_pool: [-5, -4, -3, -2, -1, 1, 2, 3, 4, 5]$ q_pool: [1/2, -1/2, 1/3, -1/3, 2/3, -2/3, 3/4, -3/4]$ U0: U0_pool[ri(1, length(U0_pool))]$ q: q_pool[ri(1, length(q_pool))]$ 
/* Valeur absolue de q pour la correction */
abs_q: abs(q)$ 
/* Réponses attendues et erreurs classiques */
ans_correct: 0$ err_u0: U0$ err_u1: U0 * q$ err_inf: inf$ err_minf: -inf$]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<div>[[feedback:prt1]]</div> ]]></text>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>U0={@U0@}, q={@q@}, Limite=0</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p>Généré avec HéStack V2 | hestack-types:suites-geometrique-limite-q-inf-1</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans_lim</name>
      <type>algebraic</type>
      <tans>ans_correct</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>0</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords>inf</allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Réponse correcte</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_lim</sans>
        <tans>ans_correct</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-OK</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #15803d; padding: 10px 14px; background: #f0fdf4; border-radius: 4px;">✅ <strong>Parfait !</strong> La valeur absolue de la raison est strictement inférieure à 1, la suite converge bien vers 0.</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>PRT-NOK</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Erreur : A répondu U_0</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_lim</sans>
        <tans>err_u0</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-U0</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Premier terme !</strong> Vous avez donné la valeur de \(U_0\). On ne cherche pas la valeur initiale, mais le comportement de la suite lorsque \(n\) devient infiniment grand.</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>PRT-ERR-NOT-U0</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Erreur : A répondu U_1</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_lim</sans>
        <tans>err_u1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-U1</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Terme suivant !</strong> Vous avez calculé \(U_0 \times q = U_1\). La limite correspond à \(n \to +\infty\), il ne faut pas remplacer \(n\) par 1, mais appliquer le théorème sur la limite de \(q^n\).</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>PRT-ERR-NOT-U1</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Erreur : A répondu +inf</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_lim</sans>
        <tans>err_inf</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-INF</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>La suite ne diverge pas !</strong> Vous avez répondu \(+\infty\). Regardez la raison : \(|q| = {@abs_q@} < 1\). Les termes de la suite se rapprochent de 0, ils ne partent pas vers l'infini.</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>4</falsenextnode>
        <falseanswernote>PRT-ERR-NOT-INF</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description>Erreur : A répondu -inf</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_lim</sans>
        <tans>err_minf</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-MINF</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>La suite ne diverge pas !</strong> Vous avez répondu \(-\infty\). Regardez la raison : \(|q| = {@abs_q@} < 1\). Les termes de la suite se rapprochent de 0, ils ne partent pas vers l'infini.</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>PRT-ERR-CALC</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #dc2626; padding: 10px 14px; background: #fff0f0; border-radius: 4px;">❌ <strong>Incorrect.</strong> Rappel de cours :<br />- Si \(|q| < 1\) alors \(\lim\limits_{n \to +\infty} q^n = 0\)<br />- Si \(|q| > 1\) alors la suite n'a pas de limite finie.<br />Ici, que vaut \(|q|\) ?</div>]]></text>
        </falsefeedback>
      </node>
    </prt>
  </question>

</quiz>