<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 396276  -->
  <question type="stack">
    <name>
      <text>Suites : Terme général d'une suite géométrique</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<div style="background: #475569; border-left: 5px solid #334155; border-radius: 0 8px 8px 0; padding: 10px 16px; margin-bottom: 12px; display: flex; align-items: center; gap: 10px; flex-wrap: wrap;"><strong style="font-weight: 800; color: #fff; font-size: .95rem;">Suites numériques — Terme général</strong> <span style="background: #334155; color: #fff; padding: 2px 9px; border-radius: 20px; font-size: .78rem; font-weight: bold;">/ 1 pt</span></div>
<p>Soit \((U_n)\) une suite géométrique de raison \(q\) telle que :</p>
<div style="display: flex; gap: 30px; justify-content: center; margin: 15px 0; flex-wrap: wrap;">
<div>\(U_0 = {@U0@}\)</div>
<div>\(U_{@k@} = {@Uk@}\)</div>
</div>
<p><strong>1.</strong> Déterminer l'expression du terme général \(U_n\) en fonction de \(n\).</p>
<div style="margin: 10px 0;">\(U_n = \) [[input:ans_un]] [[validation:ans_un]]</div>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<div style="margin-top: 20px; padding: 15px; background: #f8fafc; border: 1px solid #cbd5e1; border-radius: 8px;">
<div style="font-weight: bold; margin-bottom: 10px;">🔑 Correction détaillée</div>
<div style="font-size: .9rem;">
Pour une suite géométrique, le terme général est donné par la formule :<br /><br />
\[
U_n = U_0 \times q^n
\]<br />
<strong>1. Calcul de la raison \(q\) :</strong><br />
On sait que \(U_{@k@} = U_0 \times q^{@k@}\), donc :<br />
\[
{@Uk@} = {@U0@} \times q^{@k@} \iff q^{@k@} = \frac{@Uk@}}{@U0@}} = {@Uk_sur_U0@}
\]
\[
q = \left({@Uk_sur_U0@}\right)^{\frac{1}{@k@}} = {@q@}
\]<br />
<strong>2. Expression de \(U_n\) :</strong><br />
En remplaçant dans la formule générale, on obtient :<br />
\[
U_n = {@U0@} \times \left({@q@}\right)^n
\]
</div>
</div>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2024092500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* Redéfinition locale du tirage aléatoire inclusif (Pôle 2.9) */
ri(a,b) := a + rand(b - a + 1)$ 
/* Génération des paramètres */
U0_pool: [1, 2, 3, 4, 5, -1, -2, -3]$ q_pool: [2, -2, 3, -3, 1/2, -1/2]$ U0: U0_pool[ri(1, length(U0_pool))]$ k: ri(2, 4)$ q: q_pool[ri(1, length(q_pool))]$ 
/* Calcul des valeurs */
Uk: U0 * q^k$ Uk_sur_U0: Uk / U0$ 
/* Réponses attendues et erreurs classiques (Cascade Diagnostique Pôle 4.9) */
ans_correct: U0 * q^n$ 
/* Erreur 1 : A utilisé la formule d'une suite arithmétique (U0 + n*r) */
r_arith: (Uk - U0) / k$ err_arith: U0 + n*r_arith$ 
/* Erreur 2 : A oublié la racine k-ième (q calculé comme Uk/U0) */
err_q_faux: U0 * (Uk/U0)^n$ 
/* Erreur 3 : A fait une addition au lieu d'une multiplication (U0 + q^n) */
err_add: U0 + q^n$ 
/* Erreur 4 : A oublié de mettre l'exposant n à q */
err_nok: U0 * q$]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<div>[[feedback:prt1]]</div> ]]></text>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>U0={@U0@}, k={@k@}, Uk={@Uk@}, q={@q@}, Un={@U0@}*({@q@})^n</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p>Généré avec HéStack V2 | hestack-types:suites-geometrique-un</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans_un</name>
      <type>algebraic</type>
      <tans>ans_correct</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>U0 * q^n</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>2</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Réponse correcte</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_un</sans>
        <tans>ans_correct</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-OK</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #15803d; padding: 10px 14px; background: #f0fdf4; border-radius: 4px;">✅ <strong>Parfait !</strong> Le terme général de la suite est correct.</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>PRT-NOK</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Erreur : A utilisé la formule arithmétique</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_un</sans>
        <tans>err_arith</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-ARITH</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Confusion de suite !</strong> Vous avez utilisé la formule d'une suite <strong>arithmétique</strong> (\(U_n = U_0 + n \times r\)) au lieu d'une suite <strong>géométrique</strong> (\(U_n = U_0 \times q^n\)). Ici, on multiplie, on n'additionne pas.</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>PRT-ERR-NOT-ARITH</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Erreur : A oublié la racine k-ième pour trouver q</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_un</sans>
        <tans>err_q_faux</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-Q-ROOT</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Erreur sur la raison \(q\) !</strong> Vous avez calculé \(q = \frac{U_{@k@}}{U_0}\) au lieu de \(q = \sqrt[{@k@}]{\frac{U_{@k@}}{U_0}}\). N'oubliez pas la racine @{k@}-ième car le terme donné est \(U_{@k@}\) (et non \(U_1\)).</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>PRT-ERR-NOT-Q-ROOT</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Erreur : Addition au lieu de la multiplication</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_un</sans>
        <tans>err_add</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-ADD</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Addition au lieu de multiplication !</strong> La formule est \(U_n = U_0 \times q^n\). Vous avez écrit une addition entre \(U_0\) et \(q^n\).</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>4</falsenextnode>
        <falseanswernote>PRT-ERR-NOT-ADD</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description>Erreur : A oublié l'exposant n</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_un</sans>
        <tans>err_nok</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-NOK</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Oubli de l'exposant !</strong> Vous avez écrit \(U_n = U_0 \times q\). La raison doit être élevée à la puissance \(n\) : \(U_n = U_0 \times q^n\).</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>PRT-ERR-CALC</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #dc2626; padding: 10px 14px; background: #fff0f0; border-radius: 4px;">❌ <strong>Incorrect.</strong> Appliquez la méthode :<br />1. Calculez la raison : \(q = \sqrt[{@k@}]{\frac{U_{@k@}}{U_0}}\)<br />2. Appliquez la formule : \(U_n = U_0 \times q^n\)<br /><em>(Attention à bien saisir le signe " * " pour la multiplication si nécessaire).</em></div>]]></text>
        </falsefeedback>
      </node>
    </prt>
  </question>

</quiz>