<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 396354  -->
  <question type="stack">
    <name>
      <text>Fonctions : Plus petite racine d'un polynôme</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<div style="background: #475569; border-left: 5px solid #334155; border-radius: 0 8px 8px 0; padding: 10px 16px; margin-bottom: 12px; display: flex; align-items: center; gap: 10px; flex-wrap: wrap;"><strong style="font-weight: 800; color: #fff; font-size: .95rem;">Fonctions — Second degré</strong> <span style="background: #334155; color: #fff; padding: 2px 9px; border-radius: 20px; font-size: .78rem; font-weight: bold;">/ 1 pt</span></div>
<p>Soit la fonction polynôme du second degré définie sur \(\mathbb{R}\) par :</p>
<div style="text-align: center; font-size: 1.2rem; margin: 25px 0; background: #f8fafc; padding: 15px; border-radius: 8px; border: 1px solid #e2e8f0;">
\(P(x) = {@poly@}\)
</div>
<p>On admet que le discriminant est strictement positif (\(\Delta > 0\)), le polynôme admet donc deux racines réelles distinctes.</p>
<p><strong>1.</strong> Déterminer la <strong>plus petite</strong> des deux racines du polynôme \(P\).</p>
<div style="margin: 15px 0 10px 0;">\(x_{\min} = \) [[input:ans_rmin]] [[validation:ans_rmin]]</div>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<div style="margin-top: 20px; padding: 15px; background: #f8fafc; border: 1px solid #cbd5e1; border-radius: 8px;">
<div style="font-weight: bold; margin-bottom: 10px;">🔑 Correction détaillée</div>
<div style="font-size: .9rem;">
On résout l'équation \(P(x) = 0\).<br><br>
<strong>1. Calcul du discriminant :</strong><br>
\[
\Delta = b^2 - 4ac = ({@b@})^2 - 4 \times ({@a@}) \times ({@c@}) = {@delta@}
\]<br>
<strong>2. Calcul des deux racines :</strong><br>
\[
x_1 = \frac{-b - \sqrt{\Delta}}{2a} = {@r1@}
\]
\[
x_2 = \frac{-b + \sqrt{\Delta}}{2a} = {@r2@}
\]<br>
<strong>3. Comparaison :</strong><br>
Puisque l'on cherche la plus petite racine, on compare {@r1@} et {@r2@} :<br>
\[
\min({@r1@}, {@r2@}) = {@r_min@}
\]
</div>
</div>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* Fonction utilitaire */
ri(a,b) := a + rand(b - a + 1) - 1;

/* Génération secrète de deux racines entières distinctes */
r1: ri(-4, 4);
r2: ri(-4, 4);
while r2 = r1 do r2: ri(-4, 4);

/* Détermination explicite de la plus petite et de la plus grande */
r_min: min(r1, r2);
r_max: max(r1, r2);

/* Génération du coefficient a (non nul) */
r_a: rand(6);
if r_a = 1 then a: -3
else if r_a = 2 then a: -2
else if r_a = 3 then a: -1
else if r_a = 4 then a: 1
else if r_a = 5 then a: 2
else a: 3;

/* Fabrication du polynôme à l'envers : a(x - r1)(x - r2) */
b: -a * (r1 + r2);
c: a * r1 * r2;
poly: a*x^2 + b*x + c;

/* Calcul du discriminant */
delta: b^2 - 4*a*c;

/* Réponse attendue */
ans_correct: r_min;

/* Erreurs classiques */
/* Erreur 1 : A donné la plus grande racine au lieu de la plus petite */
err_max: r_max;
/* Erreur 2 : A oublié le "a" au dénominateur */
err_no_a: min((-b - sqrt(delta))/2, (-b + sqrt(delta))/2);
/* Erreur 3 : A oublié le signe "-" devant b */
err_no_minus: min((b - sqrt(delta))/(2*a), (b + sqrt(delta))/(2*a));
/* Erreur 4 : A donné l'abscisse du sommet */
err_sommet: -b/(2*a);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<div>[[feedback:prt1]]</div>]]></text>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>P(x)={@poly@}, Plus petite racine={@r_min@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p>Généré avec HéStack V2 | hestack-types:fonctions-second-degre-racine-min</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans_rmin</name>
      <type>algebraic</type>
      <tans>ans_correct</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>...</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Réponse correcte</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_rmin</sans>
        <tans>ans_correct</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-OK</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #15803d; padding: 10px 14px; background: #f0fdf4; border-radius: 4px;">✅ <strong>Parfait !</strong> C'est bien la plus petite racine.</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>PRT-NOK</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Erreur : A donné la plus grande racine</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_rmin</sans>
        <tans>err_max</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-MAX</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Mauvaise racine !</strong> Votre calcul semble correct, mais vous avez donné la <strong>plus grande</strong> des deux racines. L'énoncé demande explicitement la <strong>plus petite</strong>.</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>PRT-ERR-NOT-MAX</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description><![CDATA[Erreur : Oubli du coefficient "a" au dénominateur]]></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_rmin</sans>
        <tans>err_no_a</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-NOA</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Oubli du coefficient \(a\) !</strong> Votre démarche est bonne, mais vous avez divisé par 2 au lieu de \(2a\). Le dénominateur de la formule des racines est \(2a\).</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>PRT-ERR-NOT-NOA</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description><![CDATA[Erreur : Oubli du signe "-" devant b]]></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_rmin</sans>
        <tans>err_no_minus</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-NOMINUS</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Erreur de signe sur \(b\) !</strong> La formule des racines commence par un signe <strong>moins</strong> : \(\frac{-b \pm \sqrt{\Delta}}{2a}\). Vous avez calculé avec \(+b\).</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>4</falsenextnode>
        <falseanswernote>PRT-ERR-NOT-NOMINUS</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description>Erreur : A donné l'abscisse du sommet</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_rmin</sans>
        <tans>err_sommet</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-SOMMET</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Confusion avec le sommet !</strong> Vous avez donné la valeur \(\frac{-b}{2a}\), qui est l'abscisse du sommet de la parabole (et la moyenne des deux racines). On vous demande une <strong>racine</strong> (où \(P(x) = 0\)).</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>PRT-ERR-CALC</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #dc2626; padding: 10px 14px; background: #fff0f0; border-radius: 4px;">❌ <strong>Incorrect.</strong> Appliquez la méthode :<br>1. Calculez \(\Delta = b^2 - 4ac\).<br>2. Calculez les deux racines : \(x_1 = \frac{-b - \sqrt{\Delta}}{2a}\) et \(x_2 = \frac{-b + \sqrt{\Delta}}{2a}\).<br>3. Conservez uniquement la <strong>plus petite</strong> des deux.</div>]]></text>
        </falsefeedback>
      </node>
    </prt>
  </question>

</quiz>