<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 396347  -->
  <question type="stack">
    <name>
      <text>Fonctions : Calculer le discriminant d'un polynôme</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<div style="background: #475569; border-left: 5px solid #334155; border-radius: 0 8px 8px 0; padding: 10px 16px; margin-bottom: 12px; display: flex; align-items: center; gap: 10px; flex-wrap: wrap;"><strong style="font-weight: 800; color: #fff; font-size: .95rem;">Fonctions — Second degré</strong> <span style="background: #334155; color: #fff; padding: 2px 9px; border-radius: 20px; font-size: .78rem; font-weight: bold;">/ 1 pt</span></div>
<p>Soit la fonction polynôme du second degré définie sur \(\mathbb{R}\) par :</p>
<div style="text-align: center; font-size: 1.2rem; margin: 25px 0; background: #f8fafc; padding: 15px; border-radius: 8px; border: 1px solid #e2e8f0;">
\(P(x) = {@poly@}\)
</div>
<p> Calculer le discriminant du polynôme \(P\), noté \(\Delta\).</p>
<div style="margin: 15px 0 10px 0;">\(\Delta = \) [[input:ans_delta]] [[validation:ans_delta]]</div>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<div style="margin-top: 20px; padding: 15px; background: #f8fafc; border: 1px solid #cbd5e1; border-radius: 8px;">
<div style="font-weight: bold; margin-bottom: 10px;">🔑 Correction détaillée</div>
<div style="font-size: .9rem;">
Pour un polynôme de la forme \(ax^2 + bx + c\), le discriminant est donné par la formule :<br><br>
\[
\Delta = b^2 - 4ac
\]<br>
<strong>Identification des coefficients :</strong><br>
Ici, on a \(a = {@a@}\), \(b = {@b@}\) et \(c = {@c@}\).<br><br>
<strong>Application numérique :</strong><br>
\[
\Delta = ({@b@})^2 - 4 \times ({@a@}) \times ({@c@})
\]
\[
\Delta = {@b_carre@} - {@produit_4ac@} = {@ans_correct@}
\]
</div>
</div>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>/* Fonction utilitaire */
ri(a,b) := a + rand(b - a + 1) - 1;

/* Génération des coefficients */
/* a ne doit pas être nul pour être du second degré */
r_a: rand(6);
if r_a = 1 then a: -3
else if r_a = 2 then a: -2
else if r_a = 3 then a: -1
else if r_a = 4 then a: 1
else if r_a = 5 then a: 2
else a: 3;

b: ri(-5, 5);
c: ri(-5, 5);

/* Expression du polynôme pour l'affichage */
poly: a*x^2 + b*x + c;

/* Valeurs intermédiaires pour la correction */
b_carre: b^2;
produit_4ac: 4*a*c;

/* Réponse attendue */
ans_correct: b^2 - 4*a*c;

/* Erreurs classiques */
/* Erreur 1 : A oublié le 4 */
err_no_4: b^2 - a*c;
/* Erreur 2 : A fait une addition au lieu d'une soustraction */
err_plus: b^2 + 4*a*c;
/* Erreur 3 : A inversé les termes */
err_inv: 4*a*c - b^2;
/* Erreur 4 : A oublié b^2 */
err_no_b2: -4*a*c;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<div>[[feedback:prt1]]</div>]]></text>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>P(x)={@poly@}, Delta={@ans_correct@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p>Généré avec HéStack V2 | hestack-types:fonctions-second-degre-discriminant</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans_delta</name>
      <type>algebraic</type>
      <tans>ans_correct</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>2</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Réponse correcte</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_delta</sans>
        <tans>ans_correct</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-OK</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #15803d; padding: 10px 14px; background: #f0fdf4; border-radius: 4px;">✅ <strong>Parfait !</strong> Le discriminant est correct.</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>PRT-NOK</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Erreur : A oublié le 4</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_delta</sans>
        <tans>err_no_4</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-NO4</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Oubli du coefficient 4 !</strong> Vous avez calculé \(b^2 - ac\). La formule complète du discriminant est \(\Delta = b^2 - \mathbf{4}ac\).</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>PRT-ERR-NOT-NO4</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Erreur : A fait une addition</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_delta</sans>
        <tans>err_plus</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-PLUS</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Erreur de signe !</strong> Vous avez calculé \(b^2 + 4ac\). C'est une <strong>soustraction</strong> qu'il faut faire : \(\Delta = b^2 - 4ac\).</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>PRT-ERR-NOT-PLUS</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Erreur : A inversé les termes</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_delta</sans>
        <tans>err_inv</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-INV</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Termes inversés !</strong> Vous avez calculé \(4ac - b^2\). On commence toujours par le carré de \(b\) : \(\Delta = b^2 - 4ac\).</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>4</falsenextnode>
        <falseanswernote>PRT-ERR-NOT-INV</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description>Erreur : A oublié b^2</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_delta</sans>
        <tans>err_no_b2</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-NOB2</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Il manque \(b^2\) !</strong> Vous n'avez calculé que \(-4ac\). Le discriminant comporte deux termes : \(\Delta = \mathbf{b^2} - 4ac\).</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>PRT-ERR-CALC</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #dc2626; padding: 10px 14px; background: #fff0f0; border-radius: 4px;">❌ <strong>Incorrect.</strong> Appliquez rigoureusement la formule :<br>\(\Delta = b^2 - 4ac\)<br><em>(Attention aux règles de signes lors de la multiplication et de la soustraction).</em></div>]]></text>
        </falsefeedback>
      </node>
    </prt>
  </question>

</quiz>