<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 396352  -->
  <question type="stack">
    <name>
      <text>Fonctions : Racines réelles d'un polynôme du second degré</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<div style="background: #475569; border-left: 5px solid #334155; border-radius: 0 8px 8px 0; padding: 10px 16px; margin-bottom: 12px; display: flex; align-items: center; gap: 10px; flex-wrap: wrap;"><strong style="font-weight: 800; color: #fff; font-size: .95rem;">Fonctions — Racines</strong> <span style="background: #334155; color: #fff; padding: 2px 9px; border-radius: 20px; font-size: .78rem; font-weight: bold;">/ 1 pt</span></div>
<p>Soit la fonction polynôme du second degré définie sur \(\mathbb{R}\) par :</p>
<div style="text-align: center; font-size: 1.2rem; margin: 25px 0; background: #f8fafc; padding: 15px; border-radius: 8px; border: 1px solid #e2e8f0;">
\(P(x) = {@poly@}\)
</div>
<p><strong>1.</strong> Déterminer les deux racines réelles du polynôme \(P\).</p>
<div style="margin: 15px 0 10px 0; display: flex; align-items: center; gap: 10px; flex-wrap: wrap;">
<span>\(x_1 = \)</span> [[input:ans_r1]] [[validation:ans_r1]]
<span style="font-weight: bold; margin: 0 5px;">et</span>
<span>\(x_2 = \)</span> [[input:ans_r2]] [[validation:ans_r2]]
</div>
<p style="font-size: 0.85rem; color: #64748b; margin-top: -5px;"><em>(L'ordre de saisie des racines n'a pas d'importance).</em></p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<div style="margin-top: 20px; padding: 15px; background: #f8fafc; border: 1px solid #cbd5e1; border-radius: 8px;">
<div style="font-weight: bold; margin-bottom: 10px;">🔑 Correction détaillée</div>
<div style="font-size: .9rem;">
Pour trouver les racines, on résout l'équation \(P(x) = 0\).<br><br>
<strong>1. Identification et Discriminant :</strong><br>
\[
a = {@a@}, \quad b = {@b@}, \quad c = {@c@}
\]
\[
\Delta = b^2 - 4ac = ({@b@})^2 - 4 \times ({@a@}) \times ({@c@}) = {@delta@}
\]<br>
<strong>2. Application des formules des racines :</strong><br>
Puisque \(\Delta > 0\), il y a deux racines distinctes :<br>
\[
x_1 = \frac{-b - \sqrt{\Delta}}{2a} = {@r1@}
\]
\[
x_2 = \frac{-b + \sqrt{\Delta}}{2a} = {@r2@}
\]
</div>
</div>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* Fonction utilitaire */
ri(a,b) := a + rand(b - a + 1) - 1;

/* Génération secrète de deux racines entières distinctes */
r1: ri(-4, 4);
r2: ri(-4, 4);
while r2 = r1 do r2: ri(-4, 4);

/* Ordonner les racines pour un affichage propre dans la correction */
if r1 > r2 then (
    temp: r1,
    r1: r2,
    r2: temp
);

/* Génération du coefficient a (non nul) */
r_a: rand(6);
if r_a = 1 then a: -3
else if r_a = 2 then a: -2
else if r_a = 3 then a: -1
else if r_a = 4 then a: 1
else if r_a = 5 then a: 2
else a: 3;

/* Fabrication du polynôme à l'envers : a(x - r1)(x - r2) */
b: -a * (r1 + r2);
c: a * r1 * r2;
poly: a*x^2 + b*x + c;

/* Calcul du discriminant pour la correction */
delta: b^2 - 4*a*c;

/* Ensemble de réponses pour les tests logiques */
ans_correct_set: {r1, r2};

/* Erreurs classiques (mises sous forme d'ensembles pour le test AlgEquiv) */
/* Erreur 1 : A oublié le "a" au dénominateur */
err_no_a_set: {(-b - sqrt(delta))/2, (-b + sqrt(delta))/2};
/* Erreur 2 : A oublié le signe "-" devant b */
err_no_minus_set: {(b - sqrt(delta))/(2*a), (b + sqrt(delta))/(2*a)};]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<div>[[feedback:prt1]]</div>]]></text>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>P(x)={@poly@}, Racines={@r1@} et {@r2@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p>Généré avec HéStack V2 | hestack-types:fonctions-second-degre-racines</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans_r1</name>
      <type>algebraic</type>
      <tans>r1</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>...</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans_r2</name>
      <type>algebraic</type>
      <tans>r2</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>...</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Les deux racines sont correctes</description>
        <answertest>AlgEquiv</answertest>
        <sans>{ans_r1, ans_r2}</sans>
        <tans>ans_correct_set</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-OK</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #15803d; padding: 10px 14px; background: #f0fdf4; border-radius: 4px;">✅ <strong>Parfait !</strong> Les deux racines sont correctes.</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>PRT-NOK</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description><![CDATA[Erreur : Oubli du coefficient "a" au dénominateur]]></description>
        <answertest>AlgEquiv</answertest>
        <sans>{ans_r1, ans_r2}</sans>
        <tans>err_no_a_set</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-NOA</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Oubli du coefficient \(a\) !</strong> Vos calculs sont bons, mais vous avez divisé par 2 au lieu de \(2a\). Le dénominateur de la formule des racines est \(2a\), il ne faut pas oublier de diviser par {@a@}.</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>PRT-ERR-NOT-NOA</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description><![CDATA[Erreur : Oubli du signe "-" devant b]]></description>
        <answertest>AlgEquiv</answertest>
        <sans>{ans_r1, ans_r2}</sans>
        <tans>err_no_minus_set</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-NOMINUS</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Erreur de signe sur \(b\) !</strong> Vos dénominateurs sont corrects, mais vous avez calculé \(\frac{b \pm \sqrt{\Delta}}{2a}\). La formule commence par un signe <strong>moins</strong> : \(\frac{-b \pm \sqrt{\Delta}}{2a}\).</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>PRT-ERR-CALC</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #dc2626; padding: 10px 14px; background: #fff0f0; border-radius: 4px;">❌ <strong>Incorrect.</strong> Appliquez la méthode :<br>1. Calculez le discriminant \(\Delta = b^2 - 4ac\).<br>2. Appliquez les formules : \(x_1 = \frac{-b - \sqrt{\Delta}}{2a}\) et \(x_2 = \frac{-b + \sqrt{\Delta}}{2a}\).<br><em>Vérifiez bien vos signes et vos divisions.</em></div>]]></text>
        </falsefeedback>
      </node>
    </prt>
  </question>

</quiz>