<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 396367  -->
  <question type="stack">
    <name>
      <text>Fonctions : Nombre de racines réelles d'un polynôme</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<div style="background: #475569; border-left: 5px solid #334155; border-radius: 0 8px 8px 0; padding: 10px 16px; margin-bottom: 12px; display: flex; align-items: center; gap: 10px; flex-wrap: wrap;"><strong style="font-weight: 800; color: #fff; font-size: .95rem;">Fonctions — Second degré</strong> <span style="background: #334155; color: #fff; padding: 2px 9px; border-radius: 20px; font-size: .78rem; font-weight: bold;">/ 1 pt</span></div>
<p>Soit la fonction polynôme du second degré définie sur \(\mathbb{R}\) par :</p>
<div style="text-align: center; font-size: 1.2rem; margin: 25px 0; background: #f8fafc; padding: 15px; border-radius: 8px; border: 1px solid #e2e8f0;">
\(P(x) = {@poly@}\)
</div>
<p>Sans calculer les racines, déterminer le <strong>nombre de racines réelles</strong> du polynôme \(P\).</p>
<div style="margin: 15px 0 10px 0;">Nombre de racines réelles : [[input:ans_nb]] [[validation:ans_nb]]</div>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<div style="margin-top: 20px; padding: 15px; background: #f8fafc; border: 1px solid #cbd5e1; border-radius: 8px;">
<div style="font-weight: bold; margin-bottom: 10px;">🔑 Correction détaillée</div>
<div style="font-size: .9rem;">
Pour connaître le nombre de racines, on calcule le discriminant \(\Delta = b^2 - 4ac\).<br><br>
<strong>1. Identification des coefficients :</strong><br>
\[
a = {@a@}, \quad b = {@b@}, \quad c = {@c@}
\]<br>
<strong>2. Calcul du discriminant :</strong><br>
\[
\Delta = ({@b@})^2 - 4 \times ({@a@}) \times ({@c@}) = {@b_carre@} - {@prod_4ac@} = {@real_delta@}
\]<br>
<strong>3. Conclusion :</strong><br>
Puisque \(\Delta {@signe_delta@} 0\), le polynôme admet {@txt_rep@}.
</div>
</div>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* Génération par conditions séquentielles pures (100% compatible STACK) */
r_type: rand(3)$ r_poly: rand(4)$ 
if r_type = 1 and r_poly = 1 then poly: x^2 - 3*x + 2$ if r_type = 1 and r_poly = 2 then poly: x^2 + 2*x - 3$ if r_type = 1 and r_poly = 3 then poly: 2*x^2 - x - 1$ if r_type = 1 and r_poly = 4 then poly: -x^2 + 4*x - 3$ if r_type = 1 then ans_correct: 2$ 
if r_type = 2 and r_poly = 1 then poly: x^2 - 4*x + 4$ if r_type = 2 and r_poly = 2 then poly: x^2 + 6*x + 9$ if r_type = 2 and r_poly = 3 then poly: 2*x^2 - 4*x + 2$ if r_type = 2 and r_poly = 4 then poly: -x^2 + 2*x - 1$ if r_type = 2 then ans_correct: 1$ 
if r_type = 3 and r_poly = 1 then poly: x^2 + x + 1$ if r_type = 3 and r_poly = 2 then poly: x^2 - 2*x + 5$ if r_type = 3 and r_poly = 3 then poly: 2*x^2 + 2*x + 3$ if r_type = 3 and r_poly = 4 then poly: -x^2 + x - 1$ if r_type = 3 then ans_correct: 0$ 
/* Extraction des coefficients pour la correction */
a: coeff(poly, x, 2)$ b: coeff(poly, x, 1)$ c: coeff(poly, x, 0)$ b_carre: b^2$ prod_4ac: 4*a*c$ real_delta: b^2 - 4*a*c$ 
/* Texte pour la correction */
if ans_correct = 2 then txt_rep: "deux racines réelles distinctes"$ if ans_correct = 2 then signe_delta: ">"$ 
if ans_correct = 1 then txt_rep: "une racine réelle double"$ if ans_correct = 1 then signe_delta: "="$ 
if ans_correct = 0 then txt_rep: "aucune racine réelle"$ if ans_correct = 0 then signe_delta: "<"$ 
/* Erreurs classiques */
err_delta: real_delta$ 
if ans_correct = 2 then err_inv_1_2: 1$ if ans_correct = 1 then err_inv_1_2: 2$ if ans_correct = 0 then err_inv_1_2: 999$ 
if ans_correct = 1 or ans_correct = 2 then err_0_pos: 0$ if ans_correct = 0 then err_0_pos: 999$]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<div>[[feedback:prt1]]</div>]]></text>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>P(x)={@poly@}, Delta={@real_delta@}, Nb racines={@ans_correct@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p>Généré avec HéStack V2 | hestack-types:fonctions-second-degre-nb-racines</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans_nb</name>
      <type>algebraic</type>
      <tans>ans_correct</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>0, 1 ou 2</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Réponse correcte</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_nb</sans>
        <tans>ans_correct</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-OK</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #15803d; padding: 10px 14px; background: #f0fdf4; border-radius: 4px;">✅ <strong>Parfait !</strong> Le nombre de racines est correct.</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>PRT-NOK</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Erreur : A donné la valeur de Delta</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_nb</sans>
        <tans>err_delta</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-DELTA</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Confusion !</strong> Vous avez donné la <strong>valeur du discriminant \(\Delta\)</strong>. On vous demande le <strong>nombre de racines</strong> (qui doit être 0, 1 ou 2).</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>PRT-ERR-NOT-DELTA</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description><![CDATA[Erreur : Confusion Delta = 0 et Delta > 0]]></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_nb</sans>
        <tans>err_inv_1_2</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-INV</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Confusion de cas !</strong> Vous avez confondu le cas \(\Delta = 0\) et le cas \(\Delta > 0\).<br>- Si \(\Delta > 0\) : il y a <strong>2</strong> racines distinctes.<br>- Si \(\Delta = 0\) : il y a <strong>1</strong> racine double.</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>PRT-ERR-NOT-INV</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Erreur : A répondu 0 alors qu'il y en avait</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_nb</sans>
        <tans>err_0_pos</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-0POS</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Erreur de signe probable !</strong> Vous avez répondu 0 racine, ce qui signifie que vous avez trouvé \(\Delta < 0\). Le polynôme possède pourtant des racines. Vérifiez le signe "moins" dans la formule : \(\Delta = b^2 \mathbf{-} 4ac\).</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>PRT-ERR-CALC</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #dc2626; padding: 10px 14px; background: #fff0f0; border-radius: 4px;">❌ <strong>Incorrect.</strong> Appliquez la méthode :<br>1. Calculez \(\Delta = b^2 - 4ac\).<br>2. Regardez le signe de \(\Delta\) :<br>   - Si \(\Delta > 0\) : 2 racines<br>   - Si \(\Delta = 0\) : 1 racine<br>   - Si \(\Delta < 0\) : 0 racine</div>]]></text>
        </falsefeedback>
      </node>
    </prt>
  </question>

</quiz>