<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 396472  -->
  <question type="stack">
    <name>
      <text>Parsons - Fonction Moyenne Python</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<div style="background:#475569;border-left:5px solid #334155;border-radius:0 8px 8px 0;padding:10px 16px;margin-bottom:12px;display:flex;align-items:center;gap:10px;flex-wrap:wrap;">
        <strong style="font-weight:800;color:#fff;font-size:.95rem;">Algorithmique — Parsons Puzzle</strong>
        <span style="background:#334155;color:#fff;padding:2px 9px;border-radius:20px;font-size:.78rem;font-weight:700;">/ 1 pt</span>
      </div>
      
      <p><strong>Consigne :</strong> Reconstituez la fonction Python <code>calculer_moyenne(notes)</code> qui renvoie la moyenne d'une liste. La fonction doit renvoyer <code>0</code> si la liste est vide.</p>
      
[[parsons input="ans1"]]
{# parsons_encode(proof_steps) #}
[[/parsons ]]
<p>[[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<div style="margin-top:20px; padding:15px; background:#f8fafc; border:1px solid #cbd5e1; border-radius:8px;">
        <div style="font-weight:bold; margin-bottom:10px;">🔑 Correction détaillée</div>
        <div style="font-size:.9rem;">
          Voici le code correct reconstruit :
          <div class="proof">{@proof_display_para(ta, proof_steps)@}</div>
          <br/>Les lignes <code>somme = 0</code> et <code>return somme / len(notes)</code> étaient des distracteurs.
        </div>
      </div>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2024092500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[stack_include_contrib("prooflib.mac");

/* 1. Définition de tous les blocs (bons et distracteurs) */
proof_steps: [
    ["L1", "def calculer_moyenne(notes):"],
    ["L2", "    if len(notes) == 0:"],
    ["L3", "        return 0"],
    ["L4", "    somme = sum(notes)"],
    ["L5", "    moyenne = somme / len(notes)"],
    ["L6", "    return moyenne"],
    ["D1", "    somme = 0"],
    ["D2", "    return somme / len(notes)"]
];

/* 2. Définition de la structure de la bonne réponse 
   Même pour une séquence simple, il faut utiliser proof() */
ta: proof("L1", "L2", "L3", "L4", "L5", "L6");

/* 3. Mélange aléatoire des blocs pour l'affichage élève */
proof_steps: random_permutation(proof_steps);

/* 4. Génération des alternatives (utile si on utilisait proof_iff) */
tal: proof_alternatives(ta);

/* 5. Création d'un ensemble contenant les preuves aplaties pour l'évaluation */
tas: setify(map(proof_flatten, tal));]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="html">
      <text>{@ map(first, proof_steps) @}</text>
    </questionnote>
    <questiondescription format="html">
      <text></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="moodle_auto_format">
      <text></text>
    </prtcorrect>
    <prtpartiallycorrect format="moodle_auto_format">
      <text></text>
    </prtpartiallycorrect>
    <prtincorrect format="moodle_auto_format">
      <text></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>parsons</type>
      <tans>[ta, proof_steps]</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/* Décode la réponse de l'élève et vérifie si elle appartient aux preuves valides */
sa: parsons_decode(ans1);
check: elementp(sa, tas);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Vérification de la preuve Parsons</description>
        <answertest>AlgEquiv</answertest>
        <sans>check</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left:4px solid #15803d;padding:10px 14px;background:#f0fdf4;border-radius:4px;">✅ <strong>Parfait !</strong> Vous avez reconstitué le code exact dans le bon ordre.</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<div style="border-left:4px solid #dc2626;padding:10px 14px;background:#fff0f0;border-radius:4px;">❌ <strong>Incorrect.</strong> Vérifiez l'ordre des lignes et assurez-vous de ne pas avoir utilisé les lignes inutiles.</div>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>apply(parsons_answer,[ta, proof_steps])</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>

</quiz>