<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 156  -->
  <question type="stack">
    <name>
      <text>avancement</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>On étudie la réaction d'oxydoréduction entre l'eau oxygénée et le permanganate de potassium en milieu acide :</p>
<p style="text-align: center;">\( 5H_2O_2 + 2MnO_4^- + 6H^+ \rightarrow 5O_2 + 2Mn^{2+} + 8H_2O \)</p>
<p>On dispose initialement de {@n0_val@} mol d'eau oxygénée, de 0 mol de permanganate, et d'un <b>grand excès d'ions \(H^+\)</b> que l'on fixera arbitrairement à 100 mol pour le calcul. (L'eau est le solvant, elle n'apparaît pas).</p>
<p>Complétez le tableau d'avancement ci-dessous (utilisez la variable <b>x</b> pour la ligne "En cours"). Ne mettez pas d'unités.</p>

<table border="1" style="text-align: center; margin: auto; border-collapse: collapse; font-size: 0.85em;">
  <tr style="background-color: #f2f2f2;">
    <th style="padding: 6px;">État</th>
    <th style="padding: 6px;">\(n(H_2O_2)\)</th>
    <th style="padding: 6px;">\(n(MnO_4^-)\)</th>
    <th style="padding: 6px;">\(n(H^+)\)</th>
    <th style="padding: 6px;">\(n(O_2)\)</th>
    <th style="padding: 6px;">\(n(Mn^{2+})\)</th>
  </tr>
  <tr>
    <td style="padding: 6px;">Initial</td>
    <td style="padding: 6px;">[[input:i1]]</td>
    <td style="padding: 6px;">[[input:i2]]</td>
    <td style="padding: 6px;">[[input:i3]]</td>
    <td style="padding: 6px;">[[input:i4]]</td>
    <td style="padding: 6px;">[[input:i5]]</td>
  </tr>
  <tr>
    <td style="padding: 6px;">En cours</td>
    <td style="padding: 6px;">[[input:c1]]</td>
    <td style="padding: 6px;">[[input:c2]]</td>
    <td style="padding: 6px;">[[input:c3]]</td>
    <td style="padding: 6px;">[[input:c4]]</td>
    <td style="padding: 6px;">[[input:c5]]</td>
  </tr>
  <tr>
    <td style="padding: 6px;">Final</td>
    <td style="padding: 6px;">[[input:f1]]</td>
    <td style="padding: 6px;">[[input:f2]]</td>
    <td style="padding: 6px;">[[input:f3]]</td>
    <td style="padding: 6px;">[[input:f4]]</td>
    <td style="padding: 6px;">[[input:f5]]</td>
  </tr>
</table>

<p>Avancement maximal \(x_{max}\) = [[input:xmax]]</p>

<div>[[validation:i1]] [[validation:i2]] [[validation:i3]] [[validation:i4]] [[validation:i5]] 
[[validation:c1]] [[validation:c2]] [[validation:c3]] [[validation:c4]] [[validation:c5]] 
[[validation:f1]] [[validation:f2]] [[validation:f3]] [[validation:f4]] [[validation:f5]] [[validation:xmax]]</div>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p><b>Correction du tableau d'avancement :</b></p>

<table border="1" style="text-align: center; margin: auto; border-collapse: collapse; font-size: 0.85em;">
  <tr style="background-color: #f2f2f2;">
    <th style="padding: 6px;">État</th>
    <th style="padding: 6px;">\(n(H_2O_2)\)</th>
    <th style="padding: 6px;">\(n(MnO_4^-)\)</th>
    <th style="padding: 6px;">\(n(H^+)\)</th>
    <th style="padding: 6px;">\(n(O_2)\)</th>
    <th style="padding: 6px;">\(n(Mn^{2+})\)</th>
  </tr>
  <tr>
    <td style="padding: 6px;">Initial</td>
    <td style="padding: 6px;">{@n0_val@}</td>
    <td style="padding: 6px;">0</td>
    <td style="padding: 6px;">{@nH_i@}</td>
    <td style="padding: 6px;">0</td>
    <td style="padding: 6px;">0</td>
  </tr>
  <tr>
    <td style="padding: 6px;">En cours</td>
    <td style="padding: 6px;">{@n0_val@} - 5\(x\)</td>
    <td style="padding: 6px;">2\(x\)</td>
    <td style="padding: 6px;">{@nH_i@} - 6\(x\)</td>
    <td style="padding: 6px;">5\(x\)</td>
    <td style="padding: 6px;">2\(x\)</td>
  </tr>
  <tr>
    <td style="padding: 6px;">Final</td>
    <td style="padding: 6px;">0</td>
    <td style="padding: 6px;">0</td>
    <td style="padding: 6px;">{@nf_h@}</td>
    <td style="padding: 6px;">{@nf_o2@}</td>
    <td style="padding: 6px;">{@nf_mn2@}</td>
  </tr>
</table>

<p><b>Calcul de l'avancement maximal :</b></p>
<p>Le réactif limitant est l'eau oxygénée (le permanganate est ajouté au fur et à mesure et les \(H^+\) sont en excès). À l'état final, \(n(H_2O_2) = 0\).</p>
<p>Ce qui donne : {@n0_val@} - 5\(x_{max}\) = 0</p>
<p>Donc \(x_{max} = \dfrac{@n0_val@}{5} = {@val_xmax@}\) mol.</p>]]></text>
    </generalfeedback>
    <defaultgrade>4</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>n0_val : rand(4) + 2;
nH_i : 100; /* On fixe H+ initiale à 100 pour simuler l'excès */

/* Etat initial */
ni_h2o2 : n0_val;
ni_mno4 : 0;
ni_h : nH_i;
ni_o2 : 0;
ni_mn2 : 0;

/* Etat en cours (litteral avec x) */
nc_h2o2 : n0_val - 5*x;
nc_mno4 : 2*x;
nc_h : nH_i - 6*x;
nc_o2 : 5*x;
nc_mn2 : 2*x;

/* Avancement maximal */
val_xmax : float(n0_val / 5);

/* Etat final */
nf_h2o2 : ev(nc_h2o2, x = val_xmax);
nf_mno4 : ev(nc_mno4, x = val_xmax);
nf_h : ev(nc_h, x = val_xmax);
nf_o2 : ev(nc_o2, x = val_xmax);
nf_mn2 : ev(nc_mn2, x = val_xmax);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="html">
      <text>{@ta1@}</text>
    </questionnote>
    <questiondescription format="html">
      <text></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Réponse correcte, bravo !]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Réponse partiellement correcte.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Réponse incorrecte.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>c1</name>
      <type>algebraic</type>
      <tans>nc_h2o2</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>c2</name>
      <type>algebraic</type>
      <tans>nc_mno4</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>c3</name>
      <type>algebraic</type>
      <tans>nc_h</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>c4</name>
      <type>algebraic</type>
      <tans>nc_o2</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>c5</name>
      <type>algebraic</type>
      <tans>nc_mn2</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>f1</name>
      <type>numerical</type>
      <tans>nf_h2o2</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>f2</name>
      <type>numerical</type>
      <tans>nf_mno4</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>f3</name>
      <type>numerical</type>
      <tans>nf_h</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>f4</name>
      <type>numerical</type>
      <tans>nf_o2</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>f5</name>
      <type>numerical</type>
      <tans>nf_mn2</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>i1</name>
      <type>numerical</type>
      <tans>ni_h2o2</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>i2</name>
      <type>numerical</type>
      <tans>ni_mno4</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>i3</name>
      <type>numerical</type>
      <tans>ni_h</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>i4</name>
      <type>numerical</type>
      <tans>ni_o2</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>i5</name>
      <type>numerical</type>
      <tans>ni_mn2</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>xmax</name>
      <type>numerical</type>
      <tans>val_xmax</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>2</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>ligne initiale</description>
        <answertest>AlgEquiv</answertest>
        <sans>[i1, i2, i3, i4, i5]</sans>
        <tans>[ni_h2o2, ni_mno4, ni_h, ni_o2, ni_mn2]</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1/3</truescore>
        <truepenalty></truepenalty>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<b>Ligne "Initial" correcte.</b>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<b>Ligne "Initial" incorrecte.</b>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>ligne intermédiaire</description>
        <answertest>AlgEquiv</answertest>
        <sans>[c1, c2, c3, c4, c5]</sans>
        <tans>[nc_h2o2, nc_mno4, nc_h, nc_o2, nc_mn2]</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1/3</truescore>
        <truepenalty></truepenalty>
        <truenextnode>2</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<b>Ligne "En cours"correcte.</b>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<b>Ligne "En cours" incorrecte.</b> Utilisez bien la variable x et les coefficients stœchiométriques.]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Ligne Finale</description>
        <answertest>AlgEquiv</answertest>
        <sans>[f1, f2, f3, f4, f5]</sans>
        <tans>[nf_h2o2, nf_mno4, nf_h, nf_o2, nf_mn2]</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1/3</truescore>
        <truepenalty></truepenalty>
        <truenextnode>3</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<b>Ligne "Final"correcte.</b>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<b>Ligne "Final" incorrecte.</b> Avez-vous bien remplacé x par sa valeur maximale ?]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>xmax</sans>
        <tans>val_xmax</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<b>x_max correct.</b>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<b>x_max incorrect.</b>]]></text>
        </falsefeedback>
      </node>
    </prt>
  </question>

</quiz>