<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name><text>SVT - Datation absolue (radiochronologie)</text></name>
    <questiontext format="html">
      <text><![CDATA[
<p>Dans un échantillon de roche, on mesure qu'il reste une proportion <b>Nfrac = {@n_frac@}</b> de l'isotope radioactif parent par rapport à la quantité initiale. La constante de désintégration radioactive de cet isotope est <b>lam = {@lam_a@} &times; 10<sup>-{@lam_b@}</sup> an<sup>-1</sup></b>.</p>
<p>a) Écris l'expression littérale de l'âge <i>T</i> de la roche (en années) en fonction de <i>Nfrac</i> et <i>lam</i>.<br/>[[input:ans1]] [[validation:ans1]]</p>
<p>b) Calcule la valeur numérique de cet âge <i>T</i>, en années.<br/>[[input:ans2]] [[validation:ans2]]</p>
]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[
<p>La loi de décroissance radioactive s'écrit \( N/N_0 = e^{-\lambda t} \), soit \( Nfrac = e^{-lam \cdot T} \). En isolant \( T \) : \( T = -\ln(Nfrac)/lam \). Comme \( Nfrac < 1 \), \( \ln(Nfrac) \) est négatif, d'où le signe moins devant pour obtenir un âge positif.</p>
]]> </text>
    </generalfeedback>
    <defaultgrade>2</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion><text>2024092500</text></stackversion>
    <questionvariables>
      <text><![CDATA[
frac_list: [1/2,1/4,1/8,1/16,3/5,7/10,4/5]$
n_frac: frac_list[1+random(length(frac_list))]$
lama_list: [1,2,5,8]$
lam_a: lama_list[1+random(length(lama_list))]$
lamb_list: [4,5,6]$
lam_b: lamb_list[1+random(length(lamb_list))]$
lam_val: lam_a*10^(-lam_b)$
ta_formula: -log(Nfrac)/lam$
err_sign_formula: log(Nfrac)/lam$
t_val: -log(n_frac)/lam_val$
err_sign_val: log(n_frac)/lam_val$
      ]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<div>[[feedback:prt1]] [[feedback:prt2]]</div> ]]></text>
    </specificfeedback>

    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta_formula</tans>
      <boxsize>20</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords>Nfrac,lam</allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>2</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans2</name>
      <type>numerical</type>
      <tans>t_val</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>

    <prt>
      <name>prt1</name>
      <value>1</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables><text></text></feedbackvariables>
      <nodes>
        <node>
          <name>0</name>
          <description>formule correcte</description>
          <answertest>AlgEquiv</answertest>
          <sans>ans1</sans>
          <tans>ta_formula</tans>
          <testoptions></testoptions>
          <quiet>0</quiet>
          <truescoremode>=</truescoremode>
          <truescore>1</truescore>
          <truepenalty></truepenalty>
          <truenextnode>-1</truenextnode>
          <trueanswernote>prt1-0-T</trueanswernote>
          <truefeedback format="html"><text><![CDATA[<p>Correct.</p> ]]></text></truefeedback>
          <falsescoremode>=</falsescoremode>
          <falsescore>0</falsescore>
          <falsepenalty></falsepenalty>
          <falsenextnode>1</falsenextnode>
          <falseanswernote>prt1-0-F</falseanswernote>
          <falsefeedback format="html"><text><![CDATA[ ]]></text></falsefeedback>
        </node>
        <node>
          <name>1</name>
          <description>signe moins oublié (formule)</description>
          <answertest>AlgEquiv</answertest>
          <sans>ans1</sans>
          <tans>err_sign_formula</tans>
          <testoptions></testoptions>
          <quiet>0</quiet>
          <truescoremode>=</truescoremode>
          <truescore>0</truescore>
          <truepenalty></truepenalty>
          <truenextnode>-1</truenextnode>
          <trueanswernote>prt1-1-T</trueanswernote>
          <truefeedback format="html"><text><![CDATA[<p>Il te manque le signe moins : comme <i>Nfrac</i> &lt; 1, \( \ln(Nfrac) \) est négatif, il faut le signe moins devant pour obtenir un âge <i>T</i> positif.</p> ]]></text></truefeedback>
          <falsescoremode>=</falsescoremode>
          <falsescore>0</falsescore>
          <falsepenalty></falsepenalty>
          <falsenextnode>-1</falsenextnode>
          <falseanswernote>prt1-1-F</falseanswernote>
          <falsefeedback format="html"><text><![CDATA[<p>Incorrect. Isole <i>T</i> dans \( Nfrac = e^{-lam \cdot T} \) en passant par le logarithme népérien.</p> ]]></text></falsefeedback>
        </node>
      </nodes>
    </prt>

    <prt>
      <name>prt2</name>
      <value>1</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables><text></text></feedbackvariables>
      <nodes>
        <node>
          <name>0</name>
          <description>valeur numérique correcte</description>
          <answertest>NumRelative</answertest>
          <sans>ans2</sans>
          <tans>t_val</tans>
          <testoptions>0.01</testoptions>
          <quiet>0</quiet>
          <truescoremode>=</truescoremode>
          <truescore>1</truescore>
          <truepenalty></truepenalty>
          <truenextnode>-1</truenextnode>
          <trueanswernote>prt2-0-T</trueanswernote>
          <truefeedback format="html"><text><![CDATA[<p>Correct. Toutes les écritures numériques sont acceptées (ex : notation scientifique ou décimale).</p> ]]></text></truefeedback>
          <falsescoremode>=</falsescoremode>
          <falsescore>0</falsescore>
          <falsepenalty></falsepenalty>
          <falsenextnode>1</falsenextnode>
          <falseanswernote>prt2-0-F</falseanswernote>
          <falsefeedback format="html"><text><![CDATA[ ]]></text></falsefeedback>
        </node>
        <node>
          <name>1</name>
          <description>signe moins oublié (numérique, résultat négatif)</description>
          <answertest>NumRelative</answertest>
          <sans>ans2</sans>
          <tans>err_sign_val</tans>
          <testoptions>0.01</testoptions>
          <quiet>0</quiet>
          <truescoremode>=</truescoremode>
          <truescore>0</truescore>
          <truepenalty></truepenalty>
          <truenextnode>-1</truenextnode>
          <trueanswernote>prt2-1-T</trueanswernote>
          <truefeedback format="html"><text><![CDATA[<p>Ton résultat est négatif : un âge ne peut pas être négatif. Tu as sans doute oublié le signe moins devant le logarithme.</p> ]]></text></truefeedback>
          <falsescoremode>=</falsescoremode>
          <falsescore>0</falsescore>
          <falsepenalty></falsepenalty>
          <falsenextnode>-1</falsenextnode>
          <falseanswernote>prt2-1-F</falseanswernote>
          <falsefeedback format="html"><text><![CDATA[<p>Incorrect. Reprends la formule de a) et substitue les valeurs numériques de l'énoncé.</p> ]]></text></falsefeedback>
        </node>
      </nodes>
    </prt>
  </question>
</quiz>
