<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name><text>SVT - Titrage et dilutions en série (Immunologie)</text></name>
    <questiontext format="html">
      <text><![CDATA[
<p>Pour titrer un sérum, on réalise une série de dilutions successives au <b>1/10</b> : chaque tube reçoit un dixième de la concentration du tube précédent.</p>
<p>Écris la formule du facteur de dilution du sérum dans le tube n°<b>{@n_tube@}</b> (le tube n°1 correspondant au sérum non dilué).<br/>[[input:ans1]] [[validation:ans1]]</p>
]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[
<p>Les facteurs de dilution se <b>multiplient</b> à chaque étape : après <i>n</i> dilutions au 1/10, le facteur global est \( 1/10^n \) (et non \( 1/(10n) \), qui reviendrait à additionner les facteurs).</p>
]]> </text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion><text>2024092500</text></stackversion>
    <questionvariables>
      <text><![CDATA[
ntube_list: [2,3,4,5,6]$
n_tube: ntube_list[1+random(length(ntube_list))]$
dilution_correct: 1/10^n_tube$
err_addition: 1/(10*n_tube)$
      ]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<div>[[feedback:prt1]]</div> ]]></text>
    </specificfeedback>

    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>dilution_correct</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>

    <prt>
      <name>prt1</name>
      <value>1</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables><text></text></feedbackvariables>
      <nodes>
        <node>
          <name>0</name>
          <description>facteur de dilution correct</description>
          <answertest>AlgEquiv</answertest>
          <sans>ans1</sans>
          <tans>dilution_correct</tans>
          <testoptions></testoptions>
          <quiet>0</quiet>
          <truescoremode>=</truescoremode>
          <truescore>1</truescore>
          <truepenalty></truepenalty>
          <truenextnode>-1</truenextnode>
          <trueanswernote>prt1-0-T</trueanswernote>
          <truefeedback format="html"><text><![CDATA[<p>Correct : les facteurs 1/10 se multiplient à chaque tube, soit \( 1/10^n \).</p> ]]></text></truefeedback>
          <falsescoremode>=</falsescoremode>
          <falsescore>0</falsescore>
          <falsepenalty></falsepenalty>
          <falsenextnode>1</falsenextnode>
          <falseanswernote>prt1-0-F</falseanswernote>
          <falsefeedback format="html"><text><![CDATA[ ]]></text></falsefeedback>
        </node>
        <node>
          <name>1</name>
          <description>addition au lieu de multiplication des facteurs</description>
          <answertest>AlgEquiv</answertest>
          <sans>ans1</sans>
          <tans>err_addition</tans>
          <testoptions></testoptions>
          <quiet>0</quiet>
          <truescoremode>=</truescoremode>
          <truescore>0</truescore>
          <truepenalty></truepenalty>
          <truenextnode>-1</truenextnode>
          <trueanswernote>prt1-1-T</trueanswernote>
          <truefeedback format="html"><text><![CDATA[<p>Tu as additionné les facteurs de dilution (10 × n) au lieu de les multiplier. À chaque tube, la dilution au 1/10 précédente est encore divisée par 10 : le facteur global est \( 1/10^n \), pas \( 1/(10n) \).</p> ]]></text></truefeedback>
          <falsescoremode>=</falsescoremode>
          <falsescore>0</falsescore>
          <falsepenalty></falsepenalty>
          <falsenextnode>-1</falsenextnode>
          <falseanswernote>prt1-1-F</falseanswernote>
          <falsefeedback format="html"><text><![CDATA[<p>Incorrect. Chaque tube reprend le facteur du tube précédent et le multiplie par 1/10 : après <i>n</i> dilutions, le facteur global est \( 1/10^n \).</p> ]]></text></falsefeedback>
        </node>
      </nodes>
    </prt>
  </question>
</quiz>
