<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 396092  -->
  <question type="stack">
    <name>
      <text>Q2 - Algèbre : Factorisation Maximale</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<div style="background: #0284c7; border-left: 5px solid #0369a1; border-radius: 0 8px 8px 0; padding: 10px 16px; margin-bottom: 12px; display: flex; align-items: center; gap: 10px; flex-wrap: wrap;"><strong style="font-weight: 800; color: #fff; font-size: .95rem;">Q2 — Algèbre</strong> <span style="background: #0369a1; color: #fff; padding: 2px 9px; border-radius: 20px; font-size: .78rem; font-weight: bold;">/ 1 pt</span> <span style="background: #ffffff; color: #0369a1; border: 1px solid #0369a1; padding: 2px 9px; border-radius: 20px; font-size: .75rem; font-weight: 600;">#️⃣ Factorisation maximale</span></div>
<p>Factoriser <strong>au maximum</strong> l'expression mathématique suivante : <br><br>\( A(x) = {@exp1@} \)</p>
<p>[[input:ans1]] [[validation:ans1]]</p>
<p>[[iframe width="100%" height="160px" scrolling="false"]] [[style]] body { margin: 0; padding: 0; font-family: system-ui, -apple-system, sans-serif; } .stack-keyboard-container { background: #f8fafc; border: 1px solid #e2e8f0; border-radius: 8px; padding: 16px; max-width: 600px; box-sizing: border-box; margin-top: 10px; } .stack-keyboard-title { font-size: 0.85rem; font-weight: 700; color: #475569; margin-bottom: 10px; text-transform: uppercase; letter-spacing: 0.05em; } .stack-keyboard-row { display: flex; flex-wrap: wrap; gap: 6px; } .btn-ins { background: #ffffff !important; border: 1px solid #cbd5e1 !important; border-radius: 6px !important; color: #334155 !important; padding: 6px 12px !important; font-size: 0.9rem !important; font-weight: 600 !important; cursor: pointer !important; transition: all 0.15s ease !important; min-width: 38px !important; box-shadow: 0 1px 2px rgba(0,0,0,0.05) !important; user-select: none; } .btn-ins:hover { background: #f0f9ff !important; border-color: #7dd3fc !important; color: #0284c7 !important; } .btn-ins:active { background: #e0f2fe !important; } .btn-var { background: #f1f5f9 !important; border-color: #94a3b8 !important; color: #0f172a !important; font-style: italic; } [[/style]]</p>
<div class="stack-keyboard-container">
<div class="stack-keyboard-title">🧮 Aide à la saisie algébrique</div>
<div class="stack-keyboard-row"><button class="btn-ins btn-var" type="button" data-val="x">x</button> <button class="btn-ins" type="button" data-val="+">+</button> <button class="btn-ins" type="button" data-val="-">-</button> <button class="btn-ins" type="button" data-val="*">×</button> <button class="btn-ins" type="button" data-val="^2">²</button> <button class="btn-ins" type="button" data-val="^">xⁿ</button> <button class="btn-ins" type="button" data-val="(">(</button> <button class="btn-ins" type="button" data-val=")">)</button></div>
</div>
<p>[[script type="module"]] import {stack_js} from '[[cors src="stackjsiframe.js"/]]'; stack_js.request_access_to_input("ans1", true).then(function(input_id) { var stackInput = document.getElementById(input_id); document.querySelectorAll('.btn-ins').forEach(function(btn){ btn.addEventListener('click', function(e){ e.preventDefault(); var val = this.getAttribute('data-val'); var actuel = stackInput.value || ''; var texteAvant = actuel.trimEnd(); var charPrecedent = texteAvant.length &gt; 0 ? texteAvant.slice(-1) : null; var excl = [null, '=', '+', '-', '*', '/', '^', '(']; var estOp = /^[+\-*\/^)]/.test(val); var pref = ''; if (!excl.includes(charPrecedent) &amp;&amp; !estOp &amp;&amp; actuel.length &gt; 0) { pref = '*'; } stackInput.value = actuel + pref + val; stackInput.dispatchEvent(new Event('change')); }); }); }); [[/script]] [[/iframe]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<div style="margin-top: 20px; padding: 15px; background: #f8fafc; border: 1px solid #cbd5e1; border-radius: 8px;">
<div style="font-weight: bold; margin-bottom: 10px;">🔑 Correction détaillée</div>
<div style="font-size: .9rem;">On reconnaît l'identité remarquable \( a^2 - b^2 = (a-b)(a+b) \).<br>Ici, \( a = (2x+3) \) et \( b = 5 \) (car \(25 = 5^2\)).<br><br>\( A(x) = (2x+3)^2 - 5^2 \)<br>\( A(x) = [(2x+3) - 5][(2x+3) + 5] \)<br>\( A(x) = (2x - 2)(2x + 8) \)<br><br><strong>Pour factoriser au maximum :</strong> on doit extraire le facteur commun numérique de chaque parenthèse.<br>\( 2x - 2 = 2(x - 1) \)<br>\( 2x + 8 = 2(x + 4) \)<br>En regroupant les constantes : \( A(x) = 2 \times 2 \times (x - 1)(x + 4) = 4(x - 1)(x + 4) \).</div>
</div>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>exp1: (2*x+3)^2 - 25;
ta1: 4*(x-1)*(x+4);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<ol>
<li>[[feedback:prt1]]</li>
</ol>]]></text>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>Factorisation de {@exp1@} en {@ta1@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p>Généré avec HéStack V2 | hestack-types:algebraic</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>2</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>1</truenextnode>
        <trueanswernote>PRT-EQUAL-GO-ON</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>PRT-WRONG</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #dc2626; padding: 10px 14px; background: #fff0f0; border-radius: 4px;">❌ <strong>Réponse incorrecte.</strong> Votre expression n'est pas mathématiquement égale à \( A(x) \).</div>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description></description>
        <answertest>FacForm</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions>x</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-FAC-MAX</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #15803d; padding: 10px 14px; background: #f0fdf4; border-radius: 4px;">✅ <strong>Excellent !</strong> L'expression est correctement et totalement factorisée.</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.25</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>PRT-NOT-MAX</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #eab308; padding: 10px 14px; background: #fefce8; border-radius: 4px;">⚠ <strong>Calcul correct mais incomplet.</strong> Votre expression est égale à \( A(x) \), mais elle n'est pas factorisée au maximum. Pensez à extraire les facteurs communs numériques de chaque parenthèse.</div>]]></text>
        </falsefeedback>
      </node>
    </prt>
  </question>

</quiz>