<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 89  -->
  <question type="stack">
    <name>
      <text>Télescope de Newton — construction complète (miroir primaire + miroir secondaire + oculaire)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<div class="stack-comment">
<h2>Construction complète du télescope de Newton</h2>
<p>Le télescope de Newton comporte trois éléments optiques, chacun avec sa propre figure interactive ci-dessous :
un <strong>miroir primaire</strong> concave (distance focale f₁ = SF₁' = 6, sommet S), qui reçoit la lumière d'un
objet à l'infini (faisceau parallèle, angle α avec l'axe) et en forme une image intermédiaire A₁B₁ ;
un <strong>miroir secondaire plan</strong>, incliné à 45°, placé avant le foyer du primaire pour renvoyer le
faisceau convergent sur le côté (image réelle A₂B₂, repliée à 90°) ; et un <strong>oculaire</strong> convergent
(distance focale f₂ = 3) qui reçoit A₂B₂ exactement dans son plan focal objet et rejette l'image finale à
l'infini (télescope afocal, observation à l'œil détendu).</p>
</div>

[[jsxgraph input-ref-ans1="state" width="46em" aspect-ratio="1.7"]]
const board = JXG.JSXGraph.initBoard(divid, {
    boundingbox: [-13, 6, 3, -6],
    axis: false,
    keepaspectratio: true,
    showNavigation: true,
    zoom: { enabled: true, wheel: true, needShift: false, factorX: 1.25, factorY: 1.25 },
    pan: { enabled: true, needTwoFingers: false, needShift: true }
});

const f1 = 6, tanA = 1/4, mirrorH1 = 4;
const xS1 = 0, xC1 = -12, xF1p = -6, xA1 = -6, yB1 = f1 * tanA;
const X_MIN = -13, X_MAX = 3, Y_MIN = -6, Y_MAX = 6;
const LENS_X = [xS1];
const lensHeight = mirrorH1;

let toolMode = '';
let tempPoint = null, dirPoint1 = null, dirPoint2 = null, interSeg1 = null;
let selectedSegment = null;

let allDrawnElements = [];
let raySegments = [];
let logicalRays = [];
let standaloneElements = [];
let intersectionCounter = 0;

board.create('line', [[X_MIN, 0], [X_MAX, 0]], { strokeColor: 'black', strokeWidth: 1, fixed: true, highlight: false, tabindex: null });

board.create('segment', [[xS1, -mirrorH1], [xS1, mirrorH1]], { strokeColor: 'black', strokeWidth: 3, fixed: true, highlight: false, tabindex: null });
board.create('segment', [[xS1, mirrorH1], [xS1 - 0.4, mirrorH1 + 0.4]], { strokeColor: 'black', strokeWidth: 3, fixed: true, highlight: false, tabindex: null });
board.create('segment', [[xS1, -mirrorH1], [xS1 - 0.4, -mirrorH1 - 0.4]], { strokeColor: 'black', strokeWidth: 3, fixed: true, highlight: false, tabindex: null });
for (let i = -4; i <= 4; i++) {
    let y = i * (mirrorH1 / 4);
    board.create('segment', [[xS1, y], [xS1 + 0.4, y - 0.4]], { strokeColor: 'black', strokeWidth: 1.5, fixed: true, highlight: false, tabindex: null });
}

board.create('point', [xS1, 0], { name: 'S', size: 3, fixed: true, color: 'black', highlight: false, tabindex: null });
board.create('point', [xC1, 0], { name: 'C', size: 3, fixed: true, color: 'green', highlight: false, tabindex: null });
board.create('point', [xF1p, 0], { name: "F₁'", size: 3, fixed: true, color: 'green', highlight: false, tabindex: null });

board.create('segment', [[X_MIN + 0.2, (X_MIN + 0.2 - xC1) * tanA], [-1, (-1 - xC1) * tanA]], { strokeColor: '#c0392b', strokeWidth: 2, fixed: true, highlight: false, tabindex: null });
board.create('text', [X_MIN + 0.3, (X_MIN + 0.3 - xC1) * tanA + 0.5, "Objet à l'infini (angle α avec l'axe)"], { fixed: true, fontSize: 11, color: '#c0392b', highlight: false, tabindex: null });

const handlePoint = board.create('point', [0, -5.7], { visible: false, fixed: true, name: '', tabindex: null });

const DEFAULT_MSG = "Cliquez sur un tronçon de droite pour basculer son statut (réel/virtuel) ; il reste sélectionné pour le bouton Effacer. Cliquez sur un point pour le supprimer. Molette : zoom. Maj + glisser : déplacer la vue.";

const instructionsEl = document.createElement('p');
instructionsEl.style.cssText = 'margin:.6em 0 0;font-size:.85em;color:#333;';
instructionsEl.textContent = DEFAULT_MSG;

function setInstructions(msg) {
    instructionsEl.textContent = msg;
}

const toolbarDiv = document.createElement('div');
toolbarDiv.style.cssText = 'display:flex;flex-wrap:wrap;gap:.4em;margin-top:.6em;';
const toolButtons = {};

function addToolButton(label, mode, msgOrHandler) {
    const btn = document.createElement('button');
    btn.type = 'button';
    btn.textContent = label;
    btn.style.cssText = 'padding:.35em .7em;font-size:.85em;cursor:pointer;';
    if (mode === null) {
        btn.addEventListener('click', msgOrHandler);
    } else {
        btn.addEventListener('click', () => activateTool(mode, msgOrHandler));
        toolButtons[mode] = btn;
    }
    toolbarDiv.appendChild(btn);
    return btn;
}

function setActiveButton(mode) {
    for (let m in toolButtons) {
        toolButtons[m].style.background = (m === mode) ? '#dbeafe' : '';
        toolButtons[m].style.fontWeight = (m === mode) ? 'bold' : 'normal';
    }
}

function resetTool() {
    toolMode = '';
    tempPoint = null; dirPoint1 = null; dirPoint2 = null; interSeg1 = null;
    board.defaultCursor = 'default';
    setInstructions(DEFAULT_MSG);
    setActiveButton(null);
    deselectSegment();
}

function paintSegment(seg) {
    let status = seg.__status || 'defaut';
    if (status === 'reel') {
        seg.setAttribute({ strokeColor: '#e67e22', dash: 0, strokeWidth: seg === selectedSegment ? 2.5 : 1.5 });
    } else if (status === 'virtuel') {
        seg.setAttribute({ strokeColor: '#2980b9', dash: 2, strokeWidth: seg === selectedSegment ? 2.5 : 1.5 });
    } else {
        seg.setAttribute({ strokeColor: '#555555', dash: 0, strokeWidth: seg === selectedSegment ? 2.5 : 1.5 });
    }
}

function statusCode(seg) {
    let s = seg.__status || 'defaut';
    return s === 'reel' ? 1 : (s === 'virtuel' ? 2 : 0);
}

function statusFromCode(code) {
    return code === 1 ? 'reel' : (code === 2 ? 'virtuel' : 'defaut');
}

function deselectSegment() {
    let prev = selectedSegment;
    selectedSegment = null;
    if (prev) { paintSegment(prev); }
}

function onSegmentClick(seg) {
    let status = seg.__status || 'defaut';
    seg.__status = (status === 'reel') ? 'virtuel' : 'reel';
    let prev = selectedSegment;
    selectedSegment = seg;
    if (prev && prev !== seg) { paintSegment(prev); }
    paintSegment(seg);
    setInstructions(DEFAULT_MSG);
    syncState();
}

function deleteSelectedSegment() {
    if (!selectedSegment) { setInstructions("Cliquez d'abord sur un tronçon de rayon pour le sélectionner."); return; }
    let seg = selectedSegment;
    selectedSegment = null;
    board.removeObject(seg);
    let idx = raySegments.indexOf(seg);
    if (idx > -1) raySegments.splice(idx, 1);
    for (let lr of logicalRays) {
        let sIdx = lr.segments.indexOf(seg);
        if (sIdx > -1) {
            lr.segments.splice(sIdx, 1);
            if (lr.segments.length === 0) {
                logicalRays = logicalRays.filter(l => l !== lr);
            }
            break;
        }
    }
    setInstructions(DEFAULT_MSG);
    syncState();
}

function activateTool(mode, msg) {
    if (toolMode === mode) { resetTool(); return; }
    resetTool(); toolMode = mode;
    board.defaultCursor = 'crosshair';
    setInstructions(msg);
    setActiveButton(mode);
}

function syncState() {
    handlePoint.trigger(['update']);
    board.update();
}

function snapToPoint(x, y) {
    let threshold = 0.6, closestX = x, closestY = y, minDist = Infinity;
    for (let id in board.objects) {
        let obj = board.objects[id];
        if (obj.elType === 'point' && obj.visProp.visible !== false && obj.visProp.hidden !== true) {
            let dx = obj.X() - x, dy = obj.Y() - y, dist = Math.sqrt(dx * dx + dy * dy);
            if (dist < threshold && dist < minDist) { minDist = dist; closestX = obj.X(); closestY = obj.Y(); }
        }
    }
    return { x: closestX, y: closestY };
}

function getClickedSegment(x, y, threshold, excludeSeg) {
    let closestSeg = null, minDist = Infinity;
    for (let i = 0; i < raySegments.length; i++) {
        let seg = raySegments[i];
        if (seg === excludeSeg) continue;
        if (!seg.point1 || !seg.point2) continue;
        let x1 = seg.point1.X(), y1 = seg.point1.Y();
        let x2 = seg.point2.X(), y2 = seg.point2.Y();
        let dx = x2 - x1, dy = y2 - y1;
        let lengthSq = dx * dx + dy * dy;
        if (lengthSq === 0) continue;
        let t = ((x - x1) * dx + (y - y1) * dy) / lengthSq;
        t = Math.max(0, Math.min(1, t));
        let projX = x1 + t * dx, projY = y1 + t * dy;
        let dist = Math.sqrt((x - projX) ** 2 + (y - projY) ** 2);
        if (dist < threshold && dist < minDist) { minDist = dist; closestSeg = seg; }
    }
    return closestSeg;
}

function getClickedPoint(x, y, threshold) {
    let closest = null, minDist = Infinity;
    for (let el of standaloneElements) {
        if (el.elType === 'point' && el.visProp.visible !== false) {
            let d = Math.hypot(el.X() - x, el.Y() - y);
            if (d < threshold && d < minDist) { minDist = d; closest = el; }
        }
    }
    return closest;
}

function addCustomRayFromEq(m, p, isVert, xVert, xOrigin) {
    let currentLogicalRay = { segments: [], points: [], eq: isVert ? { isVert: true, x: xVert } : { isVert: false, m: m, p: p } };
    if (isVert) {
        let pA = board.create('point', [xVert, Y_MIN], { visible: false, fixed: true, highlight: false, name: '', tabindex: null });
        let pB = board.create('point', [xVert, Y_MAX], { visible: false, fixed: true, highlight: false, name: '', tabindex: null });
        let seg = board.create('segment', [pA, pB], { strokeColor: '#555555', strokeWidth: 1.5, fixed: true, highlight: false, dash: 0, tabindex: null });
        raySegments.push(seg); allDrawnElements.push(pA, pB, seg);
        currentLogicalRay.segments.push(seg); currentLogicalRay.points.push(pA, pB);
    } else {
        let boundsX = [X_MIN, X_MAX];
        if (typeof xOrigin === 'number') boundsX.push(xOrigin);
        for (let lx of LENS_X) {
            let y = m * lx + p;
            if (Math.abs(y) <= lensHeight) {
                boundsX.push(lx);
                let impactPoint = board.create('point', [lx, y], { name: '', size: 4, color: 'black', fixed: true, highlight: false, tabindex: null });
                allDrawnElements.push(impactPoint);
                currentLogicalRay.points.push(impactPoint);
            }
        }
        boundsX.sort((a, b) => a - b);
        for (let i = 0; i < boundsX.length - 1; i++) {
            let xa = boundsX[i], xb = boundsX[i + 1];
            if (xa === xb) continue;
            let ya = m * xa + p, yb = m * xb + p;
            let pA = board.create('point', [xa, ya], { visible: false, fixed: true, highlight: false, name: '', tabindex: null });
            let pB = board.create('point', [xb, yb], { visible: false, fixed: true, highlight: false, name: '', tabindex: null });
            let seg = board.create('segment', [pA, pB], { strokeColor: '#555555', strokeWidth: 1.5, fixed: true, highlight: false, dash: 0, tabindex: null });
            raySegments.push(seg); allDrawnElements.push(pA, pB, seg);
            currentLogicalRay.segments.push(seg); currentLogicalRay.points.push(pA, pB);
        }
    }
    logicalRays.push(currentLogicalRay);
    return currentLogicalRay;
}

function addRayFromPieces(m, p, isVert, xVert, pieces) {
    let currentLogicalRay = { segments: [], points: [], eq: isVert ? { isVert: true, x: xVert } : { isVert: false, m: m, p: p } };
    for (let piece of pieces) {
        let pA, pB;
        if (isVert) {
            pA = board.create('point', [xVert, piece[0]], { visible: false, fixed: true, highlight: false, name: '', tabindex: null });
            pB = board.create('point', [xVert, piece[1]], { visible: false, fixed: true, highlight: false, name: '', tabindex: null });
        } else {
            pA = board.create('point', [piece[0], m * piece[0] + p], { visible: false, fixed: true, highlight: false, name: '', tabindex: null });
            pB = board.create('point', [piece[1], m * piece[1] + p], { visible: false, fixed: true, highlight: false, name: '', tabindex: null });
        }
        let seg = board.create('segment', [pA, pB], { strokeColor: '#555555', strokeWidth: 1.5, fixed: true, highlight: false, dash: 0, tabindex: null });
        seg.__status = statusFromCode(piece[2]);
        paintSegment(seg);
        raySegments.push(seg); allDrawnElements.push(pA, pB, seg);
        currentLogicalRay.segments.push(seg); currentLogicalRay.points.push(pA, pB);
    }
    logicalRays.push(currentLogicalRay);
    return currentLogicalRay;
}

function addCustomRay(x1, y1, x2, y2) {
    let dx = x2 - x1, dy = y2 - y1;
    let m = (Math.abs(dx) < 0.01) ? Infinity : dy / dx;
    if (m === Infinity) {
        addCustomRayFromEq(null, null, true, x1);
    } else {
        let p = y1 - m * x1;
        addCustomRayFromEq(m, p, false, null, x1);
    }
    syncState();
}

function onPointClick(pt) {
    board.removeObject(pt);
    allDrawnElements = allDrawnElements.filter(el => el !== pt);
    standaloneElements = standaloneElements.filter(el => el !== pt);
    syncState();
}

board.on('down', function(evt) {
    if (evt.target && evt.target.closest && evt.target.closest('.JXG_navigation_button')) return;
    if (evt.shiftKey) return;
    let coords = board.getUsrCoordsOfMouse(evt);
    let x = coords[0], y = coords[1];

    if (toolMode !== '') {
        let snapped = snapToPoint(x, y); x = snapped.x; y = snapped.y;

        if (toolMode === 'pt1') { tempPoint = { x: x, y: y }; toolMode = 'pt2'; setInstructions("RAYON (2/2) : cliquez le 2e point."); }
        else if (toolMode === 'pt2') { addCustomRay(tempPoint.x, tempPoint.y, x, y); resetTool(); }
        else if (toolMode === 'par1') { dirPoint1 = { x: x, y: y }; toolMode = 'par2'; setInstructions("PARALLÈLE (2/3) : 2e point de direction."); }
        else if (toolMode === 'par2') { dirPoint2 = { x: x, y: y }; toolMode = 'par3'; setInstructions("PARALLÈLE (3/3) : point de passage."); }
        else if (toolMode === 'par3') {
            let dx = dirPoint2.x - dirPoint1.x, dy = dirPoint2.y - dirPoint1.y;
            addCustomRay(x, y, x + dx, y + dy);
            resetTool();
        }
        else if (toolMode === 'inter1') {
            interSeg1 = getClickedSegment(x, y, 0.4, null);
            if (interSeg1) { toolMode = 'inter2'; setInstructions("INTERSECTION (2/2) : cliquez le 2e rayon."); }
        }
        else if (toolMode === 'inter2') {
            let seg2 = getClickedSegment(x, y, 0.4, interSeg1);
            if (seg2) {
                let x1 = interSeg1.point1.X(), y1 = interSeg1.point1.Y(), x2 = interSeg1.point2.X(), y2 = interSeg1.point2.Y();
                let x3 = seg2.point1.X(), y3 = seg2.point1.Y(), x4 = seg2.point2.X(), y4 = seg2.point2.Y();
                let m1 = (x2 - x1) === 0 ? Infinity : (y2 - y1) / (x2 - x1);
                let m2 = (x4 - x3) === 0 ? Infinity : (y4 - y3) / (x4 - x3);
                if (m1 !== Infinity && m2 !== Infinity && Math.abs(m1 - m2) > 0.001) {
                    let p1 = y1 - m1 * x1, p2 = y3 - m2 * x3;
                    let xi = (p2 - p1) / (m1 - m2), yi = m1 * xi + p1;
                    let in1 = xi >= Math.min(x1, x2) - 0.5 && xi <= Math.max(x1, x2) + 0.5 && yi >= Math.min(y1, y2) - 0.5 && yi <= Math.max(y1, y2) + 0.5;
                    let in2 = xi >= Math.min(x3, x4) - 0.5 && xi <= Math.max(x3, x4) + 0.5 && yi >= Math.min(y3, y4) - 0.5 && yi <= Math.max(y3, y4) + 0.5;
                    if (in1 && in2) {
                        intersectionCounter++;
                        let name = (intersectionCounter === 1) ? "B1" : "I" + (intersectionCounter - 1);
                        let pInt = board.create('point', [xi, yi], { name: name, size: 4, color: 'black', fixed: true, highlight: false, tabindex: null });
                        allDrawnElements.push(pInt); standaloneElements.push(pInt);
                        syncState();
                    } else { setInstructions("Les droites se croisent hors des segments."); }
                } else { setInstructions("Segments parallèles."); }
                resetTool();
            }
        }
        else if (toolMode === 'perp') {
            let srcPt = getClickedPoint(x, y, 0.5);
            if (srcPt) {
                let xB1 = srcPt.X(), yB1v = srcPt.Y();
                let boundsY = [Y_MIN, 0, yB1v, Y_MAX].sort((a, b) => a - b).filter((v, i, arr) => i === 0 || v !== arr[i - 1]);
                let pieces = [];
                for (let i = 0; i < boundsY.length - 1; i++) pieces.push([boundsY[i], boundsY[i + 1], 0]);
                addRayFromPieces(null, null, true, xB1, pieces);
                let pA1 = board.create('point', [xB1, 0], { name: 'A1', size: 4, color: 'black', fixed: true, highlight: false, tabindex: null });
                allDrawnElements.push(pA1); standaloneElements.push(pA1);
                syncState();
                resetTool();
            } else {
                setInstructions("Cliquez sur le point B1 pour tracer le segment A1B1 et placer A1.");
            }
        }
        return;
    }

    let clickedSeg = getClickedSegment(x, y, 0.3, null);
    if (clickedSeg) { onSegmentClick(clickedSeg); return; }
    let clickedPt = getClickedPoint(x, y, 0.5);
    if (clickedPt) { onPointClick(clickedPt); return; }
});

addToolButton('Rayon (2 clics)', 'pt1', "RAYON (1/2) : cliquez le point de départ.");
addToolButton('Rayon parallèle (3 clics)', 'par1', "PARALLÈLE (1/3) : 1er point de direction.");
addToolButton("Intersection (2 clics) → point B1", 'inter1', "INTERSECTION (1/2) : cliquez le 1er rayon.");
addToolButton("Tracer A1 (perpendiculaire)", 'perp', "TRACÉ A1 : cliquez sur le point B1 pour tracer le segment A1B1 et placer A1.");
addToolButton('Effacer le tronçon sélectionné', null, () => deleteSelectedSegment());
addToolButton('Tout effacer', null, () => {
    selectedSegment = null;
    allDrawnElements.forEach(el => board.removeObject(el));
    allDrawnElements = []; raySegments = []; logicalRays = []; standaloneElements = []; intersectionCounter = 0;
    resetTool();
    syncState();
});

document.body.appendChild(toolbarDiv);
document.body.appendChild(instructionsEl);
stack_js.resize_containing_frame("46em", document.documentElement.offsetHeight + "px");

const serialiser = () => {
    let rayList = logicalRays.map(lr => {
        let pieces = lr.segments.map(seg => {
            let ends = lr.eq.isVert ? [seg.point1.Y(), seg.point2.Y()] : [seg.point1.X(), seg.point2.X()];
            return [ends[0], ends[1], statusCode(seg)];
        });
        return lr.eq.isVert ? ["vert", lr.eq.x, pieces] : [lr.eq.m, lr.eq.p, pieces];
    });
    let ptList = standaloneElements
        .filter(el => el.elType === 'point' && el.visProp.visible !== false && el.name)
        .map(el => [el.name, el.X(), el.Y()]);
    return "newton_primaire_construction(" + JSON.stringify(rayList) + "," + JSON.stringify(ptList) + ")";
};

function clearAll() {
    allDrawnElements.forEach(el => board.removeObject(el));
    allDrawnElements = []; raySegments = []; logicalRays = []; standaloneElements = []; intersectionCounter = 0;
}

const deserialiser = (value) => {
    clearAll();
    let newState = JSON.parse(value.replace('newton_primaire_construction(', '[').replace(/\)\s*$/, ']'));
    let rayList = newState[0], ptList = newState[1];
    for (let eq of rayList) {
        if (eq[0] === 'vert') addRayFromPieces(null, null, true, eq[1], eq[2] || [[Y_MIN, Y_MAX, 0]]);
        else addRayFromPieces(eq[0], eq[1], false, null, eq[2] || [[X_MIN, X_MAX, 0]]);
    }
    for (let p of ptList) {
        let pt = board.create('point', [p[1], p[2]], { name: p[0], size: 4, color: 'black', fixed: true, highlight: false, tabindex: null });
        allDrawnElements.push(pt);
        standaloneElements.push(pt);
    }
    board.update();
};

resetTool();
stack_jxg.custom_bind(state, serialiser, deserialiser, [handlePoint]);
board.update();

const inputEl = document.getElementById(state);
function freezeIfReadonly() {
    if (inputEl && (inputEl.hasAttribute('readonly') || inputEl.hasAttribute('disabled'))) {
        board.containerObj.style.pointerEvents = 'none';
        toolbarDiv.querySelectorAll('button').forEach(b => b.disabled = true);
        setInstructions("Construction validée : la figure est figée.");
        return true;
    }
    return false;
}
if (inputEl && !freezeIfReadonly()) {
    new MutationObserver(freezeIfReadonly).observe(inputEl, { attributes: true, attributeFilter: ['readonly', 'disabled'] });
}
[[/jsxgraph]]

<p><strong>Étape 1 — miroir primaire (image intermédiaire A₁B₁) :</strong> le faisceau parallèle incident (angle α,
tan α = 1/4) contient deux rayons remarquables. Le rayon qui passe par le centre C se réfléchit sur lui-même (il
n'est pas dévié) : tracez-le d'un seul coup avec <strong>« Rayon (2 clics) »</strong> en cliquant deux points
alignés avec C et le sommet S. Le rayon qui passe par S se réfléchit symétriquement par rapport à l'axe optique :
tracez son trajet incident (parallèle au premier, passant par S) puis son trajet réfléchi (pente opposée, passant
par S) avec deux usages de <strong>« Rayon (2 clics) »</strong> ou de <strong>« Rayon parallèle (3 clics) »</strong>.
Utilisez <strong>« Intersection (2 clics) → point B1 »</strong> sur les deux rayons réfléchis pour obtenir B₁, puis
<strong>« Tracer A1 (perpendiculaire) »</strong> en cliquant B₁ : vous devez constater que A₁ tombe exactement sur
F₁'.</p>

<p>[[input:ans1]] [[validation:ans1]]</p>

<div class="stack-comment">
<p><strong>Étape 2 — miroir secondaire (image réelle A₂B₂, faisceau replié à 90°) :</strong> le miroir secondaire,
incliné à 45° (face réfléchissante tournée vers le primaire), son centre se trouve sur l'axe optique, avant le foyer
du primaire. Les points A₁ et B₁ (obtenus à l'étape précédente), à gauche du miroir sur la figure ci-dessous, avec
B₁ indiqué en rouge, sont les points <em>visés</em> par le faisceau convergent (le miroir l'intercepte avant qu'il
ne s'y forme réellement). Pour construire l'image réelle A₂B₂ :</p>
<ol>
<li>Avec <strong>« Rayon (2 clics) »</strong>, tracez un rayon incident de B₁ vers un point d'incidence I de votre
choix sur le miroir secondaire.</li>
<li>Avec <strong>« Tracer la normale (1 clic) »</strong>, cliquez sur ce point I pour tracer la normale au miroir en
ce point.</li>
<li>Avec <strong>« Symétrique par rapport à... (3 clics) »</strong>, cliquez B₁, puis les deux extrémités de la
normale : le point symétrique obtenu est aligné avec le rayon réfléchi.</li>
<li>Avec <strong>« Rayon (2 clics) »</strong>, tracez le rayon réfléchi de I vers ce point symétrique : ici, à la
différence d'un miroir plan classique, ce rayon réfléchi est <strong>réel</strong> (le faisceau convergeait vers
B₁, donc après réflexion il converge réellement vers l'image A₂B₂, sans prolongement virtuel nécessaire).</li>
<li>Recommencez avec un second point d'incidence, distinct du premier, pour obtenir par
<strong>« Intersection (2 clics) → point B2 »</strong> le point B₂.</li>
<li>Enfin, avec <strong>« Tracer A2 (perpendiculaire au nouvel axe) »</strong>, cliquez sur B₂ : cela plante A₂ sur
le nouvel axe (replié à 90°, la droite verticale en pointillés gris) et trace le segment A₂B₂.</li>
</ol>
</div>

[[jsxgraph input-ref-ans2="state" width="46em" aspect-ratio="1.3"]]
const board2 = JXG.JSXGraph.initBoard(divid, {
    boundingbox: [-8, 4, 1, -6],
    axis: false,
    keepaspectratio: true,
    showNavigation: true,
    zoom: { enabled: true, wheel: true, needShift: false, factorX: 1.25, factorY: 1.25 },
    pan: { enabled: true, needTwoFingers: false, needShift: true }
});

const X2_MIN = -8, X2_MAX = 1, Y2_MIN = -6, Y2_MAX = 4;
const M1b = { x: -0.5, y: 2.5 }, M2b = { x: -4, y: -1 };
const B1pt = { x: -6, y: 1.5 };
const xA1pt = -6;
const axisX2 = -3;

let toolMode2 = '';
let tempPoint2 = null, interSeg1b = null;
let symP2 = null, symQ1b = null;
let selectedSegment2 = null;

let allDrawnElements2 = [];
let raySegments2 = [];
let logicalRays2 = [];
let standaloneElements2 = [];
let intersectionCounter2 = 0;
let symCounter2 = 0;

board2.create('segment', [[M1b.x, M1b.y], [M2b.x, M2b.y]], { strokeColor: 'black', strokeWidth: 4, fixed: true, highlight: false, tabindex: null });
{
    const nBackX = -0.6, nBackY = 0.6;
    for (let frac of [0.1, 0.3, 0.5, 0.7, 0.9]) {
        let hx = M1b.x + frac * (M2b.x - M1b.x), hy = M1b.y + frac * (M2b.y - M1b.y);
        board2.create('segment', [[hx, hy], [hx + 0.5 * nBackX, hy + 0.5 * nBackY]], { strokeColor: 'black', strokeWidth: 2, fixed: true, highlight: false, tabindex: null });
    }
}

board2.create('point', [B1pt.x, B1pt.y], { name: 'B1', size: 4, color: 'red', fixed: true, highlight: false, tabindex: null });
board2.create('point', [xA1pt, 0], { name: 'A1', size: 3, fixed: true, color: 'black', highlight: false, tabindex: null });
board2.create('segment', [[xA1pt, 0], [B1pt.x, B1pt.y]], { strokeColor: '#888888', strokeWidth: 1, dash: 2, fixed: true, highlight: false, tabindex: null });
board2.create('line', [[X2_MIN, 0], [X2_MAX, 0]], { strokeColor: '#aaaaaa', strokeWidth: 1, dash: 2, fixed: true, highlight: false, tabindex: null });
board2.create('segment', [[axisX2, -5.5], [axisX2, 2.5]], { strokeColor: '#888888', strokeWidth: 1, dash: 2, fixed: true, highlight: false, tabindex: null });
board2.create('text', [axisX2 + 0.15, -5.2, "nouvel axe"], { fixed: true, fontSize: 10, color: '#888888', highlight: false, tabindex: null });

const handlePoint2 = board2.create('point', [-3.5, -2.7], { visible: false, fixed: true, name: '', tabindex: null });

const DEFAULT_MSG2 = "Cliquez sur un tronçon de droite pour basculer son statut (réel/virtuel) ; il reste sélectionné pour le bouton Effacer. Cliquez sur un point pour le supprimer. Molette : zoom. Maj + glisser : déplacer la vue.";

const instructionsEl2 = document.createElement('p');
instructionsEl2.style.cssText = 'margin:.6em 0 0;font-size:.85em;color:#333;';
instructionsEl2.textContent = DEFAULT_MSG2;

function setInstructions2(msg) {
    instructionsEl2.textContent = msg;
}

const toolbarDiv2 = document.createElement('div');
toolbarDiv2.style.cssText = 'display:flex;flex-wrap:wrap;gap:.4em;margin-top:.6em;';
const toolButtons2 = {};

function addToolButton2(label, mode, msgOrHandler) {
    const btn = document.createElement('button');
    btn.type = 'button';
    btn.textContent = label;
    btn.style.cssText = 'padding:.35em .7em;font-size:.85em;cursor:pointer;';
    if (mode === null) {
        btn.addEventListener('click', msgOrHandler);
    } else {
        btn.addEventListener('click', () => activateTool2(mode, msgOrHandler));
        toolButtons2[mode] = btn;
    }
    toolbarDiv2.appendChild(btn);
    return btn;
}

function setActiveButton2(mode) {
    for (let m in toolButtons2) {
        toolButtons2[m].style.background = (m === mode) ? '#dbeafe' : '';
        toolButtons2[m].style.fontWeight = (m === mode) ? 'bold' : 'normal';
    }
}

function resetTool2() {
    toolMode2 = '';
    tempPoint2 = null; interSeg1b = null; symP2 = null; symQ1b = null;
    board2.defaultCursor = 'default';
    setInstructions2(DEFAULT_MSG2);
    setActiveButton2(null);
    deselectSegment2();
}

function paintSegment2(seg) {
    if (seg.__isNormal) {
        seg.setAttribute({ strokeColor: '#27ae60', dash: 2, strokeWidth: seg === selectedSegment2 ? 2.5 : 1.5 });
        return;
    }
    let status = seg.__status || 'defaut';
    if (status === 'reel') {
        seg.setAttribute({ strokeColor: '#e67e22', dash: 0, strokeWidth: seg === selectedSegment2 ? 2.5 : 1.5 });
    } else if (status === 'virtuel') {
        seg.setAttribute({ strokeColor: '#2980b9', dash: 2, strokeWidth: seg === selectedSegment2 ? 2.5 : 1.5 });
    } else {
        seg.setAttribute({ strokeColor: '#555555', dash: 0, strokeWidth: seg === selectedSegment2 ? 2.5 : 1.5 });
    }
}

function statusCode2(seg) {
    let s = seg.__status || 'defaut';
    return s === 'reel' ? 1 : (s === 'virtuel' ? 2 : 0);
}

function statusFromCode2(code) {
    return code === 1 ? 'reel' : (code === 2 ? 'virtuel' : 'defaut');
}

function deselectSegment2() {
    let prev = selectedSegment2;
    selectedSegment2 = null;
    if (prev) { paintSegment2(prev); }
}

function onSegmentClick2(seg) {
    if (!seg.__isNormal) {
        let status = seg.__status || 'defaut';
        seg.__status = (status === 'reel') ? 'virtuel' : 'reel';
    }
    let prev = selectedSegment2;
    selectedSegment2 = seg;
    if (prev && prev !== seg) { paintSegment2(prev); }
    paintSegment2(seg);
    setInstructions2(DEFAULT_MSG2);
    syncState2();
}

function deleteSelectedSegment2() {
    if (!selectedSegment2) { setInstructions2("Cliquez d'abord sur un tronçon de rayon pour le sélectionner."); return; }
    let seg = selectedSegment2;
    selectedSegment2 = null;
    if (seg.__isNormal) {
        let p1 = seg.point1, p2 = seg.point2;
        board2.removeObject(seg);
        board2.removeObject(p1); board2.removeObject(p2);
        allDrawnElements2 = allDrawnElements2.filter(el => el !== p1 && el !== p2);
    } else {
        board2.removeObject(seg);
    }
    let idx = raySegments2.indexOf(seg);
    if (idx > -1) raySegments2.splice(idx, 1);
    for (let lr of logicalRays2) {
        let sIdx = lr.segments.indexOf(seg);
        if (sIdx > -1) {
            lr.segments.splice(sIdx, 1);
            if (lr.segments.length === 0) {
                logicalRays2 = logicalRays2.filter(l => l !== lr);
            }
            break;
        }
    }
    setInstructions2(DEFAULT_MSG2);
    syncState2();
}

function activateTool2(mode, msg) {
    if (toolMode2 === mode) { resetTool2(); return; }
    resetTool2(); toolMode2 = mode;
    board2.defaultCursor = 'crosshair';
    setInstructions2(msg);
    setActiveButton2(mode);
}

function syncState2() {
    handlePoint2.trigger(['update']);
    board2.update();
}

function snapToPoint2(x, y) {
    let threshold = 0.6, closestX = x, closestY = y, minDist = Infinity;
    for (let id in board2.objects) {
        let obj = board2.objects[id];
        if (obj.elType === 'point' && obj.visProp.visible !== false && obj.visProp.hidden !== true) {
            let dx = obj.X() - x, dy = obj.Y() - y, dist = Math.sqrt(dx * dx + dy * dy);
            if (dist < threshold && dist < minDist) { minDist = dist; closestX = obj.X(); closestY = obj.Y(); }
        }
    }
    return { x: closestX, y: closestY };
}

function getClickedSegment2(x, y, threshold, excludeSeg) {
    let closestSeg = null, minDist = Infinity;
    for (let i = 0; i < raySegments2.length; i++) {
        let seg = raySegments2[i];
        if (seg === excludeSeg) continue;
        if (!seg.point1 || !seg.point2) continue;
        let x1 = seg.point1.X(), y1 = seg.point1.Y();
        let x2 = seg.point2.X(), y2 = seg.point2.Y();
        let dx = x2 - x1, dy = y2 - y1;
        let lengthSq = dx * dx + dy * dy;
        if (lengthSq === 0) continue;
        let t = ((x - x1) * dx + (y - y1) * dy) / lengthSq;
        t = Math.max(0, Math.min(1, t));
        let projX = x1 + t * dx, projY = y1 + t * dy;
        let dist = Math.sqrt((x - projX) ** 2 + (y - projY) ** 2);
        if (dist < threshold && dist < minDist) { minDist = dist; closestSeg = seg; }
    }
    return closestSeg;
}

function getClickedPoint2(x, y, threshold) {
    let closest = null, minDist = Infinity;
    for (let el of standaloneElements2) {
        if (el.elType === 'point' && el.visProp.visible !== false) {
            let d = Math.hypot(el.X() - x, el.Y() - y);
            if (d < threshold && d < minDist) { minDist = d; closest = el; }
        }
    }
    return closest;
}

const mM2_js = (M2b.y - M1b.y) / (M2b.x - M1b.x);
const pM2_js = M1b.y - mM2_js * M1b.x;
const xMmin2_js = Math.min(M1b.x, M2b.x), xMmax2_js = Math.max(M1b.x, M2b.x);

function addCustomRayFromEq2(m, p, isVert, xVert, xOrigin) {
    let currentLogicalRay = { segments: [], points: [], eq: isVert ? { isVert: true, x: xVert } : { isVert: false, m: m, p: p } };
    if (isVert) {
        let boundsY = [Y2_MIN, Y2_MAX];
        let iy = mM2_js * xVert + pM2_js;
        let nearExistingY = boundsY.some(by => Math.abs(by - iy) < 1e-6);
        if (!nearExistingY && xVert >= xMmin2_js - 0.3 && xVert <= xMmax2_js + 0.3) {
            boundsY.push(iy);
            let impactPoint = board2.create('point', [xVert, iy], { name: '', size: 4, color: 'black', fixed: true, highlight: false, tabindex: null });
            allDrawnElements2.push(impactPoint);
            currentLogicalRay.points.push(impactPoint);
        }
        boundsY.sort((a, b) => a - b);
        for (let i = 0; i < boundsY.length - 1; i++) {
            let ya = boundsY[i], yb = boundsY[i + 1];
            if (Math.abs(ya - yb) < 1e-6) continue;
            let pA = board2.create('point', [xVert, ya], { visible: false, fixed: true, highlight: false, name: '', tabindex: null });
            let pB = board2.create('point', [xVert, yb], { visible: false, fixed: true, highlight: false, name: '', tabindex: null });
            let seg = board2.create('segment', [pA, pB], { strokeColor: '#555555', strokeWidth: 1.5, fixed: true, highlight: false, dash: 0, tabindex: null });
            raySegments2.push(seg); allDrawnElements2.push(pA, pB, seg);
            currentLogicalRay.segments.push(seg); currentLogicalRay.points.push(pA, pB);
        }
    } else {
        let boundsX = [X2_MIN, X2_MAX];
        if (typeof xOrigin === 'number') boundsX.push(xOrigin);
        if (Math.abs(m - mM2_js) > 0.0001) {
            let ix = (pM2_js - p) / (m - mM2_js);
            let nearExistingX = boundsX.some(bx => Math.abs(bx - ix) < 1e-6);
            if (!nearExistingX && ix >= xMmin2_js - 0.3 && ix <= xMmax2_js + 0.3) {
                boundsX.push(ix);
                let impactPoint = board2.create('point', [ix, m * ix + p], { name: '', size: 4, color: 'black', fixed: true, highlight: false, tabindex: null });
                allDrawnElements2.push(impactPoint);
                currentLogicalRay.points.push(impactPoint);
            }
        }
        boundsX.sort((a, b) => a - b);
        for (let i = 0; i < boundsX.length - 1; i++) {
            let xa = boundsX[i], xb = boundsX[i + 1];
            if (Math.abs(xa - xb) < 1e-6) continue;
            let ya = m * xa + p, yb = m * xb + p;
            let pA = board2.create('point', [xa, ya], { visible: false, fixed: true, highlight: false, name: '', tabindex: null });
            let pB = board2.create('point', [xb, yb], { visible: false, fixed: true, highlight: false, name: '', tabindex: null });
            let seg = board2.create('segment', [pA, pB], { strokeColor: '#555555', strokeWidth: 1.5, fixed: true, highlight: false, dash: 0, tabindex: null });
            raySegments2.push(seg); allDrawnElements2.push(pA, pB, seg);
            currentLogicalRay.segments.push(seg); currentLogicalRay.points.push(pA, pB);
        }
    }
    logicalRays2.push(currentLogicalRay);
    return currentLogicalRay;
}

function addRayFromPieces2(m, p, isVert, xVert, pieces) {
    let currentLogicalRay = { segments: [], points: [], eq: isVert ? { isVert: true, x: xVert } : { isVert: false, m: m, p: p } };
    for (let piece of pieces) {
        let pA, pB;
        if (isVert) {
            pA = board2.create('point', [xVert, piece[0]], { visible: false, fixed: true, highlight: false, name: '', tabindex: null });
            pB = board2.create('point', [xVert, piece[1]], { visible: false, fixed: true, highlight: false, name: '', tabindex: null });
        } else {
            pA = board2.create('point', [piece[0], m * piece[0] + p], { visible: false, fixed: true, highlight: false, name: '', tabindex: null });
            pB = board2.create('point', [piece[1], m * piece[1] + p], { visible: false, fixed: true, highlight: false, name: '', tabindex: null });
        }
        let seg = board2.create('segment', [pA, pB], { strokeColor: '#555555', strokeWidth: 1.5, fixed: true, highlight: false, dash: 0, tabindex: null });
        seg.__status = statusFromCode2(piece[2]);
        paintSegment2(seg);
        raySegments2.push(seg); allDrawnElements2.push(pA, pB, seg);
        currentLogicalRay.segments.push(seg); currentLogicalRay.points.push(pA, pB);
    }
    logicalRays2.push(currentLogicalRay);
    return currentLogicalRay;
}

function addCustomRay2(x1, y1, x2, y2, isNormal) {
    let dx = x2 - x1, dy = y2 - y1;
    let m = (Math.abs(dx) < 0.01) ? Infinity : dy / dx;
    let lr;
    if (m === Infinity) {
        lr = addCustomRayFromEq2(null, null, true, x1);
    } else {
        let p = y1 - m * x1;
        lr = addCustomRayFromEq2(m, p, false, null, x1);
    }
    if (isNormal) {
        for (let seg of lr.segments) { seg.__isNormal = true; paintSegment2(seg); }
    }
    syncState2();
    return lr;
}

function onPointClick2(pt) {
    board2.removeObject(pt);
    allDrawnElements2 = allDrawnElements2.filter(el => el !== pt);
    standaloneElements2 = standaloneElements2.filter(el => el !== pt);
    syncState2();
}

board2.on('down', function(evt) {
    if (evt.target && evt.target.closest && evt.target.closest('.JXG_navigation_button')) return;
    if (evt.shiftKey) return;
    let coords = board2.getUsrCoordsOfMouse(evt);
    let x = coords[0], y = coords[1];

    if (toolMode2 !== '') {
        let snapped = snapToPoint2(x, y); x = snapped.x; y = snapped.y;

        if (toolMode2 === 'pt1') { tempPoint2 = { x: x, y: y }; toolMode2 = 'pt2'; setInstructions2("RAYON (2/2) : cliquez le 2e point."); }
        else if (toolMode2 === 'pt2') { addCustomRay2(tempPoint2.x, tempPoint2.y, x, y); resetTool2(); }
        else if (toolMode2 === 'normal1') {
            let ux = M2b.x - M1b.x, uy = M2b.y - M1b.y;
            let ulen2 = ux * ux + uy * uy;
            let t = ((x - M1b.x) * ux + (y - M1b.y) * uy) / ulen2;
            let fx = M1b.x + t * ux, fy = M1b.y + t * uy;
            let ulen = Math.sqrt(ulen2);
            let nx = -uy / ulen, ny = ux / ulen;
            let n1x = fx + 5 * nx, n1y = fy + 5 * ny, n2x = fx - 5 * nx, n2y = fy - 5 * ny;
            let nP1 = board2.create('point', [n1x, n1y], { name: '', size: 3, color: '#27ae60', fixed: true, highlight: false, tabindex: null });
            let nP2 = board2.create('point', [n2x, n2y], { name: '', size: 3, color: '#27ae60', fixed: true, highlight: false, tabindex: null });
            let nSeg = board2.create('segment', [nP1, nP2], { strokeColor: '#27ae60', dash: 2, strokeWidth: 1.5, fixed: true, highlight: false, tabindex: null });
            nSeg.__isNormal = true;
            raySegments2.push(nSeg); allDrawnElements2.push(nP1, nP2, nSeg);
            syncState2();
            resetTool2();
        }
        else if (toolMode2 === 'sym1') { symP2 = { x: x, y: y }; toolMode2 = 'sym2'; setInstructions2("SYMÉTRIE (2/3) : cliquez le 1er point de l'axe (ex : une extrémité de la normale)."); }
        else if (toolMode2 === 'sym2') { symQ1b = { x: x, y: y }; toolMode2 = 'sym3'; setInstructions2("SYMÉTRIE (3/3) : cliquez le 2e point de l'axe."); }
        else if (toolMode2 === 'sym3') {
            let dx = x - symQ1b.x, dy = y - symQ1b.y;
            let dlen2 = dx * dx + dy * dy;
            if (dlen2 < 0.0001) { setInstructions2("Les deux points de l'axe sont confondus : recommencez."); resetTool2(); return; }
            let t = ((symP2.x - symQ1b.x) * dx + (symP2.y - symQ1b.y) * dy) / dlen2;
            let fx = symQ1b.x + t * dx, fy = symQ1b.y + t * dy;
            let rx = 2 * fx - symP2.x, ry = 2 * fy - symP2.y;
            symCounter2++;
            let pSym = board2.create('point', [rx, ry], { name: 'S' + symCounter2, size: 4, color: 'black', fixed: true, highlight: false, tabindex: null });
            allDrawnElements2.push(pSym); standaloneElements2.push(pSym);
            syncState2();
            resetTool2();
        }
        else if (toolMode2 === 'inter1') {
            interSeg1b = getClickedSegment2(x, y, 0.4, null);
            if (interSeg1b) { toolMode2 = 'inter2'; setInstructions2("INTERSECTION (2/2) : cliquez le 2e rayon réfléchi."); }
        }
        else if (toolMode2 === 'inter2') {
            let seg2 = getClickedSegment2(x, y, 0.4, interSeg1b);
            if (seg2) {
                let x1 = interSeg1b.point1.X(), y1 = interSeg1b.point1.Y(), x2 = interSeg1b.point2.X(), y2 = interSeg1b.point2.Y();
                let x3 = seg2.point1.X(), y3 = seg2.point1.Y(), x4 = seg2.point2.X(), y4 = seg2.point2.Y();
                let m1 = (x2 - x1) === 0 ? Infinity : (y2 - y1) / (x2 - x1);
                let m2b = (x4 - x3) === 0 ? Infinity : (y4 - y3) / (x4 - x3);
                if (m1 !== Infinity && m2b !== Infinity && Math.abs(m1 - m2b) > 0.001) {
                    let p1 = y1 - m1 * x1, p2 = y3 - m2b * x3;
                    let xi = (p2 - p1) / (m1 - m2b), yi = m1 * xi + p1;
                    let in1 = xi >= Math.min(x1, x2) - 0.5 && xi <= Math.max(x1, x2) + 0.5 && yi >= Math.min(y1, y2) - 0.5 && yi <= Math.max(y1, y2) + 0.5;
                    let in2 = xi >= Math.min(x3, x4) - 0.5 && xi <= Math.max(x3, x4) + 0.5 && yi >= Math.min(y3, y4) - 0.5 && yi <= Math.max(y3, y4) + 0.5;
                    if (in1 && in2) {
                        intersectionCounter2++;
                        let name = (intersectionCounter2 === 1) ? "B2" : "I" + (intersectionCounter2 - 1);
                        let pInt = board2.create('point', [xi, yi], { name: name, size: 4, color: 'black', fixed: true, highlight: false, tabindex: null });
                        allDrawnElements2.push(pInt); standaloneElements2.push(pInt);
                        syncState2();
                    } else { setInstructions2("Les droites se croisent hors des segments."); }
                } else { setInstructions2("Segments parallèles."); }
                resetTool2();
            }
        }
        else if (toolMode2 === 'perp2') {
            let srcPt = getClickedPoint2(x, y, 0.5);
            if (srcPt) {
                let xB2 = srcPt.X(), yB2v = srcPt.Y();
                let boundsX = [X2_MIN, axisX2, xB2, X2_MAX].sort((a, b) => a - b).filter((v, i, arr) => i === 0 || v !== arr[i - 1]);
                let pieces = [];
                for (let i = 0; i < boundsX.length - 1; i++) pieces.push([boundsX[i], boundsX[i + 1], 0]);
                addRayFromPieces2(0, yB2v, false, null, pieces);
                let pA2 = board2.create('point', [axisX2, yB2v], { name: 'A2', size: 4, color: 'black', fixed: true, highlight: false, tabindex: null });
                allDrawnElements2.push(pA2); standaloneElements2.push(pA2);
                syncState2();
                resetTool2();
            } else {
                setInstructions2("Cliquez sur le point B2 pour tracer le segment A2B2 et placer A2.");
            }
        }
        return;
    }

    let clickedSeg = getClickedSegment2(x, y, 0.3, null);
    if (clickedSeg) { onSegmentClick2(clickedSeg); return; }
    let clickedPt = getClickedPoint2(x, y, 0.5);
    if (clickedPt) { onPointClick2(clickedPt); return; }
});

addToolButton2('Rayon (2 clics)', 'pt1', "RAYON (1/2) : cliquez le point de départ (ex : B1, ou un point d'incidence I).");
addToolButton2('Tracer la normale (1 clic)', 'normal1', "NORMALE : cliquez un point du miroir pour tracer la normale en ce point.");
addToolButton2("Symétrique par rapport à... (3 clics)", 'sym1', "SYMÉTRIE (1/3) : cliquez le point à réfléchir (ex : B1).");
addToolButton2("Intersection (2 clics) → point B2", 'inter1', "INTERSECTION (1/2) : cliquez le 1er rayon réfléchi.");
addToolButton2("Tracer A2 (perpendiculaire au nouvel axe)", 'perp2', "TRACÉ A2 : cliquez sur le point B2 pour tracer le segment A2B2 et placer A2.");
addToolButton2('Effacer le tronçon sélectionné', null, () => deleteSelectedSegment2());
addToolButton2('Tout effacer', null, () => {
    selectedSegment2 = null;
    allDrawnElements2.forEach(el => board2.removeObject(el));
    allDrawnElements2 = []; raySegments2 = []; logicalRays2 = []; standaloneElements2 = []; intersectionCounter2 = 0; symCounter2 = 0;
    resetTool2();
    syncState2();
});

document.body.appendChild(toolbarDiv2);
document.body.appendChild(instructionsEl2);
stack_js.resize_containing_frame("46em", document.documentElement.offsetHeight + "px");

const serialiser2 = () => {
    let rayList = logicalRays2.map(lr => {
        let pieces = lr.segments.map(seg => {
            let ends = lr.eq.isVert ? [seg.point1.Y(), seg.point2.Y()] : [seg.point1.X(), seg.point2.X()];
            return [ends[0], ends[1], statusCode2(seg)];
        });
        return lr.eq.isVert ? ["vert", lr.eq.x, pieces] : [lr.eq.m, lr.eq.p, pieces];
    });
    let ptList = standaloneElements2
        .filter(el => el.elType === 'point' && el.visProp.visible !== false && el.name)
        .map(el => [el.name, el.X(), el.Y()]);
    return "newton_secondaire_construction(" + JSON.stringify(rayList) + "," + JSON.stringify(ptList) + ")";
};

function clearAll2() {
    allDrawnElements2.forEach(el => board2.removeObject(el));
    allDrawnElements2 = []; raySegments2 = []; logicalRays2 = []; standaloneElements2 = []; intersectionCounter2 = 0; symCounter2 = 0;
}

const deserialiser2 = (value) => {
    clearAll2();
    let newState = JSON.parse(value.replace('newton_secondaire_construction(', '[').replace(/\)\s*$/, ']'));
    let rayList = newState[0], ptList = newState[1];
    for (let eq of rayList) {
        if (eq[0] === 'vert') addRayFromPieces2(null, null, true, eq[1], eq[2] || [[Y2_MIN, Y2_MAX, 0]]);
        else addRayFromPieces2(eq[0], eq[1], false, null, eq[2] || [[X2_MIN, X2_MAX, 0]]);
    }
    for (let p of ptList) {
        let pt = board2.create('point', [p[1], p[2]], { name: p[0], size: 4, color: 'black', fixed: true, highlight: false, tabindex: null });
        allDrawnElements2.push(pt);
        standaloneElements2.push(pt);
    }
    board2.update();
};

resetTool2();
stack_jxg.custom_bind(state, serialiser2, deserialiser2, [handlePoint2]);
board2.update();

const inputEl2 = document.getElementById(state);
function freezeIfReadonly2() {
    if (inputEl2 && (inputEl2.hasAttribute('readonly') || inputEl2.hasAttribute('disabled'))) {
        board2.containerObj.style.pointerEvents = 'none';
        toolbarDiv2.querySelectorAll('button').forEach(b => b.disabled = true);
        setInstructions2("Construction validée : la figure est figée.");
        return true;
    }
    return false;
}
if (inputEl2 && !freezeIfReadonly2()) {
    new MutationObserver(freezeIfReadonly2).observe(inputEl2, { attributes: true, attributeFilter: ['readonly', 'disabled'] });
}
[[/jsxgraph]]

<p>[[input:ans2]] [[validation:ans2]]</p>

<div class="stack-comment">
<p><strong>Étape 3 — oculaire (image finale à l'infini) :</strong> A₂B₂ se trouve exactement dans le plan focal
objet de l'oculaire (distance focale f₂ = 3, déjà marqué F₂). C'est un objet réel fini pour l'oculaire, présenté ici
sur un axe horizontal « déplié » pour plus de clarté. Tracez le rayon issu de B₂ passant par O₂ (non dévié), puis le
rayon issu de B₂ parallèle à l'axe (bouton <strong>« Rayon parallèle (3 clics) »</strong>) : il émerge en passant
par F₂'. Les deux rayons émergents doivent ressortir parallèles entre eux : l'image finale est rejetée à l'infini,
le télescope est afocal.</p>
</div>

[[jsxgraph input-ref-ans3="state" width="46em" aspect-ratio="1.9"]]
const board3 = JXG.JSXGraph.initBoard(divid, {
    boundingbox: [-4, 6, 4, -6],
    axis: false,
    keepaspectratio: true,
    showNavigation: true,
    zoom: { enabled: true, wheel: true, needShift: false, factorX: 1.25, factorY: 1.25 },
    pan: { enabled: true, needTwoFingers: false, needShift: true }
});

const f2b = 3, lensHeight3 = 3;
const xO3 = 0, xA2b = -f2b, yB2b = 3 / 2;
const xF3p = xO3 + f2b;
const X3_MIN = -6, X3_MAX = 6, Y3_MIN = -4, Y3_MAX = 4;
const LENS_X3 = [xO3];

// Repère visuel : axe optique vertical, orienté vers le bas.
// Coordonnée logique "s" (axiale, le long de l'axe optique) et "t" (transverse) ;
// position visuelle à l'écran = (t, -s).
function toVisual(s, t) { return [t, -s]; }

let toolMode3 = '';
let tempPoint3 = null, dirPoint1c = null, dirPoint2c = null, interSeg1c = null;
let selectedSegment3 = null;

let allDrawnElements3 = [];
let raySegments3 = [];
let logicalRays3 = [];
let standaloneElements3 = [];
let intersectionCounter3 = 0;

board3.create('line', [[0, -X3_MIN], [0, -X3_MAX]], { strokeColor: 'black', strokeWidth: 1, fixed: true, highlight: false, tabindex: null });
board3.create('arrow', [[0, -(X3_MAX - 0.5)], [0, -X3_MAX]], { fixed: true, highlight: false, tabindex: null });

function drawLens3(sPos) {
    let by = -sPos;
    board3.create('segment', [[-lensHeight3, by], [lensHeight3, by]], { strokeColor: 'black', strokeWidth: 3, fixed: true, highlight: false, tabindex: null });
    board3.create('segment', [[-lensHeight3, by], [-lensHeight3 + 0.4, by + 0.4]], { strokeColor: 'black', strokeWidth: 3, fixed: true, highlight: false, tabindex: null });
    board3.create('segment', [[-lensHeight3, by], [-lensHeight3 + 0.4, by - 0.4]], { strokeColor: 'black', strokeWidth: 3, fixed: true, highlight: false, tabindex: null });
    board3.create('segment', [[lensHeight3, by], [lensHeight3 - 0.4, by + 0.4]], { strokeColor: 'black', strokeWidth: 3, fixed: true, highlight: false, tabindex: null });
    board3.create('segment', [[lensHeight3, by], [lensHeight3 - 0.4, by - 0.4]], { strokeColor: 'black', strokeWidth: 3, fixed: true, highlight: false, tabindex: null });
}
drawLens3(xO3);

board3.create('point', toVisual(xO3, 0), { name: 'O₂', size: 3, fixed: true, color: 'black', highlight: false, tabindex: null });
board3.create('point', toVisual(xO3 - f2b, 0), { name: 'F₂', size: 3, fixed: true, color: 'green', highlight: false, tabindex: null });
board3.create('point', toVisual(xF3p, 0), { name: "F₂'", size: 3, fixed: true, color: 'green', highlight: false, tabindex: null });

board3.create('point', toVisual(xA2b, 0), { name: 'A2', size: 2, fixed: true, color: 'black', highlight: false, tabindex: null });
board3.create('point', toVisual(xA2b, yB2b), { name: 'B2', size: 3, fixed: true, color: '#c0392b', highlight: false, tabindex: null });
board3.create('arrow', [toVisual(xA2b, 0), toVisual(xA2b, yB2b)], { strokeColor: '#c0392b', strokeWidth: 2.5, fixed: true, highlight: false, tabindex: null });
board3.create('text', [0.3, -xA2b + 0.6, "A2B2 (objet réel, issu du secondaire)"], { fixed: true, fontSize: 11, color: '#c0392b', highlight: false, tabindex: null });

const handlePoint3 = board3.create('point', [3.5, -5.5], { visible: false, fixed: true, name: '', tabindex: null });

const DEFAULT_MSG3 = "Cliquez sur un tronçon de droite pour basculer son statut (réel/virtuel) ; il reste sélectionné pour le bouton Effacer. Cliquez sur un point pour le supprimer. Molette : zoom. Maj + glisser : déplacer la vue.";

const instructionsEl3 = document.createElement('p');
instructionsEl3.style.cssText = 'margin:.6em 0 0;font-size:.85em;color:#333;';
instructionsEl3.textContent = DEFAULT_MSG3;

function setInstructions3(msg) {
    instructionsEl3.textContent = msg;
}

const toolbarDiv3 = document.createElement('div');
toolbarDiv3.style.cssText = 'display:flex;flex-wrap:wrap;gap:.4em;margin-top:.6em;';
const toolButtons3 = {};

function addToolButton3(label, mode, msgOrHandler) {
    const btn = document.createElement('button');
    btn.type = 'button';
    btn.textContent = label;
    btn.style.cssText = 'padding:.35em .7em;font-size:.85em;cursor:pointer;';
    if (mode === null) {
        btn.addEventListener('click', msgOrHandler);
    } else {
        btn.addEventListener('click', () => activateTool3(mode, msgOrHandler));
        toolButtons3[mode] = btn;
    }
    toolbarDiv3.appendChild(btn);
    return btn;
}

function setActiveButton3(mode) {
    for (let m in toolButtons3) {
        toolButtons3[m].style.background = (m === mode) ? '#dbeafe' : '';
        toolButtons3[m].style.fontWeight = (m === mode) ? 'bold' : 'normal';
    }
}

function resetTool3() {
    toolMode3 = '';
    tempPoint3 = null; dirPoint1c = null; dirPoint2c = null; interSeg1c = null;
    board3.defaultCursor = 'default';
    setInstructions3(DEFAULT_MSG3);
    setActiveButton3(null);
    deselectSegment3();
}

function paintSegment3(seg) {
    let status = seg.__status || 'defaut';
    if (status === 'reel') {
        seg.setAttribute({ strokeColor: '#e67e22', dash: 0, strokeWidth: seg === selectedSegment3 ? 2.5 : 1.5 });
    } else if (status === 'virtuel') {
        seg.setAttribute({ strokeColor: '#2980b9', dash: 2, strokeWidth: seg === selectedSegment3 ? 2.5 : 1.5 });
    } else {
        seg.setAttribute({ strokeColor: '#555555', dash: 0, strokeWidth: seg === selectedSegment3 ? 2.5 : 1.5 });
    }
}

function statusCode3(seg) {
    let s = seg.__status || 'defaut';
    return s === 'reel' ? 1 : (s === 'virtuel' ? 2 : 0);
}

function statusFromCode3(code) {
    return code === 1 ? 'reel' : (code === 2 ? 'virtuel' : 'defaut');
}

function deselectSegment3() {
    let prev = selectedSegment3;
    selectedSegment3 = null;
    if (prev) { paintSegment3(prev); }
}

function onSegmentClick3(seg) {
    let status = seg.__status || 'defaut';
    seg.__status = (status === 'reel') ? 'virtuel' : 'reel';
    let prev = selectedSegment3;
    selectedSegment3 = seg;
    if (prev && prev !== seg) { paintSegment3(prev); }
    paintSegment3(seg);
    setInstructions3(DEFAULT_MSG3);
    syncState3();
}

function deleteSelectedSegment3() {
    if (!selectedSegment3) { setInstructions3("Cliquez d'abord sur un tronçon de rayon pour le sélectionner."); return; }
    let seg = selectedSegment3;
    selectedSegment3 = null;
    board3.removeObject(seg);
    let idx = raySegments3.indexOf(seg);
    if (idx > -1) raySegments3.splice(idx, 1);
    for (let lr of logicalRays3) {
        let sIdx = lr.segments.indexOf(seg);
        if (sIdx > -1) {
            lr.segments.splice(sIdx, 1);
            if (lr.segments.length === 0) {
                logicalRays3 = logicalRays3.filter(l => l !== lr);
            }
            break;
        }
    }
    setInstructions3(DEFAULT_MSG3);
    syncState3();
}

function activateTool3(mode, msg) {
    if (toolMode3 === mode) { resetTool3(); return; }
    resetTool3(); toolMode3 = mode;
    board3.defaultCursor = 'crosshair';
    setInstructions3(msg);
    setActiveButton3(mode);
}

function syncState3() {
    handlePoint3.trigger(['update']);
    board3.update();
}

function snapToPoint3(x, y) {
    let threshold = 0.6, closestX = x, closestY = y, minDist = Infinity;
    for (let id in board3.objects) {
        let obj = board3.objects[id];
        if (obj.elType === 'point' && obj.visProp.visible !== false && obj.visProp.hidden !== true) {
            let dx = obj.X() - x, dy = obj.Y() - y, dist = Math.sqrt(dx * dx + dy * dy);
            if (dist < threshold && dist < minDist) { minDist = dist; closestX = obj.X(); closestY = obj.Y(); }
        }
    }
    return { x: closestX, y: closestY };
}

function getClickedSegment3(x, y, threshold, excludeSeg) {
    let closestSeg = null, minDist = Infinity;
    for (let i = 0; i < raySegments3.length; i++) {
        let seg = raySegments3[i];
        if (seg === excludeSeg) continue;
        if (!seg.point1 || !seg.point2) continue;
        let x1 = seg.point1.X(), y1 = seg.point1.Y();
        let x2 = seg.point2.X(), y2 = seg.point2.Y();
        let dx = x2 - x1, dy = y2 - y1;
        let lengthSq = dx * dx + dy * dy;
        if (lengthSq === 0) continue;
        let t = ((x - x1) * dx + (y - y1) * dy) / lengthSq;
        t = Math.max(0, Math.min(1, t));
        let projX = x1 + t * dx, projY = y1 + t * dy;
        let dist = Math.sqrt((x - projX) ** 2 + (y - projY) ** 2);
        if (dist < threshold && dist < minDist) { minDist = dist; closestSeg = seg; }
    }
    return closestSeg;
}

function getClickedPoint3(x, y, threshold) {
    let closest = null, minDist = Infinity;
    for (let el of standaloneElements3) {
        if (el.elType === 'point' && el.visProp.visible !== false) {
            let d = Math.hypot(el.X() - x, el.Y() - y);
            if (d < threshold && d < minDist) { minDist = d; closest = el; }
        }
    }
    return closest;
}

function addCustomRayFromEq3(m, p, isVert, xVert, xOrigin) {
    let currentLogicalRay = { segments: [], points: [], eq: isVert ? { isVert: true, x: xVert } : { isVert: false, m: m, p: p } };
    if (isVert) {
        let pA = board3.create('point', toVisual(xVert, Y3_MIN), { visible: false, fixed: true, highlight: false, name: '', tabindex: null });
        let pB = board3.create('point', toVisual(xVert, Y3_MAX), { visible: false, fixed: true, highlight: false, name: '', tabindex: null });
        pA.__s = xVert; pB.__s = xVert;
        let seg = board3.create('segment', [pA, pB], { strokeColor: '#555555', strokeWidth: 1.5, fixed: true, highlight: false, dash: 0, tabindex: null });
        raySegments3.push(seg); allDrawnElements3.push(pA, pB, seg);
        currentLogicalRay.segments.push(seg); currentLogicalRay.points.push(pA, pB);
    } else {
        let boundsX = [X3_MIN, X3_MAX];
        if (typeof xOrigin === 'number') boundsX.push(xOrigin);
        for (let lx of LENS_X3) {
            let y = m * lx + p;
            if (Math.abs(y) <= lensHeight3) {
                boundsX.push(lx);
                let impactPoint = board3.create('point', toVisual(lx, y), { name: '', size: 4, color: 'black', fixed: true, highlight: false, tabindex: null });
                allDrawnElements3.push(impactPoint);
                currentLogicalRay.points.push(impactPoint);
            }
        }
        boundsX.sort((a, b) => a - b);
        for (let i = 0; i < boundsX.length - 1; i++) {
            let xa = boundsX[i], xb = boundsX[i + 1];
            if (xa === xb) continue;
            let ya = m * xa + p, yb = m * xb + p;
            let pA = board3.create('point', toVisual(xa, ya), { visible: false, fixed: true, highlight: false, name: '', tabindex: null });
            let pB = board3.create('point', toVisual(xb, yb), { visible: false, fixed: true, highlight: false, name: '', tabindex: null });
            pA.__s = xa; pB.__s = xb;
            let seg = board3.create('segment', [pA, pB], { strokeColor: '#555555', strokeWidth: 1.5, fixed: true, highlight: false, dash: 0, tabindex: null });
            raySegments3.push(seg); allDrawnElements3.push(pA, pB, seg);
            currentLogicalRay.segments.push(seg); currentLogicalRay.points.push(pA, pB);
        }
    }
    logicalRays3.push(currentLogicalRay);
    return currentLogicalRay;
}

function addRayFromPieces3(m, p, isVert, xVert, pieces) {
    let currentLogicalRay = { segments: [], points: [], eq: isVert ? { isVert: true, x: xVert } : { isVert: false, m: m, p: p } };
    for (let piece of pieces) {
        let pA, pB;
        if (isVert) {
            pA = board3.create('point', toVisual(xVert, piece[0]), { visible: false, fixed: true, highlight: false, name: '', tabindex: null });
            pB = board3.create('point', toVisual(xVert, piece[1]), { visible: false, fixed: true, highlight: false, name: '', tabindex: null });
            pA.__s = xVert; pB.__s = xVert;
        } else {
            pA = board3.create('point', toVisual(piece[0], m * piece[0] + p), { visible: false, fixed: true, highlight: false, name: '', tabindex: null });
            pB = board3.create('point', toVisual(piece[1], m * piece[1] + p), { visible: false, fixed: true, highlight: false, name: '', tabindex: null });
            pA.__s = piece[0]; pB.__s = piece[1];
        }
        let seg = board3.create('segment', [pA, pB], { strokeColor: '#555555', strokeWidth: 1.5, fixed: true, highlight: false, dash: 0, tabindex: null });
        seg.__status = statusFromCode3(piece[2]);
        paintSegment3(seg);
        raySegments3.push(seg); allDrawnElements3.push(pA, pB, seg);
        currentLogicalRay.segments.push(seg); currentLogicalRay.points.push(pA, pB);
    }
    logicalRays3.push(currentLogicalRay);
    return currentLogicalRay;
}

function addCustomRay3(s1, t1, s2, t2) {
    let ds = s2 - s1, dt = t2 - t1;
    let m = (Math.abs(ds) < 0.01) ? Infinity : dt / ds;
    if (m === Infinity) {
        addCustomRayFromEq3(null, null, true, s1);
    } else {
        let p = t1 - m * s1;
        addCustomRayFromEq3(m, p, false, null, s1);
    }
    syncState3();
}

function onPointClick3(pt) {
    board3.removeObject(pt);
    allDrawnElements3 = allDrawnElements3.filter(el => el !== pt);
    standaloneElements3 = standaloneElements3.filter(el => el !== pt);
    syncState3();
}

board3.on('down', function(evt) {
    if (evt.target && evt.target.closest && evt.target.closest('.JXG_navigation_button')) return;
    if (evt.shiftKey) return;
    let coords = board3.getUsrCoordsOfMouse(evt);
    let x = coords[0], y = coords[1];

    if (toolMode3 !== '') {
        let snapped = snapToPoint3(x, y); x = snapped.x; y = snapped.y;

        if (toolMode3 === 'pt1') { tempPoint3 = { s: -y, t: x }; toolMode3 = 'pt2'; setInstructions3("RAYON (2/2) : cliquez le 2e point."); }
        else if (toolMode3 === 'pt2') { addCustomRay3(tempPoint3.s, tempPoint3.t, -y, x); resetTool3(); }
        else if (toolMode3 === 'par1') { dirPoint1c = { s: -y, t: x }; toolMode3 = 'par2'; setInstructions3("PARALLÈLE (2/3) : 2e point de direction."); }
        else if (toolMode3 === 'par2') { dirPoint2c = { s: -y, t: x }; toolMode3 = 'par3'; setInstructions3("PARALLÈLE (3/3) : point de passage."); }
        else if (toolMode3 === 'par3') {
            let ds = dirPoint2c.s - dirPoint1c.s, dt = dirPoint2c.t - dirPoint1c.t;
            let s3 = -y, t3 = x;
            addCustomRay3(s3, t3, s3 + ds, t3 + dt);
            resetTool3();
        }
        else if (toolMode3 === 'inter1') {
            interSeg1c = getClickedSegment3(x, y, 0.4, null);
            if (interSeg1c) { toolMode3 = 'inter2'; setInstructions3("INTERSECTION (2/2) : cliquez le 2e rayon."); }
        }
        else if (toolMode3 === 'inter2') {
            let seg2 = getClickedSegment3(x, y, 0.4, interSeg1c);
            if (seg2) {
                let x1 = interSeg1c.point1.X(), y1 = interSeg1c.point1.Y(), x2 = interSeg1c.point2.X(), y2 = interSeg1c.point2.Y();
                let x3 = seg2.point1.X(), y3 = seg2.point1.Y(), x4 = seg2.point2.X(), y4 = seg2.point2.Y();
                let m1 = (x2 - x1) === 0 ? Infinity : (y2 - y1) / (x2 - x1);
                let m2 = (x4 - x3) === 0 ? Infinity : (y4 - y3) / (x4 - x3);
                if (m1 !== Infinity && m2 !== Infinity && Math.abs(m1 - m2) > 0.001) {
                    let p1 = y1 - m1 * x1, p2 = y3 - m2 * x3;
                    let xi = (p2 - p1) / (m1 - m2), yi = m1 * xi + p1;
                    let in1 = xi >= Math.min(x1, x2) - 0.5 && xi <= Math.max(x1, x2) + 0.5 && yi >= Math.min(y1, y2) - 0.5 && yi <= Math.max(y1, y2) + 0.5;
                    let in2 = xi >= Math.min(x3, x4) - 0.5 && xi <= Math.max(x3, x4) + 0.5 && yi >= Math.min(y3, y4) - 0.5 && yi <= Math.max(y3, y4) + 0.5;
                    if (in1 && in2) {
                        intersectionCounter3++;
                        let name = "I" + intersectionCounter3;
                        let pInt = board3.create('point', [xi, yi], { name: name, size: 4, color: 'black', fixed: true, highlight: false, tabindex: null });
                        allDrawnElements3.push(pInt); standaloneElements3.push(pInt);
                        syncState3();
                    } else { setInstructions3("Les droites se croisent hors des segments."); }
                } else { setInstructions3("Segments parallèles : c'est le résultat attendu pour les rayons émergents (image à l'infini)."); }
                resetTool3();
            }
        }
        return;
    }

    let clickedSeg = getClickedSegment3(x, y, 0.3, null);
    if (clickedSeg) { onSegmentClick3(clickedSeg); return; }
    let clickedPt = getClickedPoint3(x, y, 0.5);
    if (clickedPt) { onPointClick3(clickedPt); return; }
});

addToolButton3('Rayon (2 clics)', 'pt1', "RAYON (1/2) : cliquez le point de départ.");
addToolButton3('Rayon parallèle (3 clics)', 'par1', "PARALLÈLE (1/3) : 1er point de direction.");
addToolButton3("Intersection (2 clics)", 'inter1', "INTERSECTION (1/2) : cliquez le 1er rayon.");
addToolButton3('Effacer le tronçon sélectionné', null, () => deleteSelectedSegment3());
addToolButton3('Tout effacer', null, () => {
    selectedSegment3 = null;
    allDrawnElements3.forEach(el => board3.removeObject(el));
    allDrawnElements3 = []; raySegments3 = []; logicalRays3 = []; standaloneElements3 = []; intersectionCounter3 = 0;
    resetTool3();
    syncState3();
});

document.body.appendChild(toolbarDiv3);
document.body.appendChild(instructionsEl3);
stack_js.resize_containing_frame("46em", document.documentElement.offsetHeight + "px");

const serialiser3 = () => {
    let rayList = logicalRays3.map(lr => {
        let pieces = lr.segments.map(seg => {
            let ends = lr.eq.isVert ? [seg.point1.X(), seg.point2.X()] : [seg.point1.__s, seg.point2.__s];
            return [ends[0], ends[1], statusCode3(seg)];
        });
        return lr.eq.isVert ? ["vert", lr.eq.x, pieces] : [lr.eq.m, lr.eq.p, pieces];
    });
    let ptList = standaloneElements3
        .filter(el => el.elType === 'point' && el.visProp.visible !== false && el.name)
        .map(el => [el.name, el.X(), el.Y()]);
    return "newton_oculaire_construction(" + JSON.stringify(rayList) + "," + JSON.stringify(ptList) + ")";
};

function clearAll3() {
    allDrawnElements3.forEach(el => board3.removeObject(el));
    allDrawnElements3 = []; raySegments3 = []; logicalRays3 = []; standaloneElements3 = []; intersectionCounter3 = 0;
}

const deserialiser3 = (value) => {
    clearAll3();
    let newState = JSON.parse(value.replace('newton_oculaire_construction(', '[').replace(/\)\s*$/, ']'));
    let rayList = newState[0], ptList = newState[1];
    for (let eq of rayList) {
        if (eq[0] === 'vert') addRayFromPieces3(null, null, true, eq[1], eq[2] || [[Y3_MIN, Y3_MAX, 0]]);
        else addRayFromPieces3(eq[0], eq[1], false, null, eq[2] || [[X3_MIN, X3_MAX, 0]]);
    }
    for (let p of ptList) {
        let pt = board3.create('point', [p[1], p[2]], { name: p[0], size: 4, color: 'black', fixed: true, highlight: false, tabindex: null });
        allDrawnElements3.push(pt);
        standaloneElements3.push(pt);
    }
    board3.update();
};

resetTool3();
stack_jxg.custom_bind(state, serialiser3, deserialiser3, [handlePoint3]);
board3.update();

const inputEl3 = document.getElementById(state);
function freezeIfReadonly3() {
    if (inputEl3 && (inputEl3.hasAttribute('readonly') || inputEl3.hasAttribute('disabled'))) {
        board3.containerObj.style.pointerEvents = 'none';
        toolbarDiv3.querySelectorAll('button').forEach(b => b.disabled = true);
        setInstructions3("Construction validée : la figure est figée.");
        return true;
    }
    return false;
}
if (inputEl3 && !freezeIfReadonly3()) {
    new MutationObserver(freezeIfReadonly3).observe(inputEl3, { attributes: true, attributeFilter: ['readonly', 'disabled'] });
}
[[/jsxgraph]]

<p>[[input:ans3]] [[validation:ans3]]</p>

<div class="stack-comment">
<p>Cliquez sur un tronçon de droite pour basculer son statut (réel/virtuel) ; cliquez sur un point construit pour le
supprimer. Par défaut, un tronçon non cliqué est considéré réel.</p>
</div>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p><strong>Étape 1 (miroir primaire) :</strong> le rayon passant par le centre C n'est pas dévié (incidence
normale) : c'est une seule et même droite d'équation y = x/4 + 3. Le rayon passant par le sommet S se réfléchit
symétriquement par rapport à l'axe : incident y = x/4, réfléchi y = -x/4. Ces deux rayons réfléchis se coupent en
B1 = (-6&nbsp;;&nbsp;1,5), et son projeté A1 = (-6&nbsp;;&nbsp;0) coïncide exactement avec F1'.</p>
<p><strong>Étape 2 (miroir secondaire) :</strong> quel que soit le point d'incidence choisi sur le miroir (incliné à
45°), le rayon réfléchi issu d'un rayon incident visant B1 passe par le même point B2 = (-1,5&nbsp;;&nbsp;-3), image
symétrique de B1 par rapport au plan du miroir. Le nouvel axe (replié à 90° par rapport à l'axe du primaire) est la
droite x = -3 ; le projeté de B2 sur ce nouvel axe donne A2 = (-3&nbsp;;&nbsp;-3).</p>
<p><strong>Étape 3 (oculaire) :</strong> la distance O2A2 vaut exactement -f2 = -3 : A2B2 est dans le plan focal
objet de l'oculaire. Le rayon issu de B2 par O2 (pente -1/2) et le rayon issu de B2 parallèle à l'axe, émergent par
F2' (pente également -1/2), ressortent donc parallèles entre eux sans se confondre : l'image finale est rejetée à
l'infini, le télescope de Newton est afocal.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[f1: 6$
tanA: 1/4$
xS1: 0$
xC1: -12$
xF1p: -6$
xA1: -6$
yB1: f1*tanA$

m_center1: tanA$
p_center1: 2*f1*tanA$
m_par1: tanA$
p_par1: 0$
m_ref1: -tanA$
p_ref1: 0$

X1_MIN: -13$

M1x2: -1/2$
M1y2: 5/2$
M2x2: -4$
M2y2: -1$
mM2: (M2y2-M1y2)/(M2x2-M1x2)$
pM2: M1y2 - mM2*M1x2$
xMmin2: min(M1x2, M2x2)$
xMmax2: max(M1x2, M2x2)$

B1x2: xA1$
B1y2: yB1$
axisX2: -3$

B2x: -3/2$
B2y: -3$
A2x: -3$
A2y: -3$

X2_MIN: -8$
X2_MAX: 1$

f2: 3$
xO3: 0$
xA2b: -f2$
yB2b: yB1$
xF3p: xO3 + f2$

m_center3: (0 - yB2b)/(xO3 - xA2b)$
p_center3: 0 - m_center3*xO3$
m_par3: 0$
p_par3: yB2b$
m_ref3: (0 - yB2b)/(xF3p - xO3)$
p_ref3: yB2b - m_ref3*xO3$
X3_MAX: 6$

xtol: 0.5$
GEOM_WEIGHT: 0.7$
line_score(full_ok, geom_ok) := if full_ok then 1 else (if geom_ok then GEOM_WEIGHT else 0)$

found_ray(rayList, mExp, pExp, mtol, ptol) := block([simp, found, r],
  simp: true,
  found: false,
  for r in rayList do (
    if listp(r) and length(r) >= 2 and numberp(r[1]) and numberp(r[2]) then (
      if abs(r[1] - mExp) < mtol and abs(r[2] - pExp) < ptol then found: true
    )
  ),
  found
)$

found_point(ptList, xExp, yExp, tol) := block([simp, found, p],
  simp: true,
  found: false,
  for p in ptList do (
    if listp(p) and length(p) = 3 and numberp(p[2]) and numberp(p[3]) then (
      if abs(p[2] - xExp) < tol and abs(p[3] - yExp) < tol then found: true
    )
  ),
  found
)$

line_pieces(rayList, m, p, mtol, ptol) := block([simp, result, r],
  simp: true,
  result: [],
  for r in rayList do (
    if listp(r) and length(r) >= 3 and numberp(r[1]) and numberp(r[2]) and listp(r[3]) then (
      if abs(r[1] - m) < mtol and abs(r[2] - p) < ptol then result: append(result, r[3])
    )
  ),
  result
)$

status_covering_x(pieces, x, tol) := block([simp, pc, lo, hi, found],
  simp: true,
  found: -1,
  for pc in pieces do (
    if listp(pc) and length(pc) >= 3 then (
      lo: min(pc[1], pc[2]),
      hi: max(pc[1], pc[2]),
      if x >= lo and x <= hi then found: pc[3]
    )
  ),
  if found = -1 then (
    for pc in pieces do (
      if listp(pc) and length(pc) >= 3 then (
        lo: min(pc[1], pc[2]),
        hi: max(pc[1], pc[2]),
        if x >= lo - tol and x <= hi + tol then found: pc[3]
      )
    )
  ),
  found
)$

eff_status(s) := if s = 0 then 1 else s$

seg_required_status(rayList, m, p, mtol, ptol, xlo, xhi, tol, want) := block([simp, pieces, s1, s2],
  simp: true,
  pieces: line_pieces(rayList, m, p, mtol, ptol),
  s1: status_covering_x(pieces, xlo + tol, tol),
  s2: status_covering_x(pieces, xhi - tol, tol),
  is(eff_status(s1) = want) and is(eff_status(s2) = want)
)$

seg_required_geom(rayList, m, p, mtol, ptol, xlo, xhi, tol) := block([simp, pieces, s1, s2],
  simp: true,
  pieces: line_pieces(rayList, m, p, mtol, ptol),
  s1: status_covering_x(pieces, xlo + tol, tol),
  s2: status_covering_x(pieces, xhi - tol, tol),
  is(s1 # -1) and is(s2 # -1)
)$

line_thru(m, p, px, py, tol) := is(abs(m*px + p - py) < tol)$

mirror_intersect2(m, p) := if abs(m - mM2) < 0.0001 then false else block([ix, iy],
  ix: (pM2 - p)/(m - mM2),
  iy: m*ix + p,
  [ix, iy]
)$

on_mirror_segment2(ix, tol) := is(ix >= xMmin2 - tol and ix <= xMmax2 + tol)$

require_mc1(mc) := block(
  if not is(safe_op(mc) = "newton_primaire_construction") or not is(length(mc) = 2) then
    "La réponse doit être générée par la construction graphique ci-dessus (fonction newton_primaire_construction(rayons, points))."
  else
    true
)$
require_mc2(mc) := block(
  if not is(safe_op(mc) = "newton_secondaire_construction") or not is(length(mc) = 2) then
    "La réponse doit être générée par la construction graphique ci-dessus (fonction newton_secondaire_construction(rayons, points))."
  else
    true
)$
require_mc3(mc) := block(
  if not is(safe_op(mc) = "newton_oculaire_construction") or not is(length(mc) = 2) then
    "La réponse doit être générée par la construction graphique ci-dessus (fonction newton_oculaire_construction(rayons, points))."
  else
    true
)$

ans1_validator(ex) := stack_seq_validator(ex, [require_mc1])$
ans2_validator(ex) := stack_seq_validator(ex, [require_mc2])$
ans3_validator(ex) := stack_seq_validator(ex, [require_mc3])$]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]][[feedback:prt2]][[feedback:prt3]]</text>
    </specificfeedback>
    <questionnote format="html">
      <text></text>
    </questionnote>
    <questiondescription format="html">
      <text></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans><![CDATA[newton_primaire_construction([[1/4,3,[[-12,-6,1]]],[1/4,0,[[-13,0,1]]],[-1/4,0,[[-6,0,1]]]],[["B1",-6,3/2],["A1",-6,0]])]]></tans>
      <boxsize>60</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords>newton_primaire_construction</allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>2</showvalidation>
      <options>validator:ans1_validator</options>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans><![CDATA[newton_secondaire_construction([[1/10,21/10,[[-6,-1,1]]],[10,12,[[-3/2,-1,1]]],[-1/8,3/4,[[-6,-2,1]]],[-8,-15,[[-2,-3/2,1]]],[0,-3,[[-8,-3,0],[-3,-3/2,0],[-3/2,1,0]]]],[["B2",-3/2,-3],["A2",-3,-3]])]]></tans>
      <boxsize>60</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords>newton_secondaire_construction</allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>2</showvalidation>
      <options>validator:ans2_validator</options>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>newton_oculaire_construction([[-1/2,0,[[-3,6,1]]],[0,3/2,[[-3,0,1]]],[-1/2,3/2,[[0,6,1]]]],[])</tans>
      <boxsize>60</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords>newton_oculaire_construction</allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>2</showvalidation>
      <options>validator:ans3_validator</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>0.3500000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text><![CDATA[[rayList, ptList]: args(ans1)$

l1_ok: is(seg_required_status(rayList, m_center1, p_center1, 0.05, 0.3, xC1, xA1, xtol, 1))$
l1_geom: is(seg_required_geom(rayList, m_center1, p_center1, 0.05, 0.3, xC1, xA1, xtol))$
l2_ok: is(seg_required_status(rayList, m_par1, p_par1, 0.05, 0.3, X1_MIN, xS1, xtol, 1))$
l2_geom: is(seg_required_geom(rayList, m_par1, p_par1, 0.05, 0.3, X1_MIN, xS1, xtol))$
l3_ok: is(seg_required_status(rayList, m_ref1, p_ref1, 0.05, 0.3, xS1, xA1, xtol, 1))$
l3_geom: is(seg_required_geom(rayList, m_ref1, p_ref1, 0.05, 0.3, xS1, xA1, xtol))$

l1_score: line_score(l1_ok, l1_geom)$
l2_score: line_score(l2_ok, l2_geom)$
l3_score: line_score(l3_ok, l3_geom)$

att1a: is(found_ray(rayList, m_center1, p_center1, 0.05, 0.3))$
att1b: is(found_ray(rayList, m_par1, p_par1, 0.05, 0.3) or found_ray(rayList, m_ref1, p_ref1, 0.05, 0.3))$
any_attempt1: is(att1a or att1b)$
c1a: is(l1_ok)$
c1a_geom: is(l1_geom)$
c1b: is(l2_ok and l3_ok)$
c1b_geom: is(l2_geom and l3_geom)$

c_pointB1: is(found_point(ptList, xA1, yB1, 0.3))$
c_pointA1: is(found_point(ptList, xA1, 0, 0.3))$

nb_attempts1: length(sublist(rayList, lambda([r], not stringp(r[1]))))$
nb_total_L1: max(nb_attempts1, 3)$
extra1: is(nb_attempts1 > 3)$
score_rayons_L1: min((l1_score + l2_score + l3_score) / nb_total_L1, 1)$]]></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Miroir primaire — rayons particuliers (objet à l'infini)</description>
        <answertest>AlgEquiv</answertest>
        <sans>true</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>score_rayons_L1*3/5</truescore>
        <truepenalty></truepenalty>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt1-0-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p><strong>Miroir primaire :</strong></p>
[[if test="any_attempt1"]]<ul>
[[if test="att1a"]]<li>Rayon passant par C (non dévié) : [[if test="c1a"]]<span style="color:#1e7e34;">&#10003; correct</span>[[else]][[if test="c1a_geom"]]<span style="color:#d68910;">bien positionné, mais pas entièrement marqué réel</span>[[else]]<span style="color:#c0392b;">tracé mais pas encore correct géométriquement</span>[[/if]][[/if]]</li>[[/if]]
[[if test="att1b"]]<li>Rayon passant par S (incident + réfléchi) : [[if test="c1b"]]<span style="color:#1e7e34;">&#10003; correct</span>[[else]][[if test="c1b_geom"]]<span style="color:#d68910;">bien positionné, mais un tronçon n'a pas le bon statut réel/virtuel</span>[[else]]<span style="color:#c0392b;">incomplet ou pas encore correct géométriquement</span>[[/if]][[/if]]</li>[[/if]]
</ul>[[if test="extra1"]]<p style="color:#c0392b;">{@nb_attempts1@} droites tracées alors que 3 suffisent : une droite en trop
fait baisser le score — mieux vaut la supprimer.</p>[[/if]][[else]]<p><span style="color:#c0392b;">Aucun rayon n'a encore été tracé.</span></p>[[/if]]]]></text>
        </truefeedback>
        <falsescoremode>+</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-0-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p></p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Points B1 et A1</description>
        <answertest>AlgEquiv</answertest>
        <sans>c_pointB1 and c_pointA1</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>2/5</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>B1 est bien à l'intersection des deux rayons réfléchis, et son projeté A1 coïncide avec F1'.</p>]]></text>
        </truefeedback>
        <falsescoremode>+</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>B1 et/ou A1 ne sont pas encore placés correctement : tracez les deux rayons réfléchis, placez B1 avec
« Intersection », puis utilisez « Tracer A1 (perpendiculaire) » en cliquant sur B1.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>0.3500000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text><![CDATA[[rayList, ptList]: args(ans2)$

incident_candidates2: []$
for r in rayList do (
  if listp(r) and length(r) >= 3 and numberp(r[1]) and numberp(r[2]) and listp(r[3]) and length(r[3]) > 0 then (
    if line_thru(r[1], r[2], B1x2, B1y2, 0.3) then (
      inter: mirror_intersect2(r[1], r[2]),
      if inter # false and on_mirror_segment2(inter[1], xtol) then
        incident_candidates2: append(incident_candidates2, [[r[3], inter[1]]])
    )
  )
)$

reflected_candidates2: []$
for r in rayList do (
  if listp(r) and length(r) >= 3 and numberp(r[1]) and numberp(r[2]) and listp(r[3]) and length(r[3]) > 0 then (
    if line_thru(r[1], r[2], B2x, B2y, 0.3) then (
      inter: mirror_intersect2(r[1], r[2]),
      if inter # false and on_mirror_segment2(inter[1], xtol) then
        reflected_candidates2: append(reflected_candidates2, [[r[3], inter[1]]])
    )
  )
)$

any_attempt2: is(length(incident_candidates2) > 0 or length(reflected_candidates2) > 0)$

all_pairs2: []$
for ic in incident_candidates2 do (
  for rc in reflected_candidates2 do (
    Ix: ic[2],
    Ix2: rc[2],
    if abs(Ix - Ix2) < xtol then (
      inc_pieces: ic[1],
      refl_pieces: rc[1],
      xNear: min(B1x2, Ix), xFar: max(B1x2, Ix),
      dirsign1: if Ix >= B1x2 then 1 else -1,
      xBeyond1: Ix + dirsign1*1.0,
      inc_real_ok: is(eff_status(status_covering_x(inc_pieces, xNear + xtol, xtol)) = 1
                and eff_status(status_covering_x(inc_pieces, xFar - xtol, xtol)) = 1),
      inc_absent_ok: is(status_covering_x(inc_pieces, xBeyond1, xtol) = -1),
      inc_full: is(inc_real_ok and inc_absent_ok),
      xRNear: min(Ix, B2x) + xtol, xRFar: max(Ix, B2x) - xtol,
      refl_full: is(eff_status(status_covering_x(refl_pieces, xRNear, xtol)) = 1
                and eff_status(status_covering_x(refl_pieces, xRFar, xtol)) = 1),
      inc_score: if inc_full then 1 else GEOM_WEIGHT,
      refl_score: if refl_full then 1 else GEOM_WEIGHT,
      pair_score: (inc_score + refl_score)/2,
      all_pairs2: append(all_pairs2, [[Ix, pair_score]])
    )
  )
)$

accepted2: []$
for pr in all_pairs2 do (
  is_dup: false,
  for a in accepted2 do (if abs(a[1] - pr[1]) < xtol then is_dup: true),
  if not is_dup and length(accepted2) < 2 then accepted2: append(accepted2, [pr])
)$

nb_accepted2: length(accepted2)$
nb_bonnes2: if nb_accepted2 = 0 then 0 elseif nb_accepted2 = 1 then accepted2[1][2] else accepted2[1][2] + accepted2[2][2]$
score_rayons2: min(nb_bonnes2/2, 1)$

have2b: is(nb_accepted2 >= 2)$

c_pointB2: is(found_point(ptList, B2x, B2y, 0.3))$
c_pointA2: is(found_point(ptList, A2x, A2y, 0.3))$]]></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Constructions incidence/réflexion à deux points d'incidence distincts sur le secondaire</description>
        <answertest>AlgEquiv</answertest>
        <sans>true</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>score_rayons2*3/5</truescore>
        <truepenalty></truepenalty>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt2-0-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p><strong>Miroir secondaire :</strong></p>
[[if test="any_attempt2"]]<p>Constructions valides détectées (rayon incident visant B1, associé à un rayon réfléchi
réel passant par le même point d'incidence) : {@nb_accepted2@} sur 2 nécessaires.</p>
[[if test="have2b"]]<p style="color:#1e7e34;">&#10003; Deux points d'incidence distincts ont été correctement
traités (vérifiez que chaque rayon réfléchi est bien marqué réel entre le miroir et B2).</p>[[else]]<p style="color:#c0392b;">Il
manque au moins une construction valide à un second point d'incidence, distinct du premier.</p>[[/if]]
[[else]]<p style="color:#c0392b;">Aucune construction (rayon incident issu de B1, associé à un rayon réfléchi passant
par le même point du miroir) n'a encore été détectée.</p>[[/if]]]]></text>
        </truefeedback>
        <falsescoremode>+</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-0-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p></p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Point B2</description>
        <answertest>AlgEquiv</answertest>
        <sans>c_pointB2</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1/5</truescore>
        <truepenalty></truepenalty>
        <truenextnode>2</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>B2 est correctement placé à l'intersection des deux rayons réfléchis réels.</p>]]></text>
        </truefeedback>
        <falsescoremode>+</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Le point B2 n'est pas encore placé correctement : il ne peut l'être que si deux constructions
(incidence + réflexion) sont d'abord correctement réalisées en deux points distincts du miroir secondaire (voir le
bilan ci-dessus).</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Point A2 (projeté sur le nouvel axe)</description>
        <answertest>AlgEquiv</answertest>
        <sans>c_pointA2</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>1/5</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>A2 est correctement placé sur le nouvel axe (replié à 90°), au pied du projeté de B2.</p>]]></text>
        </truefeedback>
        <falsescoremode>+</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>A2 n'est pas encore placé correctement : utilisez « Tracer A2 (perpendiculaire au nouvel axe) » en
cliquant sur B2 (après l'avoir correctement construit).</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>0.3000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text><![CDATA[[rayList, ptList]: args(ans3)$

l1_ok3: is(seg_required_status(rayList, m_center3, p_center3, 0.05, 0.3, xA2b, X3_MAX, xtol, 1))$
l1_geom3: is(seg_required_geom(rayList, m_center3, p_center3, 0.05, 0.3, xA2b, X3_MAX, xtol))$
l2_ok3: is(seg_required_status(rayList, m_par3, p_par3, 0.05, 0.3, xA2b, xO3, xtol, 1))$
l2_geom3: is(seg_required_geom(rayList, m_par3, p_par3, 0.05, 0.3, xA2b, xO3, xtol))$
l3_ok3: is(seg_required_status(rayList, m_ref3, p_ref3, 0.05, 0.3, xO3, X3_MAX, xtol, 1))$
l3_geom3: is(seg_required_geom(rayList, m_ref3, p_ref3, 0.05, 0.3, xO3, X3_MAX, xtol))$

l1_score3: line_score(l1_ok3, l1_geom3)$
l2_score3: line_score(l2_ok3, l2_geom3)$
l3_score3: line_score(l3_ok3, l3_geom3)$

att3a: is(found_ray(rayList, m_center3, p_center3, 0.05, 0.3))$
att3b: is(found_ray(rayList, m_par3, p_par3, 0.05, 0.3) or found_ray(rayList, m_ref3, p_ref3, 0.05, 0.3))$
any_attempt3: is(att3a or att3b)$
c3a: is(l1_ok3)$
c3a_geom: is(l1_geom3)$
c3b: is(l2_ok3 and l3_ok3)$
c3b_geom: is(l2_geom3 and l3_geom3)$

nb_attempts3: length(sublist(rayList, lambda([r], not stringp(r[1]))))$
nb_total_L3: max(nb_attempts3, 3)$
extra3: is(nb_attempts3 > 3)$
score_rayons_L3: min((l1_score3 + l2_score3 + l3_score3) / nb_total_L3, 1)$]]></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Oculaire — rayons particuliers (objet réel dans le plan focal objet, résultat afocal)</description>
        <answertest>AlgEquiv</answertest>
        <sans>true</sans>
        <tans>true</tans>
        <testoptions></testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>score_rayons_L3</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-0-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p><strong>Oculaire :</strong></p>
[[if test="any_attempt3"]]<ul>
[[if test="att3a"]]<li>Rayon issu de B2 par O2 : [[if test="c3a"]]<span style="color:#1e7e34;">&#10003; correct</span>[[else]][[if test="c3a_geom"]]<span style="color:#d68910;">bien positionné, mais pas entièrement marqué réel</span>[[else]]<span style="color:#c0392b;">tracé mais pas encore correct géométriquement</span>[[/if]][[/if]]</li>[[/if]]
[[if test="att3b"]]<li>Rayon issu de B2 parallèle à l'axe, émergent vers F2' : [[if test="c3b"]]<span style="color:#1e7e34;">&#10003; correct (incident et émergent)</span>[[else]][[if test="c3b_geom"]]<span style="color:#d68910;">bien positionné, mais un tronçon n'a pas le bon statut réel/virtuel</span>[[else]]<span style="color:#c0392b;">incomplet ou pas encore correct géométriquement</span>[[/if]][[/if]]</li>[[/if]]
</ul>[[if test="c3a and c3b"]]<p style="color:#1e7e34;">Les deux rayons émergents sont bien parallèles entre eux (même pente
-1/2) : l'image finale est rejetée à l'infini, le télescope est afocal.</p>[[/if]][[if test="extra3"]]<p style="color:#c0392b;">{@nb_attempts3@}
droites tracées alors que 3 suffisent : une droite en trop fait baisser le score — mieux vaut la supprimer.</p>[[/if]][[else]]<p><span style="color:#c0392b;">Aucun rayon n'a encore été tracé.</span></p>[[/if]]]]></text>
        </truefeedback>
        <falsescoremode>+</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-0-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p></p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <qtest>
      <testcase>1</testcase>
      <description>Construction complète et correcte des trois étapes : score plein</description>
      <testinput>
        <name>ans1</name>
        <value><![CDATA[newton_primaire_construction([[1/4,3,[[-12,-6,1]]],[1/4,0,[[-13,0,1]]],[-1/4,0,[[-6,0,1]]]],[["B1",-6,3/2],["A1",-6,0]])]]></value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value><![CDATA[newton_secondaire_construction([[1/10,21/10,[[-6,-1,1]]],[10,12,[[-3/2,-1,1]]],[-1/8,3/4,[[-6,-2,1]]],[-8,-15,[[-2,-3/2,1]]],[0,-3,[[-8,-3,0],[-3,-3/2,0],[-3/2,1,0]]]],[["B2",-3/2,-3],["A2",-3,-3]])]]></value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>newton_oculaire_construction([[-1/2,0,[[-3,6,1]]],[0,3/2,[[-3,0,1]]],[-1/2,3/2,[[0,6,1]]]],[])</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.3500000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-0-T|prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.3500000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-0-T|prt2-1-T|prt2-2-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>0.3000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-0-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description>Construction entièrement vide sur les trois étapes</description>
      <testinput>
        <name>ans1</name>
        <value>newton_primaire_construction([],[])</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>newton_secondaire_construction([],[])</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>newton_oculaire_construction([],[])</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-0-F|prt1-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-0-F|prt2-1-F|prt2-2-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-0-F</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description>Étape 1 (primaire) seule complète, étapes 2 et 3 vides</description>
      <testinput>
        <name>ans1</name>
        <value><![CDATA[newton_primaire_construction([[1/4,3,[[-12,-6,1]]],[1/4,0,[[-13,0,1]]],[-1/4,0,[[-6,0,1]]]],[["B1",-6,3/2],["A1",-6,0]])]]></value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>newton_secondaire_construction([],[])</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>newton_oculaire_construction([],[])</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.3500000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-0-T|prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-0-F|prt2-1-F|prt2-2-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-0-F</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>4</testcase>
      <description>Miroir secondaire : un seul point d'incidence correctement traité (pas de second point, pas de B2/A2)</description>
      <testinput>
        <name>ans1</name>
        <value>newton_primaire_construction([],[])</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>newton_secondaire_construction([[1/10,21/10,[[-6,-1,1]]],[10,12,[[-3/2,-1,1]]]],[])</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>newton_oculaire_construction([],[])</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-0-F|prt1-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.1050000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-0-T|prt2-1-F|prt2-2-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-0-F</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>5</testcase>
      <description>Tout correct sauf un rayon superflu non supprimé sur le primaire (dilution du score) et un tronçon réfléchi du secondaire marqué virtuel au lieu de réel (crédit géométrie partiel 70%)</description>
      <testinput>
        <name>ans1</name>
        <value><![CDATA[newton_primaire_construction([[1/4,3,[[-12,-6,1]]],[1/4,0,[[-13,0,1]]],[-1/4,0,[[-6,0,1]]],[1/2,1,[[-6,0,1]]]],[["B1",-6,3/2],["A1",-6,0]])]]></value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value><![CDATA[newton_secondaire_construction([[1/10,21/10,[[-6,-1,1]]],[10,12,[[-3/2,-1,2]]],[-1/8,3/4,[[-6,-2,1]]],[-8,-15,[[-2,-3/2,1]]],[0,-3,[[-8,-3,0],[-3,-3/2,0],[-3/2,1,0]]]],[["B2",-3/2,-3],["A2",-3,-3]])]]></value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>newton_oculaire_construction([[-1/2,0,[[-3,6,1]]],[0,3/2,[[-3,0,1]]],[-1/2,3/2,[[0,6,1]]]],[])</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.3150000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-0-T|prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.3342500</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-0-T|prt2-1-T|prt2-2-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>0.3000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-0-T</expectedanswernote>
      </expected>
    </qtest>
  </question>

</quiz>