<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 396264  -->
  <question type="stack">
    <name>
      <text>Géométrie : Ordonnée à l'origine d'une droite</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<div style="background: #475569; border-left: 5px solid #334155; border-radius: 0 8px 8px 0; padding: 10px 16px; margin-bottom: 12px; display: flex; align-items: center; gap: 10px; flex-wrap: wrap;"><strong style="font-weight: 800; color: #fff; font-size: .95rem;">Géométrie — Équation de droite</strong> <span style="background: #334155; color: #fff; padding: 2px 9px; border-radius: 20px; font-size: .78rem; font-weight: bold;">/ 1 pt</span></div>
<p>Soit \((d)\) la droite dont l'équation cartésienne est :</p>
<div style="text-align: center; margin: 15px 0; font-size: 1.1rem; background: #f8fafc; padding: 10px; border-radius: 6px; border: 1px solid #e2e8f0; display: inline-block;">\({@a*x + b*y + c@} = 0\)</div>
<p><strong>1.</strong> Déterminer l'ordonnée à l'origine de la droite \((d)\).</p>
<p style="font-size: .85rem; color: #64748b; margin-bottom: 8px;"><em>Saisir la valeur exacte (ex: <code style="background: #f1f5f9; padding: 2px 6px; border-radius: 4px;">3/2</code> ou <code style="background: #f1f5f9; padding: 2px 6px; border-radius: 4px;">-4</code>).</em></p>
<div>L'ordonnée à l'origine est : [[input:ans_oao]] [[validation:ans_oao]]</div>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<div style="margin-top: 20px; padding: 15px; background: #f8fafc; border: 1px solid #cbd5e1; border-radius: 8px;">
<div style="font-weight: bold; margin-bottom: 10px;">🔑 Correction détaillée</div>
<div style="font-size: .9rem;">
L'ordonnée à l'origine correspond à la valeur de \(y\) lorsque la droite coupe l'axe des ordonnées, c'est-à-dire lorsque <strong>\(x = 0\)</strong>.<br><br>
On remplace \(x\) par \(0\) dans l'équation {@a*x + b*y + c@} = 0 :<br><br>
\[
{@a@}(0) + {@b@}y + {@c@} = 0
\]
\[
{@b@}y + {@c@} = 0 \implies {@b@}y = -{@c@} \implies y = \frac{-{@c@}}{{@b@}}
\]
\[
y = {@oao@}
\]<br>
L'ordonnée à l'origine de la droite est donc {@oao@}.
</div>
</div>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2024092500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* Fonctions utilitaires */
ri(a,b) := a + rand(b - a + 1)$ rnz(a,b) := block([v], v: ri(a,b), if v=0 then rnz(a,b) else v)$ 
/* Génération des coefficients */
/* a et b non nuls pour éviter les droites horizontales/verticales */
a: rnz(-5,5)$ b: rnz(-5,5)$ 
/* c peut être nul (droite passant par l'origine) */
c: ri(-5,5)$ 
/* Ordonnée à l'origine correcte : y quand x=0 => b*y + c = 0 */
oao: -c/b$ 
/* Erreurs classiques */
err_constante: c$ err_sans_moins: c/b$ err_pente: -a/b$ err_abscisse: -c/a$]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<div>[[feedback:prt1]]</div> ]]></text>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>Eq: {@a*x+b*y+c@}=0, OAQ={@oao@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p>Généré avec HéStack V2 | hestack-types:geometry-y-intercept</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans_oao</name>
      <type>algebraic</type>
      <tans>oao</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Réponse correcte</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_oao</sans>
        <tans>oao</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-OK</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #15803d; padding: 10px 14px; background: #f0fdf4; border-radius: 4px;">✅ <strong>Parfait !</strong> C'est bien la valeur de \(y\) lorsque \(x = 0\).</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>PRT-NOK</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Erreur : A juste lu la constante c</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_oao</sans>
        <tans>err_constante</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-CONST</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Confusion avec la constante !</strong> Vous avez donné la valeur de \(c\). L'ordonnée à l'origine n'est pas \(c\) (sauf si l'équation est réduite \(y = mx + p\)). Ici, il faut isoler \(y\) pour trouver \(p\), ou poser \(x=0\).</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>PRT-ERR-NOT-CONST</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Erreur : Oubli du signe moins</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_oao</sans>
        <tans>err_sans_moins</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-SIGN</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Erreur de signe !</strong> Vous avez calculé \(\frac{c}{b}\). Lorsque vous isolez \(y\) (\({@b@}y = -{@c@}\)), n'oubliez pas que le signe "moins" passe au numérateur : \(y = \frac{-{@c@}}{{@b@}}\).</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>PRT-ERR-NOT-SIGN</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Erreur : A calculé le coefficient directeur</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_oao</sans>
        <tans>err_pente</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-PENTE</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Confusion pente / ordonnée !</strong> Vous avez calculé le <strong>coefficient directeur</strong> (la pente) \(-\frac{a}{b}\). L'ordonnée à l'origine est l'endroit où la droite <strong>coupe l'axe vertical</strong> (quand \(x=0\)).</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>4</falsenextnode>
        <falseanswernote>PRT-ERR-NOT-PENTE</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description>Erreur : A calculé l'abscisse à l'origine</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_oao</sans>
        <tans>err_abscisse</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-ABSCISSE</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Axe faux !</strong> Vous avez calculé l'<strong>abscisse</strong> à l'origine (quand \(y=0\)). On vous demande l'<strong>ordonnée</strong> à l'origine, c'est-à-dire la valeur de \(y\) quand \(x=0\).</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>5</falsenextnode>
        <falseanswernote>PRT-ERR-NOT-ABSCISSE</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>5</name>
        <description>Erreur générique (Fallback)</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_oao</sans>
        <tans>oao</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-GENERIC</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #dc2626; padding: 10px 14px; background: #fff0f0; border-radius: 4px;">❌ <strong>Incorrect.</strong> Pour trouver l'ordonnée à l'origine, il faut remplacer \(x\) par \(0\) dans l'équation et résoudre pour \(y\).</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>PRT-ERR-FALLBACK</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>