<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 396257  -->
  <question type="stack">
    <name>
      <text>Géométrie : Milieu d'un segment</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<div style="background: #475569; border-left: 5px solid #334155; border-radius: 0 8px 8px 0; padding: 10px 16px; margin-bottom: 12px; display: flex; align-items: center; gap: 10px; flex-wrap: wrap;"><strong style="font-weight: 800; color: #fff; font-size: .95rem;">Géométrie — Milieu</strong> <span style="background: #334155; color: #fff; padding: 2px 9px; border-radius: 20px; font-size: .78rem; font-weight: bold;">/ 1 pt</span></div>
<p>Dans un repère orthonormé \((O; \vec{i}, \vec{j})\), on considère les points :</p>
<div style="display: flex; gap: 40px; justify-content: center; margin: 15px 0; flex-wrap: wrap;">
<div>\(A \left( {@xa@} \; ; \; {@ya@} \right)\)</div>
<div>\(B \left( {@xb@} \; ; \; {@yb@} \right)\)</div>
</div>
<p><strong>1.</strong> Déterminer les coordonnées du milieu \(I\) du segment \([AB]\).</p>
<p style="font-size: .85rem; color: #64748b; margin-bottom: 8px;"><em>Saisir sous la forme : <code style="background: #f1f5f9; padding: 2px 6px; border-radius: 4px;">matrix([x],[y])</code> (les valeurs exactes comme 3/2 sont acceptées).</em></p>
<div>\(I \left( \right.\) [[input:ans_mid]] [[validation:ans_mid]] \(\left. \right)\)</div>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<div style="margin-top: 20px; padding: 15px; background: #f8fafc; border: 1px solid #cbd5e1; border-radius: 8px;">
<div style="font-weight: bold; margin-bottom: 10px;">🔑 Correction détaillée</div>
<div style="font-size: .9rem;">
Les coordonnées du milieu d'un segment sont données par la <strong>moyenne arithmétique</strong> des coordonnées des extrémités :<br><br>
\[
I \left( \frac{x_A + x_B}{2} \; ; \; \frac{y_A + y_B}{2} \right)
\]<br>
<strong>Application :</strong><br><br>
\[
x_I = \frac{{@xa@} + ({@xb@})}{2} = \frac{{@somme_x@}}{2} = {@x_mid@}
\]
\[
y_I = \frac{{@ya@} + ({@yb@})}{2} = \frac{{@somme_y@}}{2} = {@y_mid@}
\]<br>
Donc \(I \left( {@x_mid@} \; ; \; {@y_mid@} \right)\).
</div>
</div>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>/* Fonction utilitaire */
ri(a,b) := a + rand(b - a + 1);
/* Génération des coordonnées */
xa: ri(-5,5);
 ya: ri(-5,5);
 xb: ri(-5,5);
 yb: ri(-5,5); 
/* On s'assure que A et B ne sont pas confondus */
while xa=xb and ya=yb do (
    xb: ri(-5,5),
    yb: ri(-5,5)
)$ 
/* Calculs pour la correction */
somme_x: xa + xb;
 somme_y: ya + yb;
 x_mid: somme_x / 2;
 y_mid: somme_y / 2;
 M_mid: matrix([x_mid], [y_mid]); 
/* Erreurs classiques */
/* Erreur 1 : Inversion des coordonnées x et y */
err_swap: matrix([y_mid], [x_mid]);
 /* Erreur 2 : A fait une différence au lieu d'une somme */
err_diff: matrix([(xa - xb)/2], [(ya - yb)/2]);
 /* Erreur 3 : A oublié de diviser la deuxième coordonnée par 2 */
err_no_div_y: matrix([x_mid], [somme_y]);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<div>[[feedback:prt1]]</div>]]></text>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>A({@xa@};{@ya@}), B({@xb@};{@yb@}), I({@x_mid@};{@y_mid@})</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p>Généré avec HéStack V2 | hestack-types:geometry-midpoint</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans_mid</name>
      <type>matrix</type>
      <tans>M_mid</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>matrix([x],[y])</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Réponse correcte</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_mid</sans>
        <tans>M_mid</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-OK</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #15803d; padding: 10px 14px; background: #f0fdf4; border-radius: 4px;">✅ <strong>Parfait !</strong> Les coordonnées du milieu sont correctes.</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>PRT-NOK</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Erreur : Inversion de x et y</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_mid</sans>
        <tans>err_swap</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-SWAP</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Coordonnées inversées !</strong> Vous avez mis la moyenne des \(y\) en abscisse (\(x\)) et la moyenne des \(x\) en ordonnée (\(y\)). Attention à l'ordre : \(I \left( \frac{x_A+x_B}{2} \; ; \; \frac{y_A+y_B}{2} \right)\).</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>PRT-ERR-NOT-SWAP</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Erreur : Différence au lieu de la somme</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_mid</sans>
        <tans>err_diff</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-DIFF</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Erreur de signe !</strong> Vous avez calculé \(\frac{x_A - x_B}{2}\) au lieu de \(\frac{x_A + x_B}{2}\). Pour le milieu, il faut <strong>additionner</strong> les coordonnées, pas les soustraire.</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>PRT-ERR-NOT-DIFF</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Erreur : Oubli de la division par 2 sur y</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_mid</sans>
        <tans>err_no_div_y</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-NO-DIV-Y</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Calcul incomplet !</strong> Votre abscisse (\(x\)) est correcte, mais pour l'ordonnée (\(y\)), vous avez oublié de <strong>diviser par 2</strong>. La formule est une moyenne, il faut donc diviser par 2 les <strong>deux</strong> coordonnées.</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>PRT-ERR-CALC</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #dc2626; padding: 10px 14px; background: #fff0f0; border-radius: 4px;">❌ <strong>Incorrect.</strong> Appliquez la formule du milieu :<br>\(I \left( \frac{x_A + x_B}{2} \; ; \; \frac{y_A + y_B}{2} \right)\)</div>]]></text>
        </falsefeedback>
      </node>
    </prt>
  </question>

</quiz>