<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 396261  -->
  <question type="stack">
    <name>
      <text>Géométrie 3D : Pente d'une droite</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<div style="background: #475569; border-left: 5px solid #334155; border-radius: 0 8px 8px 0; padding: 10px 16px; margin-bottom: 12px; display: flex; align-items: center; gap: 10px; flex-wrap: wrap;"><strong style="font-weight: 800; color: #fff; font-size: .95rem;">Géométrie dans l'espace — Pente</strong> <span style="background: #334155; color: #fff; padding: 2px 9px; border-radius: 20px; font-size: .78rem; font-weight: bold;">/ 1 pt</span></div>
<p>Dans un repère orthonormé de l'espace \((O; \vec{i}, \vec{j}, \vec{k})\), on considère la droite \((d)\) passant par le point :</p>
<div style="text-align: center; margin: 10px 0;">\(A \left( {@xa@} \; ; \; {@ya@} \; ; \; {@za@} \right)\)</div>
<p>et de vecteur directeur :</p>
<div style="text-align: center; margin: 10px 0;">\(\vec{u} \begin{pmatrix} {@a@} \\ {@b@} \\ {@c@} \end{pmatrix}\)</div>
<p><strong>Calculer</strong> la pente de la droite \((d)\), c'est-à-dire la pente de sa projection orthogonale dans le plan horizontal \((O; \vec{i}, \vec{j})\).</p>
<p style="font-size: .85rem; color: #64748b; margin-bottom: 8px;"><em>Saisir le résultat sous forme de fraction exacte (ex: <code style="background: #f1f5f9; padding: 2px 6px; border-radius: 4px;">3/2</code> ou <code style="background: #f1f5f9; padding: 2px 6px; border-radius: 4px;">-4/5</code>).</em></p>
<div>\(p = \) [[input:ans_pente]] [[validation:ans_pente]]</div>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<div style="margin-top: 20px; padding: 15px; background: #f8fafc; border: 1px solid #cbd5e1; border-radius: 8px;">
<div style="font-weight: bold; margin-bottom: 10px;">🔑 Correction détaillée</div>
<div style="font-size: .9rem;">
Le plan horizontal \((O; \vec{i}, \vec{j})\) est le plan \((xOy)\).<br>
La projection de la droite \((d)\) dans ce plan est une droite 2D dont le coefficient directeur est la <strong>pente</strong> cherchée.<br><br>
Une représentation paramétrique de \((d)\) est :<br>
\(x = {@xa@} + {@a@}t\)<br>
\(y = {@ya@} + {@b@}t\)<br>
\(z = {@za@} + {@c@}t\)<br><br>
Dans le plan \((xOy)\), on a \(x = {@xa@} + {@a@}t\) et \(y = {@ya@} + {@b@}t\).<br>
On isole \(t\) : \(t = \frac{x - {@xa@}}{{@a@}}\)<br><br>
On injecte dans \(y\) :<br>
\(y = {@ya@} + {@b@} \left( \frac{x - {@xa@}}{{@a@}} \right)\)<br><br>
Le coefficient directeur (la pente) est donc le quotient des coordonnées du vecteur directeur :<br>
\[
p = \frac{{@b@}}{{@a@}} = {@pente@}
\]
<em style="color: #64748b;">Remarque : La composante en \(z\) ({@c@}) n'intervient pas dans le calcul de la pente "2D" de la droite.</em>
</div>
</div>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>/* Fonctions utilitaires */
ri(a,b) := a + rand(b - a + 1)$ rnz(a,b) := block([v], v: ri(a,b), if v=0 then rnz(a,b) else v)$ 
/* Point A */
xa: ri(-4,4)$ ya: ri(-4,4)$ za: ri(-4,4)$ 
/* Vecteur directeur u(a,b,c) */
/* a et b doivent être non nuls pour éviter une pente nulle ou infinie */
a: rnz(-4,4)$ b: rnz(-4,4)$ c: ri(-4,4)$ 
/* Pente correcte (Projection sur xOy : dy/dx = b/a) */
pente: b/a$ 
/* Erreurs classiques */
/* Erreur 1 : A inversé la fraction (dx/dy au lieu de dy/dx) */
err_inv: a/b$ /* Erreur 2 : A utilisé la composante z au lieu de y */
err_z: c/a$ /* Erreur 3 : A calculé la pente de la projection sur xOz */
err_xz: c/a$ /* Erreur 4 : Erreur de signe global */
err_moins: -b/a$</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<div>[[feedback:prt1]]</div>]]></text>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>A({@xa@};{@ya@};{@za@}), u({@a@};{@b@};{@c@}), pente={@pente@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p>Généré avec HéStack V2 | hestack-types:geometry-3d-slope</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans_pente</name>
      <type>algebraic</type>
      <tans>pente</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Réponse correcte</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_pente</sans>
        <tans>pente</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-OK</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #15803d; padding: 10px 14px; background: #f0fdf4; border-radius: 4px;">✅ <strong>Parfait !</strong> Vous avez bien identifié que la pente correspond au rapport des variations en \(y\) sur celles en \(x\).</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>PRT-NOK</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Erreur : Fraction inversée</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_pente</sans>
        <tans>err_inv</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-INV</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Fraction inversée !</strong> Vous avez calculé \(\frac{a}{b}\) (ou \(\frac{\Delta x}{\Delta y}\)). La pente d'une droite est toujours l'élévation (\(\Delta y\)) divisée par l'avancement horizontal (\(\Delta x\)), soit \(\frac{b}{a}\).</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>PRT-ERR-NOT-INV</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Erreur : A utilisé la composante z</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_pente</sans>
        <tans>err_z</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-Z</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Mauvaise composante !</strong> Vous avez utilisé la coordonnée en \(z\) (la hauteur) au lieu de la coordonnée en \(y\). La pente dans le plan horizontal \((xOy)\) se calcule uniquement avec \(x\) et \(y\) : \(\frac{b}{a}\).</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>PRT-ERR-NOT-Z</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Erreur : Erreur de signe</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_pente</sans>
        <tans>err_moins</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-MOINS</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Erreur de signe !</strong> Le calcul des numérateurs/dénominateurs est bon, mais le signe final est incorrect. Vérifiez la règle des signes lors du quotient.</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>PRT-ERR-CALC</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #dc2626; padding: 10px 14px; background: #fff0f0; border-radius: 4px;">❌ <strong>Incorrect.</strong> La pente de la droite est égale au rapport de la composante en \(y\) sur la composante en \(x\) du vecteur directeur :<br>\(p = \frac{b}{a}\)</div>]]></text>
        </falsefeedback>
      </node>
    </prt>
  </question>

</quiz>