<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 245108  -->
  <question type="stack">
    <name>
      <text>conversion d'unités autogénérées</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Convertit {@val_depart@}  en {@y@}</p>
<p>[[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text></text>
    </generalfeedback>
    <defaultgrade>2</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[val: float(rand(990)/10);

multiple: [["f",1E-15],["p",1E-12],["n",1E-9],["µ",1E-6],["m",1E-3],["c",1E-2],["d",1E-1],[" ",1],["da",1E1],["h",1E2],["k",1E3],["M",1E6],["G",1E9],["T",1E12]];
/* Choisir une sous-liste au hasard dans la liste multiple */
choix_liste1: rand(multiple);
choix_liste2: rand(multiple);

while choix_liste1 = choix_liste2 do
 choix_liste2: rand(multiple);
end;

/* Extraire le préfixe (premier élément de la sous-liste) */
choix1: choix_liste1[1];
choix2: choix_liste2[1];
/* Extraire le facteur (deuxième élément de la sous-liste) */
choix_a: choix_liste1[2];
choix_b: choix_liste2[2];

unités: ["W", "Pa", "J", "cd", "K", "g", "N", "V", "A", "s", "m"];
unité_choisie: rand(unités);
liste_combine1: [choix1, unité_choisie];
liste_combine2: [choix2, unité_choisie];

/* On utilise 'sconcat' pour créer les unités sous forme de chaînes de caractères */
x: sconcat(choix1, unité_choisie);
y: sconcat(choix2, unité_choisie);

/* On définit la variable de départ */
val_depart: val * x;

/* Création de la variable 'ta1' */
ta1: val *choix_a /choix_b;
z: choix_a /choix_b;
ta5:?;]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>[[feedback:prt1]]</p>]]></text>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>convertir {@val_depart@}  en {@y@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<p><i class="fa fa-check"></i> Réponse correcte, bravo !</p>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<p><i class="fa fa-adjust"></i> Réponse partiellement correcte.</p>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<p><i class="fa fa-times"></i> Réponse incorrecte.</p>]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>2</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>test</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions>0.01</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Vous avez bien fais la conversion.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Votre conversion n'est pas correcte.</p>
<p>Il faut savoir qu'un {@choix1@} vaut {@choix_a@} et qu'un {@choix2@} vaut {@choix_b@}. Pour convertir il faut multiplier la valeur de départ {@val@} par le rapport de la puissance de départ sur celle qu'on désire donc \( \frac{@choix_a@}{@choix_b@} ={@z@} \) . la réponse attendue était {@ta1@} {@y@}</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>906296259</deployedseed>
    <deployedseed>1406497739</deployedseed>
    <deployedseed>2127103204</deployedseed>
    <deployedseed>1058425125</deployedseed>
    <deployedseed>1492037624</deployedseed>
    <deployedseed>1761702351</deployedseed>
    <deployedseed>395493478</deployedseed>
    <deployedseed>1300998652</deployedseed>
    <deployedseed>1253244899</deployedseed>
    <deployedseed>2043554874</deployedseed>
    <deployedseed>801741530</deployedseed>
    <deployedseed>577534430</deployedseed>
    <deployedseed>1121281817</deployedseed>
    <deployedseed>167076098</deployedseed>
    <deployedseed>1170963614</deployedseed>
    <deployedseed>1619724595</deployedseed>
    <deployedseed>2097354227</deployedseed>
    <deployedseed>1212610653</deployedseed>
    <deployedseed>1248607471</deployedseed>
    <deployedseed>1448596817</deployedseed>
    <deployedseed>851108648</deployedseed>
    <deployedseed>194529666</deployedseed>
    <deployedseed>746745245</deployedseed>
    <deployedseed>689871691</deployedseed>
    <deployedseed>1361759429</deployedseed>
    <deployedseed>767985140</deployedseed>
    <deployedseed>166588837</deployedseed>
    <deployedseed>69446262</deployedseed>
    <deployedseed>696530357</deployedseed>
    <deployedseed>60431331</deployedseed>
    <deployedseed>1842159654</deployedseed>
    <deployedseed>2133740303</deployedseed>
    <deployedseed>1972064721</deployedseed>
    <deployedseed>1908357876</deployedseed>
    <deployedseed>1921208829</deployedseed>
    <deployedseed>1866397206</deployedseed>
    <deployedseed>168241871</deployedseed>
    <deployedseed>2000528664</deployedseed>
    <deployedseed>813750891</deployedseed>
    <deployedseed>1439206967</deployedseed>
    <deployedseed>1440889593</deployedseed>
    <deployedseed>968640007</deployedseed>
    <deployedseed>1634970366</deployedseed>
    <deployedseed>350502389</deployedseed>
    <deployedseed>2097927956</deployedseed>
    <deployedseed>1463100035</deployedseed>
    <deployedseed>1648249259</deployedseed>
    <deployedseed>1844410973</deployedseed>
    <deployedseed>1449915914</deployedseed>
    <deployedseed>208252752</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <tags>
      <tag><text>stack</text>
</tag>
      <tag><text>conversion</text>
</tag>
      <tag><text>calcul</text>
</tag>
    </tags>
  </question>

</quiz>