<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 396129  -->
  <question type="stack">
    <name>
      <text>Q10 - Complexes : Équation du 2nd degré</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<div style="background: #059669; border-left: 5px solid #047857; border-radius: 0 8px 8px 0; padding: 10px 16px; margin-bottom: 12px; display: flex; align-items: center; gap: 10px; flex-wrap: wrap;"><strong style="font-weight: 800; color: #fff; font-size: .95rem;">Q10 — Nombres complexes</strong> <span style="background: #047857; color: #fff; padding: 2px 9px; border-radius: 20px; font-size: .78rem; font-weight: bold;">/ 1 pt</span> <span style="background: #ffffff; color: #047857; border: 1px solid #047857; padding: 2px 9px; border-radius: 20px; font-size: .75rem; font-weight: 600;">#️⃣ Équations dans C</span></div>
<p>Résoudre dans   \( \mathbf{C} \)    l'équation suivante : <br><br>\( {@z_eq@} \) <br><br></p>
<div style="background: #f0fdf4; border: 1px solid #bbf7d0; border-radius: 6px; padding: 8px 12px; font-size: 0.85rem; color: #065f46; margin-bottom: 15px;"><em>Consigne :</em> On notera \( z_1 \) la racine dont la partie imaginaire est <strong>positive</strong>, et \( z_2 \) l'autre racine.</div>
<div style="display: flex; gap: 20px; align-items: center; flex-wrap: wrap;">
<div><strong>z₁</strong> = [[input:ans_z1]] [[validation:ans_z1]]</div>
<div><strong>z₂</strong> = [[input:ans_z2]] [[validation:ans_z2]]</div>
</div>
<p>[[iframe width="100%" height="200px" scrolling="false"]] [[style]] body { margin: 0; padding: 0; font-family: system-ui, -apple-system, sans-serif; } .stack-keyboard-container { background: #f8fafc; border: 1px solid #e2e8f0; border-radius: 8px; padding: 16px; max-width: 600px; box-sizing: border-box; margin-top: 10px; } .stack-keyboard-title { font-size: 0.85rem; font-weight: 700; color: #475569; margin-bottom: 10px; text-transform: uppercase; letter-spacing: 0.05em; } .stack-keyboard-row { display: flex; flex-wrap: wrap; gap: 6px; align-items: center; margin-bottom: 8px; } .stack-keyboard-row:last-child { margin-bottom: 0; } .btn-ins { background: #ffffff !important; border: 1px solid #cbd5e1 !important; border-radius: 6px !important; color: #334155 !important; padding: 6px 12px !important; font-size: 0.9rem !important; font-weight: 600 !important; cursor: pointer !important; transition: all 0.15s ease !important; min-width: 38px !important; box-shadow: 0 1px 2px rgba(0,0,0,0.05) !important; user-select: none; } .btn-ins:hover { background: #ecfdf5 !important; border-color: #6ee7b7 !important; color: #059669 !important; } .btn-ins:active { background: #d1fae5 !important; } .btn-i { background: #ecfdf5 !important; border-color: #a7f3d0 !important; color: #047857 !important; font-style: italic; font-weight: 800 !important; } .target-btn { background: #fff !important; border: 2px solid #cbd5e1 !important; border-radius: 6px !important; padding: 4px 10px !important; font-size: 0.8rem !important; font-weight: 700 !important; color: #64748b !important; cursor: pointer !important; transition: all 0.2s ease !important; } .active-target { background: #059669 !important; color: #ffffff !important; border-color: #059669 !important; box-shadow: 0 2px 4px rgba(5, 150, 105, 0.3) !important; } .sep { width: 1px; height: 28px; background: #e2e8f0; margin: 0 4px; } [[/style]]</p>
<div class="stack-keyboard-container">
<div class="stack-keyboard-title">🧮 Clavier cible</div>
<div class="stack-keyboard-row"><button id="target_z1" class="target-btn active-target" type="button">Saisir pour z₁</button> <button id="target_z2" class="target-btn" type="button">Saisir pour z₂</button></div>
<div class="stack-keyboard-row"><button class="btn-ins btn-i" type="button" data-val="%i">i</button>
<div class="sep"> </div>
<button class="btn-ins" type="button" data-val="+">+</button> <button class="btn-ins" type="button" data-val="-">-</button> <button class="btn-ins" type="button" data-val="*">×</button> <button class="btn-ins" type="button" data-val="/">/</button>
<div class="sep"> </div>
<button class="btn-ins" type="button" data-val="(">(</button> <button class="btn-ins" type="button" data-val=")">)</button></div>
</div>
<p>[[script type="module"]] import {stack_js} from '[[cors src="stackjsiframe.js"/]]'; let currentTarget = null; function setupInput(name, btnId) { stack_js.request_access_to_input(name, true).then(function(input_id) { var input = document.getElementById(input_id); document.getElementById(btnId).addEventListener('click', function(e) { e.preventDefault(); currentTarget = input; document.querySelectorAll('.target-btn').forEach(b =&gt; b.classList.remove('active-target')); this.classList.add('active-target'); }); if(!currentTarget) { currentTarget = input; } }); } setupInput("ans_z1", "target_z1"); setupInput("ans_z2", "target_z2"); document.querySelectorAll('.btn-ins').forEach(function(btn){ btn.addEventListener('click', function(e){ e.preventDefault(); if(!currentTarget) return; var val = this.getAttribute('data-val'); var actuel = currentTarget.value || ''; var texteAvant = actuel.trimEnd(); var charPrecedent = texteAvant.length &gt; 0 ? texteAvant.slice(-1) : null; var excl = [null, '=', '+', '-', '*', '/', '^', '(']; var estOp = /^[+\-*\/^)]/.test(val); var pref = ''; if (!excl.includes(charPrecedent) &amp;&amp; !estOp &amp;&amp; actuel.length &gt; 0) { pref = '*'; } currentTarget.value = actuel + pref + val; currentTarget.dispatchEvent(new Event('change')); }); }); [[/script]] [[/iframe]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<div style="margin-top: 20px; padding: 15px; background: #f8fafc; border: 1px solid #cbd5e1; border-radius: 8px;">
<div style="font-weight: bold; margin-bottom: 10px;">🔑 Correction détaillée</div>
<div style="font-size: .9rem;">On calcule le discriminant :<br>\( \Delta = b^2 - 4ac = (-2)^2 - 4 \times 1 \times 5 = 4 - 20 = -16 \)<br><br>Puisque \(\Delta &lt; 0\), l'équation admet deux racines complexes conjuguées :<br>\( z = \frac{-b \pm i\sqrt{-\Delta}}{2a} = \frac{2 \pm i\sqrt{16}}{2} = \frac{2 \pm 4i}{2} \)<br><br>En séparant les parties réelles et imaginaires, on obtient :<br>• <strong>z_1 = 1 + 2i</strong> (partie imaginaire positive)<br>• <strong>z_2 = 1 - 2i</strong></div>
</div>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[ta1: 1 + 2*%i; ta2: 1 - 2*%i; texput(z_eq, "z^2 - 2z + 5 = 0");]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<ol>
<li>[[feedback:prt1]]</li>
</ol>]]></text>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>Racines de z^2-2z+5=0 : {@ta1@} et {@ta2@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p>Généré avec HéStack V2 | hestack-types:complex</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans_z1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans_z2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>z1 est-il la racine à Im positif ?</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_z1</sans>
        <tans>ta1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>1</truenextnode>
        <trueanswernote>PRT-Z1-OK</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>PRT-Z1-NOK</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>z1 ok, z2 est-il la racine à Im négatif ?</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_z2</sans>
        <tans>ta2</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-Z2-OK</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #15803d; padding: 10px 14px; background: #f0fdf4; border-radius: 4px;">✅ <strong>Parfait !</strong> Les deux racines sont correctes et dans le bon ordre. Vous avez bien appliqué la formule avec le discriminant négatif.</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.5</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>PRT-Z2-NOK</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #eab308; padding: 10px 14px; background: #fefce8; border-radius: 4px;">⚠️ <strong>Partiellement correct.</strong> La première racine \( z_1 \) est bonne, mais \( z_2 \) est fausse. Vérifiez le signe de la partie imaginaire pour la deuxième racine.</div>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>z1 n'est pas le bon, a-t-il inversé les racines ?</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_z1</sans>
        <tans>ta2</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>3</truenextnode>
        <trueanswernote>PRT-SWAP-Z1</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>PRT-ALL-WRONG</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #dc2626; padding: 10px 14px; background: #fff0f0; border-radius: 4px;">❌ <strong>Incorrect.</strong> Aucune des deux racines n'est correcte. N'oubliez pas que si \(\Delta &lt; 0\), les racines sont \( z = \frac{-b \pm i\sqrt{-\Delta}}{2a} \).</div>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Il a inversé, z2 est-il donc z1 ?</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_z2</sans>
        <tans>ta1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-SWAP-TOTAL</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🔄 <strong>Racines inversées !</strong> Mathématiquement, vos calculs sont justes, mais vous avez inversé \( z_1 \) et \( z_2 \). Relisez la consigne : \( z_1 \) doit avoir une partie imaginaire <strong>positive</strong>.</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.25</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>PRT-SWAP-PARTIAL</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #eab308; padding: 10px 14px; background: #fefce8; border-radius: 4px;">⚠️ <strong>Erreur d'ordre ou de calcul.</strong> Il semble que vous ayez placé la racine \( 1-2i \) dans la case \( z_1 \), mais la deuxième racine est fausse.</div>]]></text>
        </falsefeedback>
      </node>
    </prt>
  </question>

</quiz>