<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 396122  -->
  <question type="stack">
    <name>
      <text>Complexes : Forme algébrique (Arbre détaillé)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<div style="background: #06b6d4; border-left: 5px solid #0891b2; border-radius: 0 8px 8px 0; padding: 10px 16px; margin-bottom: 12px; display: flex; align-items: center; gap: 10px; flex-wrap: wrap;"><strong style="font-weight: 800; color: #fff; font-size: .95rem;">Q7 — Nombres complexes</strong> <span style="background: #0891b2; color: #fff; padding: 2px 9px; border-radius: 20px; font-size: .78rem; font-weight: bold;">/ 1 pt</span> <span style="background: #ffffff; color: #0891b2; border: 1px solid #0891b2; padding: 2px 9px; border-radius: 20px; font-size: .75rem; font-weight: 600;">#️⃣ Forme algébrique</span></div>
<p>Déterminer la forme algébrique \( a + bi \) (avec \( a, b \in \mathbb{R} \)) de l'expression suivante : <br><br>\( {@z@} \)</p>
<p>[[input:ans1]] [[validation:ans1]]</p>
<p>[[iframe width="100%" height="160px" scrolling="false"]] [[style]] body { margin: 0; padding: 0; font-family: system-ui, -apple-system, sans-serif; } .stack-keyboard-container { background: #f8fafc; border: 1px solid #e2e8f0; border-radius: 8px; padding: 16px; max-width: 600px; box-sizing: border-box; margin-top: 10px; } .stack-keyboard-title { font-size: 0.85rem; font-weight: 700; color: #475569; margin-bottom: 10px; text-transform: uppercase; letter-spacing: 0.05em; } .stack-keyboard-row { display: flex; flex-wrap: wrap; gap: 6px; align-items: center; } .btn-ins { background: #ffffff !important; border: 1px solid #cbd5e1 !important; border-radius: 6px !important; color: #334155 !important; padding: 6px 12px !important; font-size: 0.9rem !important; font-weight: 600 !important; cursor: pointer !important; transition: all 0.15s ease !important; min-width: 38px !important; box-shadow: 0 1px 2px rgba(0,0,0,0.05) !important; user-select: none; } .btn-ins:hover { background: #ecfeff !important; border-color: #67e8f9 !important; color: #0891b2 !important; } .btn-ins:active { background: #cffafe !important; } .btn-var { background: #f1f5f9 !important; border-color: #94a3b8 !important; color: #0f172a !important; font-style: italic; } .btn-i { background: #ecfeff !important; border-color: #a5f3fc !important; color: #0e7490 !important; font-style: italic; font-weight: 800 !important; } .sep { width: 1px; height: 28px; background: #e2e8f0; margin: 0 4px; } [[/style]]</p>
<div class="stack-keyboard-container">
<div class="stack-keyboard-title">🧮 Aide à la saisie</div>
<div class="stack-keyboard-row"><button class="btn-ins btn-i" type="button" data-val="%i">i</button>
<div class="sep"> </div>
<button class="btn-ins" type="button" data-val="+">+</button> <button class="btn-ins" type="button" data-val="-">-</button> <button class="btn-ins" type="button" data-val="*">×</button> <button class="btn-ins" type="button" data-val="/">/</button>
<div class="sep"> </div>
<button class="btn-ins" type="button" data-val="(">(</button> <button class="btn-ins" type="button" data-val=")">)</button></div>
</div>
<p>[[script type="module"]] import {stack_js} from '[[cors src="stackjsiframe.js"/]]'; stack_js.request_access_to_input("ans1", true).then(function(input_id) { var stackInput = document.getElementById(input_id); document.querySelectorAll('.btn-ins').forEach(function(btn){ btn.addEventListener('click', function(e){ e.preventDefault(); var val = this.getAttribute('data-val'); var actuel = stackInput.value || ''; var texteAvant = actuel.trimEnd(); var charPrecedent = texteAvant.length &gt; 0 ? texteAvant.slice(-1) : null; var excl = [null, '=', '+', '-', '*', '/', '^', '(']; var estOp = /^[+\-*\/^)]/.test(val); var pref = ''; if (!excl.includes(charPrecedent) &amp;&amp; !estOp &amp;&amp; actuel.length &gt; 0) { pref = '*'; } stackInput.value = actuel + pref + val; stackInput.dispatchEvent(new Event('change')); }); }); }); [[/script]] [[/iframe]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<div style="margin-top: 20px; padding: 15px; background: #f8fafc; border: 1px solid #cbd5e1; border-radius: 8px;">
<div style="font-weight: bold; margin-bottom: 10px;">🔑 Correction détaillée</div>
<div style="font-size: .9rem;">Pour passer à la forme algébrique, on multiplie par le <strong>conjugué</strong> du dénominateur :<br><br>\( z = \frac{3 + 2i}{1 - i} \times \frac{1 + i}{1 + i} \)<br><br><strong>Au dénominateur (identité remarquable) :</strong><br>\( (1-i)(1+i) = 1^2 - i^2 = 1 - (-1) = 2 \)<br><br><strong>Au numérateur (distribution) :</strong><br>\( (3+2i)(1+i) = 3 + 3i + 2i + 2i^2 \)<br>\( = 3 + 5i + 2(-1) = 1 + 5i \)<br><br><strong>Résultat final :</strong><br>\( z = \frac{1 + 5i}{2} = \frac{1}{2} + \frac{5}{2}i \)</div>
</div>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* 1. La vraie réponse mathématique */
ta1: rectform((3+2*%i)/(1-%i));

/* 2. Les erreurs typiques pour l'arbre de décision */
ta_oubli_div: 1 + 5*%i;
taErreurI2: 5/2 + 5/2*%i;

/* 3. L'affichage forcé de z avec texput pour empêcher la simplification d'affichage */
texput(z, "\\dfrac{3 + 2*i}{1 - i}");]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<ol>
<li>[[feedback:prt1]]</li>
</ol>]]></text>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>Forme algébrique de (3+2i)/(1-i) : {@ta1@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p>Généré avec HéStack V2 | hestack-types:complex</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>20</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>2</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Vérifie si la réponse est exactement 1/2 + 5/2*i</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-COMPLEX-OK</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #15803d; padding: 10px 14px; background: #f0fdf4; border-radius: 4px;">✅ <strong>Parfait !</strong> La forme algébrique est parfaitement simplifiée.</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>PRT-NOT-PERFECT</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Vérifie si l'élève a trouvé 1+5i (a oublié de diviser par 2)</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta_oubli_div</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-DIV</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #eab308; padding: 10px 14px; background: #fefce8; border-radius: 4px;">⚠️ <strong>Presque juste !</strong> Vous avez correctement calculé le numérateur (\(1+5i\)), mais <strong>vous avez oublié de diviser par le dénominateur</strong>. Rappelez-vous que \((1-i)(1+i) = 2\). Il faut donc diviser votre résultat par 2.</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>PRT-NOT-ERR-DIV</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Vérifie si l'élève a calculé avec i^2=1 au lieu de i^2=-1 (donne 5/2 + 5/2*i)</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>taErreurI2</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.25</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-I2</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Erreur de cours !</strong> Il semble que vous ayez utilisé \( i^2 = 1 \). C'est <strong>FAUX</strong>. La définition du nombre complexe impose que \( i^2 = -1 \). Recalculez votre numérateur en corrigeant cette erreur.</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>PRT-ERR-UNKNOWN</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #dc2626; padding: 10px 14px; background: #fff0f0; border-radius: 4px;">❌ <strong>Incorrect.</strong> La méthode est : multiplier par le conjugué, développer (attention à \(i^2=-1\)), puis diviser par le dénominateur. <br><em>N'oubliez pas d'utiliser le bouton "<strong>i</strong>" du clavier pour que la machine reconnaisse le symbole.</em></div>]]></text>
        </falsefeedback>
      </node>
    </prt>
  </question>

</quiz>