<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 396142  -->
  <question type="stack">
    <name>
      <text>Q12 - Complexes : Géométrie et Affixes</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<div style="background: #475569; border-left: 5px solid #334155; border-radius: 0 8px 8px 0; padding: 10px 16px; margin-bottom: 12px; display: flex; align-items: center; gap: 10px; flex-wrap: wrap;"><strong style="font-weight: 800; color: #fff; font-size: .95rem;">Q12 — Nombres complexes</strong> <span style="background: #334155; color: #fff; padding: 2px 9px; border-radius: 20px; font-size: .78rem; font-weight: bold;">/ 1 pt</span> <span style="background: #ffffff; color: #334155; border: 1px solid #334155; padding: 2px 9px; border-radius: 20px; font-size: .75rem; font-weight: 600;">#️⃣ Géométrie plane</span></div>
<p>Dans un repère orthonormé direct \((O; \vec{u}, \vec{v})\), on considère les points \(A\) et \(B\) d'affixes respectives : <br><br>\( z_A = {@za_disp@} \) et \( z_B = {@zb_disp@} \) <br><br>Calculer l'affixe du milieu \(I\) du segment \([AB]\) et la longueur \(AB\). <br><br></p>
<div style="display: flex; gap: 20px; align-items: center; flex-wrap: wrap;">
<div><strong>z_I</strong> = [[input:ans_zI]] [[validation:ans_zI]]</div>
<div><strong>AB</strong> = [[input:ans_AB]] [[validation:ans_AB]]</div>
</div>
<p>[[iframe width="100%" height="200px" scrolling="false"]] [[style]] body { margin: 0; padding: 0; font-family: system-ui, -apple-system, sans-serif; } .stack-keyboard-container { background: #f8fafc; border: 1px solid #e2e8f0; border-radius: 8px; padding: 16px; max-width: 600px; box-sizing: border-box; margin-top: 10px; } .stack-keyboard-title { font-size: 0.85rem; font-weight: 700; color: #475569; margin-bottom: 10px; text-transform: uppercase; letter-spacing: 0.05em; } .stack-keyboard-row { display: flex; flex-wrap: wrap; gap: 6px; align-items: center; margin-bottom: 8px; } .stack-keyboard-row:last-child { margin-bottom: 0; } .btn-ins { background: #ffffff !important; border: 1px solid #cbd5e1 !important; border-radius: 6px !important; color: #334155 !important; padding: 6px 12px !important; font-size: 0.9rem !important; font-weight: 600 !important; cursor: pointer !important; transition: all 0.15s ease !important; min-width: 38px !important; box-shadow: 0 1px 2px rgba(0,0,0,0.05) !important; user-select: none; } .btn-ins:hover { background: #f1f5f9 !important; border-color: #94a3b8 !important; color: #0f172a !important; } .btn-ins:active { background: #e2e8f0 !important; } .btn-i { background: #f1f5f9 !important; border-color: #94a3b8 !important; color: #0f172a !important; font-style: italic; font-weight: 800 !important; } .target-btn { background: #fff !important; border: 2px solid #cbd5e1 !important; border-radius: 6px !important; padding: 4px 10px !important; font-size: 0.8rem !important; font-weight: 700 !important; color: #64748b !important; cursor: pointer !important; transition: all 0.2s ease !important; } .active-target { background: #475569 !important; color: #ffffff !important; border-color: #475569 !important; box-shadow: 0 2px 4px rgba(71, 85, 105, 0.3) !important; } .sep { width: 1px; height: 28px; background: #e2e8f0; margin: 0 4px; } [[/style]]</p>
<div class="stack-keyboard-container">
<div class="stack-keyboard-title">🧮 Clavier cible</div>
<div class="stack-keyboard-row"><button id="target_zI" class="target-btn active-target" type="button">Saisir pour z_I</button> <button id="target_AB" class="target-btn" type="button">Saisir pour AB</button></div>
<div class="stack-keyboard-row"><button class="btn-ins btn-i" type="button" data-val="%i">i</button> <button class="btn-ins" type="button" data-val="sqrt()">√</button>
<div class="sep"> </div>
<button class="btn-ins" type="button" data-val="+">+</button> <button class="btn-ins" type="button" data-val="-">-</button> <button class="btn-ins" type="button" data-val="*">×</button> <button class="btn-ins" type="button" data-val="/">/</button>
<div class="sep"> </div>
<button class="btn-ins" type="button" data-val="(">(</button> <button class="btn-ins" type="button" data-val=")">)</button></div>
</div>
<p>[[script type="module"]] import {stack_js} from '[[cors src="stackjsiframe.js"/]]'; let currentTarget = null; function setupInput(name, btnId) { stack_js.request_access_to_input(name, true).then(function(input_id) { var input = document.getElementById(input_id); document.getElementById(btnId).addEventListener('click', function(e) { e.preventDefault(); currentTarget = input; document.querySelectorAll('.target-btn').forEach(b =&gt; b.classList.remove('active-target')); this.classList.add('active-target'); }); if(!currentTarget) { currentTarget = input; } }); } setupInput("ans_zI", "target_zI"); setupInput("ans_AB", "target_AB"); document.querySelectorAll('.btn-ins').forEach(function(btn){ btn.addEventListener('click', function(e){ e.preventDefault(); if(!currentTarget) return; var val = this.getAttribute('data-val'); var actuel = currentTarget.value || ''; var texteAvant = actuel.trimEnd(); var charPrecedent = texteAvant.length &gt; 0 ? texteAvant.slice(-1) : null; var excl = [null, '=', '+', '-', '*', '/', '^', '(']; var estOp = /^[+\-*\/^)]/.test(val); var pref = ''; if (!excl.includes(charPrecedent) &amp;&amp; !estOp &amp;&amp; actuel.length &gt; 0) { pref = '*'; } currentTarget.value = actuel + pref + val; currentTarget.dispatchEvent(new Event('change')); }); }); [[/script]] [[/iframe]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<div style="margin-top: 20px; padding: 15px; background: #f8fafc; border: 1px solid #cbd5e1; border-radius: 8px;">
<div style="font-weight: bold; margin-bottom: 10px;">🔑 Correction détaillée</div>
<div style="font-size: .9rem;"><strong>1. Affixe du milieu I :</strong><br>La formule est \( z_I = \frac{z_A + z_B}{2} \).<br>\( z_I = \frac{(3+i) + (-1+3i)}{2} = \frac{2 + 4i}{2} = 1 + 2i \)<br><br><strong>2. Longueur AB :</strong><br>La distance correspond au module du vecteur \(\vec{AB}\), donc \( AB = |z_B - z_A| \).<br>\( z_B - z_A = (-1+3i) - (3+i) = -4 + 2i \)<br>\( AB = |-4 + 2i| = \sqrt{(-4)^2 + 2^2} = \sqrt{16 + 4} = \sqrt{20} = 2\sqrt{5} \)</div>
</div>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>za: 3+%i;
zb: -1+3*%i;
ta_zI: (za+zb)/2;
ta_AB: abs(zb-za);
ta_err_mid: za+zb;
ta_err_dist: abs(za)+abs(zb);
za_disp: 3+%i;
zb_disp: -1+3*%i;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<ol>
<li>[[feedback:prt1]]</li>
</ol>]]></text>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>Affixe milieu et distance AB : zI={@ta_zI@}, AB={@ta_AB@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p>Généré avec HéStack V2 | hestack-types:complex</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans_AB</name>
      <type>algebraic</type>
      <tans>ta_AB</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans_zI</name>
      <type>algebraic</type>
      <tans>ta_zI</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_zI</sans>
        <tans>ta_zI</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>1</truenextnode>
        <trueanswernote>PRT-ZI-OK</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>PRT-ZI-NOK</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_AB</sans>
        <tans>ta_AB</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-AB-OK</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #15803d; padding: 10px 14px; background: #f0fdf4; border-radius: 4px;">✅ <strong>Parfait !</strong> L'affixe du milieu et la longueur sont correctes. Vous maîtrisez le lien entre géométrie et complexes.</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.5</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>PRT-AB-NOK</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #eab308; padding: 10px 14px; background: #fefce8; border-radius: 4px;">⚠️ <strong>Partiellement correct.</strong> L'affixe du milieu est juste, mais la distance est fausse. Rappelez-vous que \( AB = |z_B - z_A| \) et non la somme des modules.</div>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_zI</sans>
        <tans>ta_err_mid</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.25</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-MID</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Oubli de la division !</strong> Vous avez calculé \( z_A + z_B \) mais vous avez oublié de diviser par 2 pour trouver le milieu.</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>PRT-ALL-WRONG</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #dc2626; padding: 10px 14px; background: #fff0f0; border-radius: 4px;">❌ <strong>Incorrect.</strong> L'affixe du milieu est faux. Formule : \( z_I = \frac{z_A + z_B}{2} \).</div>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_AB</sans>
        <tans>ta_err_dist</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.25</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-DIST</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Mauvaise formule de distance !</strong> Vous avez calculé \( |z_A| + |z_B| \). En géométrie, les longueurs ne s'additionnent pas ainsi. Il faut utiliser le module du vecteur : \( AB = |z_B - z_A| \).</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>PRT-ERR-OTHER</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #dc2626; padding: 10px 14px; background: #fff0f0; border-radius: 4px;">❌ <strong>Incorrect.</strong> Les deux valeurs sont fausses. Revoyez les formules : Milieu (\(\frac{z_1+z_2}{2}\)) et Distance (\(|z_2-z_1|\)).</div>]]></text>
        </falsefeedback>
      </node>
    </prt>
  </question>

</quiz>