<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 396124  -->
  <question type="stack">
    <name>
      <text>Q9 - Complexes : Module et Argument</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<div style="background:#7c3aed;border-left:5px solid #6d28d9;border-radius:0 8px 8px 0;padding:10px 16px;margin-bottom:12px;display:flex;align-items:center;gap:10px;flex-wrap:wrap;">
        <strong style="font-weight:800;color:#fff;font-size:.95rem;">Q9 — Nombres complexes</strong>
        <span style="background:#6d28d9;color:#fff;padding:2px 9px;border-radius:20px;font-size:.78rem;font-weight:700;">/ 1 pt</span>
        <span style="background:#ffffff;color:#6d28d9;border:1px solid #6d28d9;padding:2px 9px;border-radius:20px;font-size:.75rem;font-weight:600;">#️⃣ Module & Argument</span>
      </div>
      
      Soit le nombre complexe \( z = {@z@} \). <br/><br/>
      Déterminer le module \( |z| \) et l'argument principal \( \theta \) (en radians) du nombre \( z \).
      <br/><br/>
      <div style="display:flex; gap:20px; align-items:center; flex-wrap:wrap;">
        <div><strong>|z|</strong> = [[input:ans_mod]] [[validation:ans_mod]]</div>
        <div><strong>Arg(z)</strong> = [[input:ans_arg]] [[validation:ans_arg]]</div>
      </div>

[[iframe width="100%" height="200px" scrolling="false"]]
[[style]]
  body { margin: 0; padding: 0; font-family: system-ui, -apple-system, sans-serif; }
  .stack-keyboard-container { background: #f8fafc; border: 1px solid #e2e8f0; border-radius: 8px; padding: 16px; max-width: 600px; box-sizing: border-box; margin-top: 10px; }
  .stack-keyboard-title { font-size: 0.85rem; font-weight: 700; color: #475569; margin-bottom: 10px; text-transform: uppercase; letter-spacing: 0.05em; }
  .stack-keyboard-row { display: flex; flex-wrap: wrap; gap: 6px; align-items: center; margin-bottom: 8px; }
  .stack-keyboard-row:last-child { margin-bottom: 0; }
  .btn-ins { background: #ffffff !important; border: 1px solid #cbd5e1 !important; border-radius: 6px !important; color: #334155 !important; padding: 6px 12px !important; font-size: 0.9rem !important; font-weight: 600 !important; cursor: pointer !important; transition: all 0.15s ease !important; min-width: 38px !important; box-shadow: 0 1px 2px rgba(0,0,0,0.05) !important; user-select: none; }
  .btn-ins:hover { background: #f5f3ff !important; border-color: #c4b5fd !important; color: #7c3aed !important; }
  .btn-ins:active { background: #ede9fe !important; }
  .btn-pi { color: #7c3aed !important; font-weight: 800 !important; font-style: italic; }
  .target-btn { background: #fff !important; border: 2px solid #cbd5e1 !important; border-radius: 6px !important; padding: 4px 10px !important; font-size: 0.8rem !important; font-weight: 700 !important; color: #64748b !important; cursor: pointer !important; transition: all 0.2s ease !important; }
  .active-target { background: #7c3aed !important; color: #ffffff !important; border-color: #7c3aed !important; box-shadow: 0 2px 4px rgba(124, 58, 237, 0.3) !important; }
  .sep { width: 1px; height: 28px; background: #e2e8f0; margin: 0 4px; }
[[/style]]

<div class="stack-keyboard-container">
  <div class="stack-keyboard-title">🧮 Clavier cible</div>
  <div class="stack-keyboard-row">
    <button id="target_mod" class="target-btn active-target" type="button">Saisir pour |z|</button>
    <button id="target_arg" class="target-btn" type="button">Saisir pour Arg(z)</button>
  </div>
  <div class="stack-keyboard-row">
    <button class="btn-ins" type="button" data-val="sqrt()">√</button>
    <button class="btn-ins btn-pi" type="button" data-val="%pi">π</button>
    <div class="sep"></div>
    <button class="btn-ins" type="button" data-val="+">+</button>
    <button class="btn-ins" type="button" data-val="-">-</button>
    <button class="btn-ins" type="button" data-val="*">×</button>
    <button class="btn-ins" type="button" data-val="/">/</button>
    <div class="sep"></div>
    <button class="btn-ins" type="button" data-val="(">(</button>
    <button class="btn-ins" type="button" data-val=")">)</button>
  </div>
</div>

[[script type="module"]]
import {stack_js} from '[[cors src="stackjsiframe.js"/]]';
let currentTarget = null;
function setupInput(name, btnId) {
    stack_js.request_access_to_input(name, true).then(function(input_id) {
        var input = document.getElementById(input_id);
        document.getElementById(btnId).addEventListener('click', function(e) {
            e.preventDefault();
            currentTarget = input;
            document.querySelectorAll('.target-btn').forEach(b => b.classList.remove('active-target'));
            this.classList.add('active-target');
        });
        if(!currentTarget) { currentTarget = input; }
    });
}
setupInput("ans_mod", "target_mod");
setupInput("ans_arg", "target_arg");
document.querySelectorAll('.btn-ins').forEach(function(btn){
    btn.addEventListener('click', function(e){
        e.preventDefault();
        if(!currentTarget) return;
        var val = this.getAttribute('data-val');
        var actuel = currentTarget.value || '';
        var texteAvant = actuel.trimEnd();
        var charPrecedent = texteAvant.length > 0 ? texteAvant.slice(-1) : null;
        var excl = [null, '=', '+', '-', '*', '/', '^', '('];
        var estOp = /^[+\-*\/^)]/.test(val);
        var pref = '';
        if (!excl.includes(charPrecedent) && !estOp && actuel.length > 0) { pref = '*'; }
        currentTarget.value = actuel + pref + val;
        currentTarget.dispatchEvent(new Event('change'));
    });
});
[[/script]]
[[/iframe]]]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<div style="margin-top:20px; padding:15px; background:#f8fafc; border:1px solid #cbd5e1; border-radius:8px;"><div style="font-weight:bold; margin-bottom:10px;">🔑 Correction détaillée</div><div style="font-size:.9rem;"><strong>1. Calcul du module :</strong><br/>\( |z| = \sqrt{(-\sqrt{3})^2 + 1^2} = \sqrt{3 + 1} = \sqrt{4} = 2 \)<br/><br/><strong>2. Calcul de l'argument :</strong><br/>On utilise la formule : \( \cos(\theta) = \frac{x}{|z|} \) et \( \sin(\theta) = \frac{y}{|z|} \)<br/>\( \cos(\theta) = \frac{-\sqrt{3}}{2} \) et \( \sin(\theta) = \frac{1}{2} \)<br/>L'angle qui a un cosinus négatif et un sinus positif est dans le <strong>deuxième quadrant</strong>.<br/>L'angle de référence est \( \frac{\pi}{6} \).<br/>Donc l'argument principal est : \( \theta = \pi - \frac{\pi}{6} = \frac{5\pi}{6} \).</div></div>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text></text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[z_calc: -sqrt(3) + %i; ta_mod: cabs(z_calc); ta_arg: carg(z_calc); ta_arg_erreur_quad: -%pi/6; texput(z, "-\\sqrt{3} + i");]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<ol><li>[[feedback:prt1]]</li></ol>]]></text>
    </specificfeedback>
    <questionnote format="html">
      <text>Module et Argument de -sqrt(3)+i</text>
    </questionnote>
    <questiondescription format="html">
      <text>Généré avec HéStack V2 | hestack-types:complex</text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans_arg</name>
      <type>algebraic</type>
      <tans>ta_arg</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans_mod</name>
      <type>algebraic</type>
      <tans>ta_mod</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_mod</sans>
        <tans>ta_mod</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>1</truenextnode>
        <trueanswernote>PRT-MOD-OK</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>PRT-MOD-WRONG</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_arg</sans>
        <tans>ta_arg</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ARG-OK</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left:4px solid #15803d;padding:10px 14px;background:#f0fdf4;border-radius:4px;">✅ <strong>Parfait !</strong> Le module et l'argument principal sont corrects. Vous avez bien identifié le quadrant.</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.5</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>PRT-ARG-WRONG</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_arg</sans>
        <tans>ta_arg</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-MOD-WRONG-ARG-OK</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left:4px solid #eab308;padding:10px 14px;background:#fefce8;border-radius:4px;">⚠️ <strong>Partiellement correct.</strong> L'argument est bon, mais le module est faux. Formule : \( |z| = \sqrt{x^2 + y^2} \).</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>PRT-ALL-WRONG</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<div style="border-left:4px solid #dc2626;padding:10px 14px;background:#fff0f0;border-radius:4px;">❌ <strong>Incorrect.</strong> Le module et l'argument sont faux. Revoyez les formules de base et la position du point dans le plan complexe.</div>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_arg</sans>
        <tans>ta_arg_erreur_quad</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.25</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-QUADRANT</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left:4px solid #f97316;padding:10px 14px;background:#fff7ed;border-radius:4px;">🚨 <strong>Erreur de quadrant !</strong> Votre valeur \( -\frac{\pi}{6} \) correspond à l'angle de référence, mais le point \( z \) est dans le <strong>deuxième quadrant</strong> (partie réelle négative, partie imaginaire positive). L'argument principal doit être positif !</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>PRT-ARG-OTHER-ERR</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<div style="border-left:4px solid #eab308;padding:10px 14px;background:#fefce8;border-radius:4px;">⚠️ <strong>Module correct, mais argument incorrect.</strong> Utilisez le cercle trigonométrique pour trouver le bon angle en fonction du signe de x et y. N'oubliez pas le bouton "<strong>π</strong>" du clavier.</div>]]></text>
        </falsefeedback>
      </node>
    </prt>
  </question>

</quiz>