<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!--
  Question STACK générée par le générateur "Cinématique du point — trajectoire composée".
  Piste : ligne(3m,20°) -> arc(R=5m,30°)
  Valeurs fixes (pas de randomisation).
-->
<question type="stack">
  <name>
    <text>Cinematique_trajectoire_composee</text>
  </name>
  <questiontext format="html">
    <text><![CDATA[
<p>La figure ci-dessous montre la piste construite pour cet exercice. Le point M est
photographié à intervalles de temps réguliers \( \Delta t = {@dt@}\ \text{s} \), en
partant de son point de départ M0, y compris pendant qu'il est guidé par la piste. Une fois
la piste quittée (au point M{@kExit@}), une seule force agit sur lui : son poids.</p>

<p>On admet que la vitesse instantanée en un point M<sub>k</sub> peut être estimée par
\( \vec v_k \approx \dfrac{\overrightarrow{M_{k-1}M_{k+1}}}{2\,\Delta t} \) — valable
seulement en vol libre (accélération constante), pas pendant que le point est guidé par la
piste.</p>

<p><b>Étape 1</b> — sur la figure, fais glisser les points colorés pour construire, à
l'échelle indiquée ({@kv@} carreau(x) pour 1 m/s) :</p>
<ul>
<li>le vecteur \( \vec v_2 \) (rouge), <b>à partir de M{@kExit+2@}</b> ;</li>
<li>le vecteur \( \vec v_3 \) (bleu), <b>à partir de M{@kExit+3@}</b>.</li>
</ul>

[[jsxgraph width="600" height="440" id="board_traj"]]
var b1 = JXG.JSXGraph.initBoard('board_traj', {
  boundingbox: [{@xmin_t@}, {@ymax_t@}, {@xmax_t@}, {@ymin_t@}],
  axis: true, showCopyright: false, showNavigation: false, keepAspectRatio: true
});
b1.create('text', [{@xmin_t@}+0.3, {@ymin_t@}+0.4, 'x (m)'], {fontSize:12});
b1.create('text', [{@xmin_t@}+0.3, {@ymax_t@}-0.3, 'y (m)'], {fontSize:12});

b1.create('segment', [[0,0],[2.8191,1.0261]], {strokeColor:'#8a5a2b', strokeWidth:5, fixed:true});
b1.create('curve', [[2.8191,2.9212,3.0224,3.1227,3.2221,3.3204,3.4177,3.5139,3.609,3.7028,3.7955,3.8868,3.9769,4.0655,4.1528,4.2386,4.3229,4.4057,4.4869,4.5665,4.6445,4.7208,4.7954,4.8682,4.9392],[1.0261,1.0645,1.1051,1.148,1.193,1.2402,1.2895,1.3409,1.3944,1.45,1.5076,1.5672,1.6288,1.6923,1.7578,1.8251,1.8943,1.9653,2.0381,2.1127,2.189,2.267,2.3466,2.4278,2.5106]], {strokeColor:'#8a5a2b', strokeWidth:5, fixed:true});

var Mlist = {@Mlist@};
var kExit = {@kExit@};
var Mpts = [];
for (var i=0; i<Mlist.length; i++){
  Mpts.push(b1.create('point', Mlist[i], {name:'M'+i, fixed:true,
    size: i<kExit ? 2.5 : 3, color: i<kExit ? '#8a5a2b' : '#0033aa',
    label:{offset:[6,6], fontSize: i<kExit ? 10 : 13}}));
}

var kv = {@kv@};
var i2 = kExit+2, i3 = kExit+3;
var T2 = b1.create('point', [Mlist[i2][0], Mlist[i2][1]], {name:'v2', size:4, color:'#c0392b',
  snapToGrid:true, snapSizeX:0.05, snapSizeY:0.05});
b1.create('arrow', [Mpts[i2], T2], {strokeColor:'#c0392b', strokeWidth:2});

var T3 = b1.create('point', [Mlist[i3][0], Mlist[i3][1]], {name:'v3', size:4, color:'#2980b9',
  snapToGrid:true, snapSizeX:0.05, snapSizeY:0.05});
b1.create('arrow', [Mpts[i3], T3], {strokeColor:'#2980b9', strokeWidth:2});
[[/jsxgraph]]

<p><b>Étape 2</b> — dans le repère ci-dessous (1 carreau = 1 m/s), reporte les vecteurs
\( \vec v_2 \) et \( \vec v_3 \) que tu viens de construire à partir d'une même origine O,
puis construis le vecteur variation de vitesse
\( \Delta \vec v = \vec v_3 - \vec v_2 \) (vert) en faisant glisser sa pointe.</p>

[[jsxgraph width="380" height="380" id="board_dv"]]
var b2 = JXG.JSXGraph.initBoard('board_dv', {
  boundingbox: [-{@vmax2@}, {@vmax2@}, {@vmax2@}, -{@vmax2@}],
  axis: true, showCopyright: false, showNavigation: false, keepAspectRatio: true, grid:true
});
var O = b2.create('point', [0, 0], {name:'O', fixed:true, size:2, color:'black'});
var Tdv = b2.create('point', [0, 0], {name:'\u0394v', size:4, color:'#27ae60',
  snapToGrid:true, snapSizeX:0.05, snapSizeY:0.05});
b2.create('arrow', [O, Tdv], {strokeColor:'#27ae60', strokeWidth:2});

function syncInput(pt, origin, suffix, scale){
  var dx = ((pt.X()-origin.X())/scale).toFixed(2);
  var dy = ((pt.Y()-origin.Y())/scale).toFixed(2);
  var el = document.querySelector('input[id$="' + suffix + '"]');
  if (el){
    el.value = '[' + dx + ',' + dy + ']';
    el.dispatchEvent(new Event('input', {bubbles:true}));
    el.dispatchEvent(new Event('change', {bubbles:true}));
  }
}
T2.on('up', function(){ syncInput(T2, Mpts[i2], '_ans1', kv); });
T3.on('up', function(){ syncInput(T3, Mpts[i3], '_ans2', kv); });
Tdv.on('up', function(){ syncInput(Tdv, O, '_ans3', 1); });
[[/jsxgraph]]

<div style="display:none">
[[input:ans1]]
[[input:ans2]]
[[input:ans3]]
</div>

<p>[[validation:ans1]] [[validation:ans2]] [[validation:ans3]]</p>
]]></text>
  </questiontext>

  <generalfeedback format="html">
    <text><![CDATA[
<p>Une fois le point lancé, il n'est soumis qu'à son poids : l'accélération est constante,
verticale, dirigée vers le bas, de valeur g, quelle que soit la forme de la piste parcourue
avant. On montre que \( \Delta \vec v = \vec a \cdot \Delta t \), un vecteur
<b>vertical, dirigé vers le bas, de norme constante</b> — la signature d'un mouvement
uniformément accéléré.</p>
]]></text>
  </generalfeedback>

  <defaultgrade>1.0000000</defaultgrade>
  <penalty>0.1000000</penalty>
  <hidden>0</hidden>
  <idnumber></idnumber>

  <questionvariables>
    <text><![CDATA[
g: 9.8$

dt: 0.15$
kExit: 6$
Mlist: [[0,0],[1.0922,0.3975],[2.1135,0.7693],[3.0614,1.1215],[3.8732,1.5581],[4.5029,2.0529],[4.9392,2.5106],[5.323,2.8499],[5.6938,2.9534],[6.0646,2.8364],[6.4355,2.4989]]$

v2: [2.4722, -0.045]$
v3: [2.4722, -1.515]$
dv: [v3[1]-v2[1], v3[2]-v2[2]]$

tol: 0.30$

xmin_t: -1$
xmax_t: 9$
ymin_t: -1$
ymax_t: 4$

vmax2: 3$
kv: 0.61$
]]></text>
  </questionvariables>

  <specificfeedback format="html">
    <text><![CDATA[[[feedback:ans1]] [[feedback:ans2]] [[feedback:ans3]]]]></text>
  </specificfeedback>

  <questionnote format="html">
    <text><![CDATA[Piste : ligne(3m,20°) -&gt; arc(R=5m,30°), dt=0.15 s, v_sortie=3.8461 m/s]]></text>
  </questionnote>

  <questiondescription format="html">
    <text></text>
  </questiondescription>

  <input>
    <name>ans1</name>
    <type>algebraic</type>
    <tans>v2</tans>
    <boxsize>15</boxsize>
    <strictsyntax>1</strictsyntax>
    <insertstars>0</insertstars>
    <syntaxhint></syntaxhint>
    <forbidwords></forbidwords>
    <mustverify>1</mustverify>
    <showvalidation>1</showvalidation>
    <options></options>
  </input>

  <input>
    <name>ans2</name>
    <type>algebraic</type>
    <tans>v3</tans>
    <boxsize>15</boxsize>
    <strictsyntax>1</strictsyntax>
    <insertstars>0</insertstars>
    <syntaxhint></syntaxhint>
    <forbidwords></forbidwords>
    <mustverify>1</mustverify>
    <showvalidation>1</showvalidation>
    <options></options>
  </input>

  <input>
    <name>ans3</name>
    <type>algebraic</type>
    <tans>dv</tans>
    <boxsize>15</boxsize>
    <strictsyntax>1</strictsyntax>
    <insertstars>0</insertstars>
    <syntaxhint></syntaxhint>
    <forbidwords></forbidwords>
    <mustverify>1</mustverify>
    <showvalidation>1</showvalidation>
    <options></options>
  </input>

  <prt>
    <name>prt1</name>
    <value>0.2500000</value>
    <autosimplify>1</autosimplify>
    <feedbackstyle>1</feedbackstyle>
    <node>
      <name>0</name>
      <answertest>NumAbsolute</answertest>
      <sans>ans1</sans>
      <tans>v2</tans>
      <testoptions>tol</testoptions>
      <quiet>0</quiet>
      <truescoremode>=</truescoremode>
      <truescore>1.0000000</truescore>
      <truepenalty></truepenalty>
      <truenextnode>-1</truenextnode>
      <trueanswernote>prt1-0-T</trueanswernote>
      <truefeedback format="html">
        <text><![CDATA[<p>Vecteur \( \vec v_2 \) correct.</p>]]></text>
      </truefeedback>
      <falsescoremode>=</falsescoremode>
      <falsescore>0.0000000</falsescore>
      <falsepenalty></falsepenalty>
      <falsenextnode>-1</falsenextnode>
      <falseanswernote>prt1-0-F</falseanswernote>
      <falsefeedback format="html">
        <text><![CDATA[<p>Vecteur \( \vec v_2 \) incorrect : reprends
\( \vec v_2 \approx \overrightarrow{M_{7}M_{9}}/(2\Delta t) \), à tracer
<b>à partir de M8</b>, à l'échelle indiquée.</p>]]></text>
      </falsefeedback>
    </node>
  </prt>

  <prt>
    <name>prt2</name>
    <value>0.2500000</value>
    <autosimplify>1</autosimplify>
    <feedbackstyle>1</feedbackstyle>
    <node>
      <name>0</name>
      <answertest>NumAbsolute</answertest>
      <sans>ans2</sans>
      <tans>v3</tans>
      <testoptions>tol</testoptions>
      <quiet>0</quiet>
      <truescoremode>=</truescoremode>
      <truescore>1.0000000</truescore>
      <truepenalty></truepenalty>
      <truenextnode>-1</truenextnode>
      <trueanswernote>prt2-0-T</trueanswernote>
      <truefeedback format="html">
        <text><![CDATA[<p>Vecteur \( \vec v_3 \) correct.</p>]]></text>
      </truefeedback>
      <falsescoremode>=</falsescoremode>
      <falsescore>0.0000000</falsescore>
      <falsepenalty></falsepenalty>
      <falsenextnode>-1</falsenextnode>
      <falseanswernote>prt2-0-F</falseanswernote>
      <falsefeedback format="html">
        <text><![CDATA[<p>Vecteur \( \vec v_3 \) incorrect : reprends
\( \vec v_3 \approx \overrightarrow{M_{8}M_{10}}/(2\Delta t) \), à tracer
<b>à partir de M9</b>, à l'échelle indiquée.</p>]]></text>
      </falsefeedback>
    </node>
  </prt>

  <prt>
    <name>prt3</name>
    <value>0.5000000</value>
    <autosimplify>1</autosimplify>
    <feedbackstyle>1</feedbackstyle>
    <node>
      <name>0</name>
      <answertest>NumAbsolute</answertest>
      <sans>ans3</sans>
      <tans>dv</tans>
      <testoptions>tol</testoptions>
      <quiet>0</quiet>
      <truescoremode>=</truescoremode>
      <truescore>1.0000000</truescore>
      <truepenalty></truepenalty>
      <truenextnode>-1</truenextnode>
      <trueanswernote>prt3-0-T</trueanswernote>
      <truefeedback format="html">
        <text><![CDATA[<p>Vecteur \( \Delta \vec v \) correct : il est bien vertical, dirigé
vers le bas, de norme \( g\Delta t \). C'est la preuve que l'accélération est
constante.</p>]]></text>
      </truefeedback>
      <falsescoremode>=</falsescoremode>
      <falsescore>0.0000000</falsescore>
      <falsepenalty></falsepenalty>
      <falsenextnode>-1</falsenextnode>
      <falseanswernote>prt3-0-F</falseanswernote>
      <falsefeedback format="html">
        <text><![CDATA[<p>Vecteur \( \Delta \vec v \) incorrect : reporte \( \vec v_2 \)
et \( \vec v_3 \) à partir de la même origine O, puis relie la pointe de
\( \vec v_2 \) à la pointe de \( \vec v_3 \) : c'est \( \Delta \vec v \).</p>]]></text>
      </falsefeedback>
    </node>
  </prt>

  <deployedseed></deployedseed>
  <qtest></qtest>

</question>
</quiz>