<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 396210  -->
  <question type="stack">
    <name>
      <text>Statistiques : Calcul d'une médiane (Effectif pair)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<div style="background:#475569;border-left:5px solid #334155;border-radius:0 8px 8px 0;padding:10px 16px;margin-bottom:12px;display:flex;align-items:center;gap:10px;flex-wrap:wrap;">
        <strong style="font-weight:800;color:#fff;font-size:.95rem;">Statistiques — Médiane d'une série</strong>
        <span style="background:#334155;color:#fff;padding:2px 9px;border-radius:20px;font-size:.78rem;font-weight:700;">/ 1 pt</span>
      </div>
      
      Voici une série de 6 mesures, rangées dans l'ordre croissant :<br/><br/>
      <div style="background:#f8fafc; border:1px solid #e2e8f0; border-radius:8px; padding:15px; text-align:center; font-size:1.1rem; margin-bottom:15px;">
        \( S = \left\{ {@str_L@} \right\} \)
      </div>
      
      Déterminer la médiane de cette série.<br/>
      <em>(On donnera le résultat exact sous forme de fraction si nécessaire, ou en décimal).</em><br/><br/>
      <div>\(Me = \) [[input:ans_med]] [[validation:ans_med]]</div>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<div style="margin-top:20px; padding:15px; background:#f8fafc; border:1px solid #cbd5e1; border-radius:8px;">
        <div style="font-weight:bold; margin-bottom:10px;">🔑 Correction détaillée</div>
        <div style="font-size:.9rem;">
          L'effectif total est \(n = 6\) (pair).<br/>
          La médiane est la moyenne des valeurs situées aux rangs \(\frac{n}{2}\) et \(\frac{n}{2} + 1\), soit aux rangs 3 et 4.<br/><br/>
          Les deux valeurs centrales sont : {@L[3]@} et {@L[4]@}.<br/>
          \(Me = \frac{{@L[3]@} + {@L[4]@}}{2} = \frac{{@L[3]+L[4]@}}{2} = {@ta_med@}\)
        </div>
      </div>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text></text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* 1. Génération de 6 nombres distincts triés par ordre croissant */
L: sort([rand(9)+1, rand(8)+10, rand(9)+20, rand(9)+30, rand(8)+40, rand(9)+50])$ 
/* 2. Affichage de la liste */
str_L: simplode(L, " ; ")$ 
/* 3. Calcul de la vraie médiane (Maxima gère les effectifs pairs tout seul) */
ta_med: float(median(L))$ 
/* 4. Les erreurs classiques pour le PRT */
ta_mean: float(mean(L))$ ta_sum_mid: L[3] + L[4]$ ta_one_mid: L[3]$]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<div>[[feedback:prt1]]</div>]]></text>
    </specificfeedback>
    <questionnote format="html">
      <text>Série={@str_L@}, Med={@ta_med@}</text>
    </questionnote>
    <questiondescription format="html">
      <text>Généré avec HéStack V2 | hestack-types:stats-median</text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans_med</name>
      <type>algebraic</type>
      <tans>ta_med</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Vérification de la médiane</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_med</sans>
        <tans>ta_med</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-OK</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left:4px solid #15803d;padding:10px 14px;background:#f0fdf4;border-radius:4px;">✅ <strong>Parfait !</strong> Vous avez correctement identifié que l'effectif était pair, trouvé les deux valeurs centrales, et fait leur moyenne.</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>PRT-NOK</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Erreur : A calculé la moyenne</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_med</sans>
        <tans>ta_mean</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-MEAN</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left:4px solid #f97316;padding:10px 14px;background:#fff7ed;border-radius:4px;">🚨 <strong>Confusion !</strong> Vous avez calculé la MOYENNE de la série. On vous demande la MÉDIANE. Ce n'est pas la même formule !</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>PRT-ERR-AUTRE</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Erreur : A fait la somme au lieu de la moyenne</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_med</sans>
        <tans>ta_sum_mid</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.25</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-SUM</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left:4px solid #f97316;padding:10px 14px;background:#fff7ed;border-radius:4px;">⚠️ <strong>Presque !</strong> Vous avez bien trouvé les deux valeurs centrales ({@L[3]@} et {@L[4]@}), mais vous les avez ADDITIONNÉES. Comme l'effectif est pair, la médiane est la MOYENNE de ces deux valeurs, il faut diviser par 2 !</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>PRT-ERR-AUTRE2</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Erreur : A donné une seule valeur centrale</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_med</sans>
        <tans>ta_one_mid</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-ONE</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left:4px solid #dc2626;padding:10px 14px;background:#fff0f0;border-radius:4px;">❌ <strong>Effet pair !</strong> L'effectif est de 6 (pair). Il n'y a pas une seule valeur centrale, il y en a DEUX. Vous devez faire la moyenne de la 3ème et de la 4ème valeur.</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>PRT-ERR-CALC</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<div style="border-left:4px solid #dc2626;padding:10px 14px;background:#fff0f0;border-radius:4px;">❌ <strong>Incorrect.</strong> Rangez les valeurs (elles le sont déjà), trouvez la position du milieu (entre la 3ème et la 4ème ici), et calculez leur moyenne.</div>]]></text>
        </falsefeedback>
      </node>
    </prt>
  </question>

</quiz>