<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 396218  -->
  <question type="stack">
    <name>
      <text>Statistiques : Calcul d'une variance</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<div style="background: #475569; border-left: 5px solid #334155; border-radius: 0 8px 8px 0; padding: 10px 16px; margin-bottom: 12px; display: flex; align-items: center; gap: 10px; flex-wrap: wrap;"><strong style="font-weight: 800; color: #fff; font-size: .95rem;">Statistiques — Variance d'une série</strong> <span style="background: #334155; color: #fff; padding: 2px 9px; border-radius: 20px; font-size: .78rem; font-weight: bold;">/ 1 pt</span></div>
<p>Soit la série statistique suivante composée de 5 mesures :<br><br></p>
<div style="background: #f8fafc; border: 1px solid #e2e8f0; border-radius: 8px; padding: 15px; text-align: center; font-size: 1.1rem; margin-bottom: 15px;">\( S = \left\{ {@str_L@} \right\} \)</div>
<p>Calculer la variance de cette série.<br><em>(Arrondir le résultat à \(10^{-2}\) près).</em><br><br></p>
<div>\(V \approx\) [[input:ans_var]] [[validation:ans_var]]</div>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<div style="margin-top: 20px; padding: 15px; background: #f8fafc; border: 1px solid #cbd5e1; border-radius: 8px;">
<div style="font-weight: bold; margin-bottom: 10px;">🔑 Correction détaillée</div>
<div style="font-size: .9rem;">La formule de la variance (statistique descriptive) est :<br>\(V = \frac{1}{n} \sum_{i=1}^{n} (x_i - \bar{x})^2\)<br><br>1. La moyenne est : \(\bar{x} = \frac{{@ta_somme@}}{5} = {@ta_mean@}\)<br>2. On calcule les carrés des écarts à la moyenne, on les somme, puis on divise par l'effectif total \(n = 5\).<br>\(V \approx {@ta_var@}\)</div>
</div>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* 1. Définitions maison des fonctions stats (pour éviter le "load" interdit) */
mean(L) := apply("+", L) / length(L);
var(L) := block([m], m: mean(L), mean((L - m)^2));
 std(L) := sqrt(var(L)); 
/* 2. Génération de 5 nombres entiers aléatoires */
L: [rand(8)+1, rand(8)+1, rand(8)+1, rand(8)+1, rand(8)+1];
/* 3. Affichage de la liste */
str_L: simplode(L, " ; "); 
/* 4. Calculs pour la correction */
ta_somme: apply("+", L);
 ta_mean: float(ta_somme / 5);
 ta_var: 1.00*round(float(var(L)) * 100) / 100; 
/* 5. Les erreurs classiques pour le PRT */
ta_err_somme: round(float(var(L) * 5) * 100) / 100;
 ta_err_std: round(float(std(L)) * 100) / 100;
 ta_err_n1: round(float(var(L) * 5 / 4) * 100) / 100;]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<div>[[feedback:prt1]]</div>]]></text>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>Série={@str_L@}, V={@ta_var@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p>Généré avec HéStack V2 | hestack-types:stats-variance</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans_var</name>
      <type>numerical</type>
      <tans>1.00*ta_var</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Vérification de la variance</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans_var</sans>
        <tans>ta_var</tans>
        <testoptions>0.015</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-OK</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #15803d; padding: 10px 14px; background: #f0fdf4; border-radius: 4px;">✅ <strong>Parfait !</strong> La variance a été correctement calculée avec la formule \(\frac{1}{n}\sum(x_i - \bar{x})^2\).</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>PRT-NOK</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Erreur : A calculé l'écart-type</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans_var</sans>
        <tans>ta_err_std</tans>
        <testoptions>0.015</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-STD</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Confusion !</strong> Vous avez calculé l'ÉCART-TYPE (\(\sigma\)), c'est-à-dire la racine carrée de la variance. On demande ici la VARIANCE (\(\sigma^2\)). Ne prenez pas la racine carrée !</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>PRT-ERR-AUTRE</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Erreur : Oubli de la division par n</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans_var</sans>
        <tans>ta_err_somme</tans>
        <testoptions>0.1</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-SOMME</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Oubli de la division !</strong> Vous avez calculé la SOMME des carrés des écarts à la moyenne. C'est bien, mais il faut diviser ce résultat par l'effectif total \(n = 5\) pour obtenir la variance.</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>PRT-ERR-AUTRE2</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Erreur : Division par n-1</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans_var</sans>
        <tans>ta_err_n1</tans>
        <testoptions>0.015</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-N1</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #dc2626; padding: 10px 14px; background: #fff0f0; border-radius: 4px;">❌ <strong>Erreur de formule !</strong> Vous avez divisé par \(n-1\) (ici 4) au lieu de diviser par \(n\) (ici 5). La division par \(n-1\) est réservée à la variance d'échantillon en probabilités. En statistiques descriptives, on divise par \(n\).</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>PRT-ERR-CALC</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #dc2626; padding: 10px 14px; background: #fff0f0; border-radius: 4px;">❌ <strong>Incorrect.</strong> Appliquez strictement la formule : calculez la moyenne, puis \(\frac{1}{5}\sum(x_i - \bar{x})^2\). N'oubliez pas de diviser par 5 et de mettre au carré.</div>]]></text>
        </falsefeedback>
      </node>
    </prt>
  </question>

</quiz>