<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 396232  -->
  <question type="stack">
    <name>
      <text>Statistiques : Calcul du premier quartile Q1</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<div style="background: #475569; border-left: 5px solid #334155; border-radius: 0 8px 8px 0; padding: 10px 16px; margin-bottom: 12px; display: flex; align-items: center; gap: 10px; flex-wrap: wrap;"><strong style="font-weight: 800; color: #fff; font-size: .95rem;">Statistiques — Quartiles</strong> <span style="background: #334155; color: #fff; padding: 2px 9px; border-radius: 20px; font-size: .78rem; font-weight: bold;">/ 1 pt</span></div>
<p>Voici une série statistique de 8 mesures, rangées dans l'ordre croissant :<br><br></p>
<div style="background: #f8fafc; border: 1px solid #e2e8f0; border-radius: 8px; padding: 15px; text-align: center; font-size: 1.1rem; margin-bottom: 15px;">\( S = \left\{ {@str_L@} \right\} \)</div>
<p>Déterminer la valeur exacte du premier quartile \(Q_1\) de cette série.<br><em>(Rappel : \(Q_1\) est la plus petite valeur de la série telle qu'au moins 25 % des valeurs lui sont inférieures ou égales).</em><br><br></p>
<div>\(Q_1 = \) [[input:ans_q1]] [[validation:ans_q1]]</div>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<div style="margin-top: 20px; padding: 15px; background: #f8fafc; border: 1px solid #cbd5e1; border-radius: 8px;">
<div style="font-weight: bold; margin-bottom: 10px;">🔑 Correction détaillée</div>
<div style="font-size: .9rem;">L'effectif total est \(n = 8\).<br>La position du premier quartile est : \(\frac{1 \times 8}{4} = 2\).<br><br>Comme la position est un entier, \(Q_1\) est simplement la 2ème valeur de la série rangée dans l'ordre croissant.<br>\(Q_1 = {@L[pos_Q1]@}\)</div>
</div>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* 1. Génération de 8 nombres (dont quelques décimaux pour la variété) */
L: sort([1.0*rand(50)/10 + 5, 1.0*rand(60)/10 + 7, 1.0*rand(40)/10 + 10,1.0* rand(50)/10 + 12, 1.0*rand(70)/10 + 14, 1.0*rand(60)/10 + 16, 1.0*rand(80)/10 + 18, 1.0*(rand(50)/10 +20)]);
/* 2. Affichage de la liste */
str_L: simplode(L, " ; "); 
/* 3. Calculs pour la correction */
n: length(L);
 pos_Q1: n/4;
 ta_q1: 1.0*L[pos_Q1]; 
/* 4. Les erreurs classiques pour le PRT */
/* Erreur 1 : A donné la médiane (4ème valeur) */
ta_median: L[n/2];
 /* Erreur 2 : A donné la moyenne de toute la série */
ta_mean: float(apply("+", L) / n);
 /* Erreur 3 : A fait la moyenne des 3 premières valeurs au lieu de lire la 2ème */
ta_mean_first3: float((L[1]+L[2]+L[3])/3);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<div>[[feedback:prt1]]</div>]]></text>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>Série n=8, Pos Q1={@pos_Q1@}, Q1={@ta_q1@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p>Généré avec HéStack V2 | hestack-types:stats-quartile</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans_q1</name>
      <type>algebraic</type>
      <tans>ta_q1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Vérification de Q1</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_q1</sans>
        <tans>ta_q1</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-OK</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #15803d; padding: 10px 14px; background: #f0fdf4; border-radius: 4px;">✅ <strong>Parfait !</strong> Vous avez bien calculé la position (\(\frac{8}{4}=2\)) et lu la bonne valeur.</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>PRT-NOK</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Erreur : A donné la médiane</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_q1</sans>
        <tans>ta_median</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-MED</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Confusion des quartiles !</strong> Vous avez donné la valeur de la MÉDIANE (qui se trouve au rang \(\frac{8}{2}=4\)). On demande le PREMIER quartile \(Q_1\), qui se trouve au rang \(\frac{8}{4}=2\).</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>PRT-ERR-AUTRE</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Erreur : A donné la moyenne</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_q1</sans>
        <tans>ta_mean</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-MEAN</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Confusion des indicateurs !</strong> Vous avez calculé la MOYENNE de la série. Un quartile est une valeur de la série elle-même (repérée par son rang), pas un calcul impliquant toutes les valeurs.</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>PRT-ERR-AUTRE2</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Erreur : Moyenne des 3 premières valeurs</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_q1</sans>
        <tans>ta_mean_first3</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-MEAN3</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #dc2626; padding: 10px 14px; background: #fff0f0; border-radius: 4px;">❌ <strong>Erreur de position !</strong> Vous avez fait la moyenne des 3 premières valeurs. Pour trouver un quartile, il ne faut faire AUCUN calcul, il suffit de compter le rang (\(\frac{1 \times 8}{4}=2\)) et de lire la 2ème valeur !</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>PRT-ERR-CALC</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #dc2626; padding: 10px 14px; background: #fff0f0; border-radius: 4px;">❌ <strong>Incorrect.</strong> Calculez le rang de \(Q_1\) : \(\frac{1 \times n}{4}\). Puisque le rang est un entier, \(Q_1\) est simplement la valeur située à ce rang.</div>]]></text>
        </falsefeedback>
      </node>
    </prt>
  </question>

</quiz>