<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 396201  -->
  <question type="stack">
    <name>
      <text>Primitive : Exponentielle et Constante</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<div style="background: #475569; border-left: 5px solid #334155; border-radius: 0 8px 8px 0; padding: 10px 16px; margin-bottom: 12px; display: flex; align-items: center; gap: 10px; flex-wrap: wrap;"><strong style="font-weight: 800; color: #fff; font-size: .95rem;">Analyse — Calcul de Primitives</strong> <span style="background: #334155; color: #fff; padding: 2px 9px; border-radius: 20px; font-size: .78rem; font-weight: bold;">/ 1 pt</span></div>
<p>Soit \(f\) la fonction définie sur \(\mathbb{R}\) par :<br><br></p>
<div style="background: #f8fafc; border: 1px solid #e2e8f0; border-radius: 8px; padding: 15px; text-align: center; font-size: 1.1rem; margin-bottom: 15px;">\( f(x) = {@f@} \)</div>
<p>Déterminer <strong>une</strong> primitive \(F\) de la fonction \(f\) sur \(\mathbb{R}\).<br><em>(On n'oubliera pas la constante d'intégration, par exemple + k).</em><br><br></p>
<div>\(F(x) = \) [[input:ans_F]] [[validation:ans_F]]</div>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<div style="margin-top: 20px; padding: 15px; background: #f8fafc; border: 1px solid #cbd5e1; border-radius: 8px;">
<div style="font-weight: bold; margin-bottom: 10px;">🔑 Correction détaillée</div>
<div style="font-size: .9rem;">On décompose la fonction pour trouver les primitives de chaque terme :<br><br>- Pour {@b@}×\(e^{{@a@}x}\), on sait que \(\int e^{kx} dx = \frac{1}{k}e^{kx}\). Donc la primitive est \(\frac{{@b@}}{{@a@}}e^{{@a@}x}\).<br>- Pour {@c@}, une primitive est {@c@}×\(x\).<br><br>Une primitive de \(f\) est donc :<br>\(F(x) = \frac{{@b@}}{{@a@}}e^{{@a@}x} + {@c@}x + k\) (avec \(k \in \mathbb{R}\))</div>
</div>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>/* 1. Paramètres aléatoires */
a: rand([-2,-1,1,2])$ b: rand([2,3,4,5])$ c: rand([1,2,3]);
/* 2. Fonction de départ */
f: b*%e^(a*x) + c;
/* 3. Les erreurs classiques pour le PRT */
ta_fp_sans_k: (b/a)*%e^(a*x) + c*x$ ta_err_coef: b*%e^(a*x) + c*x;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<div>[[feedback:prt1]]</div>]]></text>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>f(x)={@f@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p>Généré avec HéStack V2 | hestack-types:primitives</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans_F</name>
      <type>algebraic</type>
      <tans>(b/a)*%e^(a*x) + c*x + k</tans>
      <boxsize>30</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>utilisez exp() pour l'exponentielle</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords>k</allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>2</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Vérification par dérivation</description>
        <answertest>AlgEquiv</answertest>
        <sans>diff(ans_F, x)</sans>
        <tans>f</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-DIFF-OK</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #15803d; padding: 10px 14px; background: #f0fdf4; border-radius: 4px;">✅ <strong>Parfait !</strong> Si on dérive votre expression, on retombe bien sur \(f(x)\). Vous avez trouvé une primitive correcte.</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>PRT-DIFF-NOK</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Détection oubli de la constante</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_F</sans>
        <tans>ta_fp_sans_k</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-K</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Presque !</strong> Votre forme ressemble à la primitive sans la constante, mais il y a une erreur de calcul dedans. De plus, n'oubliez pas le "+ k".</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>PRT-ERR-AUTRE</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Détection erreur de coefficient exp</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_F</sans>
        <tans>ta_err_coef</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.25</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-COEF</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #dc2626; padding: 10px 14px; background: #fff0f0; border-radius: 4px;">❌ <strong>Erreur de coefficient !</strong> Vous avez écrit {@b@}×\(e^{{@a@}x}\) au lieu de diviser par {@a@}. Retenez : \(\int e^{kx} dx = \frac{1}{k}e^{kx}\).</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>PRT-ERR-CALC</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #dc2626; padding: 10px 14px; background: #fff0f0; border-radius: 4px;">❌ <strong>Incorrect.</strong> Vérifiez vos formules de primitives. \(\int e^{ax} dx \neq e^{ax}\) (il faut diviser par \(a\)). N'oubliez pas la constante.</div>]]></text>
        </falsefeedback>
      </node>
    </prt>
  </question>

</quiz>