<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 396182  -->
  <question type="stack">
    <name>
      <text>Dérivée : Piège du produit (Exponentielle)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<div style="background: #475569; border-left: 5px solid #334155; border-radius: 0 8px 8px 0; padding: 10px 16px; margin-bottom: 12px; display: flex; align-items: center; gap: 10px; flex-wrap: wrap;"><strong style="font-weight: 800; color: #fff; font-size: .95rem;">Calcul diff. — Dérivée d'un produit</strong> <span style="background: #334155; color: #fff; padding: 2px 9px; border-radius: 20px; font-size: .78rem; font-weight: bold;">/ 1 pt</span></div>
<p>Soit \(f\) la fonction définie sur \(\mathbb{R}\) par :<br><br></p>
<div style="background: #f8fafc; border: 1px solid #e2e8f0; border-radius: 8px; padding: 15px; text-align: center; font-size: 1.1rem; margin-bottom: 15px;">\( f(x) = ({@u@}) \cdot e^x \)</div>
<p>Calculer la fonction dérivée \(f'\) de \(f\).<br><br></p>
<div>\(f'(x) = \) [[input:ans_fp]] [[validation:ans_fp]]</div>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<div style="margin-top: 20px; padding: 15px; background: #f8fafc; border: 1px solid #cbd5e1; border-radius: 8px;">
<div style="font-weight: bold; margin-bottom: 10px;">🔑 Correction détaillée</div>
<div style="font-size: .9rem;">On utilise la formule de la dérivée d'un produit : \((uv)' = u'v + uv'\).<br><br>Ici, on pose \(u(x) = {@u@}\) donc \(u'(x) = {@du@}\).<br>Et \(v(x) = e^x\) donc \(v'(x) = e^x\).<br><br>\(f'(x) = {@du@} \cdot e^x + ({@u@}) \cdot e^x\)<br>\(f'(x) = ({@du@} + {@u@})e^x = {@fp_simpl@}\)</div>
</div>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>/* 1. Génération aléatoire */
a: rand([1,2,3])$ b: rand([-3,-2,-1,1,2])$ 
/* 2. Définition des fonctions */
u: a*x + b$ du: a$ v: %e^x$ dv: %e^x$ 
/* 3. Calculs pour la correction */
f: u * v$ fp: diff(f, x)$ fp_simpl: ev(fp, simp)$ 
/* 4. Les erreurs classiques des élèves */
err_oublie_poly: u * dv$                /* Il dérive juste e^x */
err_oublie_exp: du * v$                  /* Il dérive juste le polynome */
err_produit: du * dv$                    /* Il fait u' * v' */
err_plus: (u + v) * (du + dv)$          /* Il fait (u+v)' * (u+v)' */</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<div>[[feedback:prt1]]</div>]]></text>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>f(x)=({@u@})e^x, f'={@fp_simpl@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p>Généré avec HéStack V2 | hestack-types:derivative-product</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans_fp</name>
      <type>algebraic</type>
      <tans>fp</tans>
      <boxsize>30</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>2</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Vérification de la dérivée</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_fp</sans>
        <tans>fp</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-OK</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #15803d; padding: 10px 14px; background: #f0fdf4; border-radius: 4px;">✅ <strong>Parfait !</strong> La dérivée d'un produit est bien la somme des deux "dérivées partielles" : \(u'v + uv'\).</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>PRT-NOK</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Erreur : Oubli dérivée polynôme</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_fp</sans>
        <tans>err_oublie_poly</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-POLY</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Erreur de formule !</strong> Vous avez écrit {@u@} \(\times e^x\). Vous avez dérivé \(e^x\), mais vous avez <strong>oublié de dériver le polynôme {@u@}</strong>. Appliquez bien \((uv)' = u'v + uv'\).</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>PRT-ERR-AUTRE</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Erreur : Oubli dérivée exponentielle</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_fp</sans>
        <tans>err_oublie_exp</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-EXP</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Erreur de formule !</strong> Vous avez écrit {@du@} \(\times e^x\). Vous avez dérivé le polynôme, mais vous avez <strong>oublié de dériver \(e^x\)</strong> (qui se dérive en lui-même). Appliquez bien \((uv)' = u'v + uv'\).</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>PRT-ERR-AUTRE2</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Erreur : Multiplication des dérivées</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_fp</sans>
        <tans>err_produit</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-PROD</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #dc2626; padding: 10px 14px; background: #fff0f0; border-radius: 4px;">❌ <strong>Grosse erreur de formule !</strong> Vous avez calculé \(u' \times v'\) (la multiplication des dérivées). La formule du produit est une <strong>SOMME</strong> : \((uv)' = u'v + uv'\).</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>PRT-ERR-CALC</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #dc2626; padding: 10px 14px; background: #fff0f0; border-radius: 4px;">❌ <strong>Incorrect.</strong> Revoyez la formule de dérivation d'un produit : \((uv)' = u'v + uv'\). N'oubliez pas les parenthèses si nécessaire.</div>]]></text>
        </falsefeedback>
      </node>
    </prt>
  </question>

</quiz>