<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 396664  -->
  <question type="stack">
    <name>
      <text>Oscilloscope - Décharge RC (Logique Vrais Boutons)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<div style="background:#0c4a6e;border-left:5px solid #0369a1;border-radius:0 8px 8px 0;padding:10px 16px;margin-bottom:12px;display:flex;align-items:center;gap:10px;flex-wrap:wrap;">
        <strong style="font-weight:800;color:#fff;font-size:.95rem;">Oscilloscope — Décharge d'un condensateur (RC)</strong>
        <span style="background:#0369a1;color:#fff;padding:2px 9px;border-radius:20px;font-size:.78rem;font-weight:700;">/ 2 pts</span>
        <span style="background:#ffffff;color:#0369a1;border:1px solid #0369a1;padding:2px 9px;border-radius:20px;font-size:.75rem;font-weight:600;">📏 Grandeur physique</span>
      </div>
      <p>On visualise sur l'oscilloscope la tension \(u_C(t)\) aux bornes d'un condensateur lors de sa décharge.</p>
      <p><strong>Consigne :</strong> 
      <ul>
        <li>Avec <span style="color:#ef4444;font-weight:bold;">■ Y</span>, mesurez l'amplitude de la tension initiale \(E\) (entre le palier haut et la ligne du bas).</li>
        <li>Avec <span style="color:#a855f7;font-weight:bold;">■ X</span>, mesurez la constante de temps \(\tau\) (entre l'instant \(t=0\) et le point correspondant à une chute de 63,2% de \(E\)).</li>
      </ul>
      </p>
      
[[jsxgraph width="775px" height="550px"]]
(function(){
var lSH=[5e-7,0.000001,0.000002,0.000005,0.00001,0.00002,0.00005,0.0001,0.0002,0.0005,0.001,0.002,0.005,0.01,0.02,0.05,0.1,0.2,0.5];
var lSV=[0.005,0.01,0.02,0.05,0.1,0.2,0.5,1,2,5,10];
var si=7, ti=11, yp=0, xp=0; 
var ev = parseFloat("{#ev_val#}") / 1000;
var ta = parseFloat("{#ta_val#}") / 1000000;
function SV(){return lSV[si];}
function SH(){return lSH[ti];}
function fV(v){return v<0.05?(v*1e3).toPrecision(3)+" mV":v+" V";}
function fT(s){return s<5e-5?(s*1e6).toPrecision(3)+" \u00B5s":s<0.05?(s*1e3).toPrecision(3)+" ms":s.toPrecision(3)+" s";}
var bd=JXG.JSXGraph.initBoard(divid,{boundingbox:[-5.4,4.5,5.4,-4.8],axis:false,grid:false,showNavigation:false,showCopyright:false,pan:{enabled:false},zoom:{enabled:false}});
bd.create("polygon",[[-5,-4],[5,-4],[5,4],[-5,4]],{fillColor:"#ffffff",fillOpacity:1,strokeWidth:0,fixed:true,layer:0,highlight:false, vertices:{visible:false, fixed:true}, borders:{visible:false}});
bd.create("segment",[[-5,-4],[5,-4]], {strokeColor:"#4b5563",strokeWidth:2,fixed:true,layer:0,highlight:false});
bd.create("segment",[[5,-4],[5,4]], {strokeColor:"#4b5563",strokeWidth:2,fixed:true,layer:0,highlight:false});
bd.create("segment",[[5,4],[-5,4]], {strokeColor:"#4b5563",strokeWidth:2,fixed:true,layer:0,highlight:false});
bd.create("segment",[[-5,4],[-5,-4]], {strokeColor:"#4b5563",strokeWidth:2,fixed:true,layer:0,highlight:false});
var gs={strokeColor:"#e5e7eb",strokeWidth:0.8,fixed:true,layer:1,highlight:false};
var gc={strokeColor:"#d1d5db",strokeWidth:1.2,fixed:true,layer:1,highlight:false};
var ts={strokeColor:"#9ca3af",strokeWidth:0.55,fixed:true,layer:1,highlight:false};
for(var xi=-4;xi<=4;xi++) bd.create("segment",[[xi,-4],[xi,4]],xi===0?gc:gs);
for(var yi=-3;yi<=3;yi++) bd.create("segment",[[-5,yi],[5,yi]],yi===0?gc:gs);
for(var xt=-5;xt<=5.01;xt+=0.2) if(Math.abs(xt%1)>0.05) bd.create("segment",[[xt,-0.07],[xt,0.07]],ts);
for(var yt=-4;yt<=4.01;yt+=0.2) if(Math.abs(yt%1)>0.05) bd.create("segment",[[-0.07,yt],[0.07,yt]],ts);
function rcSig(xd){var t=(xd+5-xp)*SH(),sv=SV(); var y = t>=0 ? ev/sv*(Math.exp(-t/ta))+yp : ev/sv+yp; return y>4?4:y<-4?-4:y;}
bd.create("functiongraph",[rcSig,-5,5],{strokeColor:"#2563eb",strokeWidth:3,layer:3,highlight:false,fixed:true, numberPoints:150, doAdvancedPlot:false});
var lbSH=bd.create("text",[0,-3.75,function(){return fT(SH())+"/div";}],{fixed:true,color:"#1f2937",fontSize:13,anchorX:"middle",layer:4,highlight:false});
var lbSV=bd.create("text",[-4.7,3.6,function(){return fV(SV())+"/div";}],{fixed:true,color:"#1f2937",fontSize:13,layer:4,highlight:false});
var cA=bd.create("point",[-5,0],{size:7,fillColor:"#a855f7",strokeColor:"#a855f7",name:"",layer:5,showInfobox:false,label:{visible:false}, visible:false, drag: ['x']});
var cB=bd.create("point",[-2,0],{size:7,fillColor:"#a855f7",strokeColor:"#a855f7",name:"",layer:5,showInfobox:false,label:{visible:false}, visible:false, drag: ['x']});
var segPA=bd.create("segment",[function(){return[cA.X(),-4];},function(){return[cA.X(),4];}],{strokeColor:"#a855f7",strokeWidth:1.5,dash:3,layer:5,highlight:false, visible:false});
var segPB=bd.create("segment",[function(){return[cB.X(),-4];},function(){return[cB.X(),4];}],{strokeColor:"#a855f7",strokeWidth:1.5,dash:3,layer:5,highlight:false, visible:false});
var txtP=bd.create("text",[0,3.4,function(){var dx=Math.abs(cB.X()-cA.X()),dts=dx*SH(); return "\u0394t="+fT(dts);}],{fixed:false,color:"#7c3aed",fontSize:14,anchorX:"middle",layer:6,highlight:false, visible:false});
var cD=bd.create("point",[0, 2],{size:7,fillColor:"#ef4444",strokeColor:"#ef4444",name:"",layer:5,showInfobox:false,label:{visible:false}, visible:false, drag: ['y']});
var cE=bd.create("point",[0, 0],{size:7,fillColor:"#ef4444",strokeColor:"#ef4444",name:"",layer:5,showInfobox:false,label:{visible:false}, visible:false, drag: ['y']});
var segRD=bd.create("segment",[function(){return[-5,cD.Y()];},function(){return[5,cD.Y()];}],{strokeColor:"#ef4444",strokeWidth:1.5,dash:3,layer:5,highlight:false, visible:false});
var segRE=bd.create("segment",[function(){return[-5,cE.Y()];},function(){return[5,cE.Y()];}],{strokeColor:"#ef4444",strokeWidth:1.5,dash:3,layer:5,highlight:false, visible:false});
var txtR=bd.create("text",[4.5,3.2,function(){var dy=Math.abs(cD.Y()-cE.Y()),dv=dy*SV(); return "\u0394V="+fV(dv);}],{fixed:false,color:"#dc2626",fontSize:14,anchorX:"right",layer:6,highlight:false, visible:false});
function toggleGroup(arr) {var isVisible = arr[0].getAttribute('visible'); arr.forEach(function(obj) { obj.setAttribute({visible: !isVisible}); });}
var purpleGroup = [cA, cB, segPA, segPB, txtP];
var redGroup = [cD, cE, segRD, segRE, txtR];
var style = document.createElement('style');
style.textContent = '.osc-btn{background:#ffffff;color:#374151;border:2px solid #6b7280;padding:5px 12px;cursor:pointer;border-radius:4px;font-size:15px;font-family:monospace;font-weight:bold;margin:0 2px;box-shadow: 0 1px 2px rgba(0,0,0,0.05);}.osc-btn:hover{background:#f3f4f6;border-color:#374151;}';
document.head.appendChild(style);
var ctrl = document.createElement('div');
ctrl.style.cssText = 'position:absolute; bottom:0; left:0; width:100%; box-sizing:border-box; display:flex; align-items:center; gap:6px; flex-wrap:wrap; padding:8px 12px; background:rgba(255,255,255,0.95); border-top:2px solid #4b5563; z-index:100; user-select:none; box-shadow: 0 -2px 10px rgba(0,0,0,0.05);';
function mkL(t,c,w){ var s=document.createElement('span'); s.textContent=t; s.style.cssText='color:'+(c||'#374151')+';font-size:14px;font-family:monospace;font-weight:bold;display:inline-block;'+(w?'min-width:'+w+'px;text-align:center;':''); return s; }
function mkS(){ var s=document.createElement('span'); s.style.cssText='border-left:1px solid #d1d5db;height:24px;display:inline-block;margin:0 6px;'; return s; }
function mkB(t,c,fn){ var b=document.createElement('button'); b.className='osc-btn'; b.textContent=t; if(c){b.style.color=c; b.style.borderColor=c;} if(fn) b.addEventListener('click', fn); return b; }
var lblV = mkL('', '#374151', '65px');
var lblT = mkL('', '#374151', '75px');
ctrl.appendChild(mkL('V/div', '#dc2626', '45px'));
ctrl.appendChild(mkB('\u25BC', '#dc2626', function(){si=Math.max(0,si-1);bd.update();}));
ctrl.appendChild(lblV);
ctrl.appendChild(mkB('\u25B2', '#dc2626', function(){si=Math.min(lSV.length-1,si+1);bd.update();}));
ctrl.appendChild(mkS());
ctrl.appendChild(mkL('t/div', '#2563eb', '45px'));
ctrl.appendChild(mkB('\u25C4', '#2563eb', function(){ti=Math.max(0,ti-1);bd.update();}));
ctrl.appendChild(lblT);
ctrl.appendChild(mkB('\u25BA', '#2563eb', function(){ti=Math.min(lSH.length-1,ti+1);bd.update();}));
ctrl.appendChild(mkL('XPOS', '#7c3aed', '50px')); 
ctrl.appendChild(mkB('\u25C4', '#7c3aed', function(){xp=Math.max(-5,parseFloat((xp-0.5).toFixed(1)));bd.update();}));
ctrl.appendChild(mkB('\u25BA', '#7c3aed', function(){xp=Math.min(5,parseFloat((xp+0.5).toFixed(1)));bd.update();}));
ctrl.appendChild(mkS());
ctrl.appendChild(mkL('YPOS', '#d97706', '50px'));
ctrl.appendChild(mkB('\u25BC', '#d97706', function(){yp=Math.max(-4,parseFloat((yp-0.5).toFixed(1)));bd.update();}));
ctrl.appendChild(mkB('\u25B2', '#d97706', function(){yp=Math.min(4,parseFloat((yp+0.5).toFixed(1)));bd.update();}));
var spacer = document.createElement('div');
spacer.style.cssText = 'flex-grow:1;';
ctrl.appendChild(spacer);
ctrl.appendChild(mkB('\u25A0 X', '#7c3aed', function(){ toggleGroup(purpleGroup); }));
ctrl.appendChild(mkB('\u25A0 Y', '#dc2626', function(){ toggleGroup(redGroup); }));
var boardContainer = document.getElementById(divid);
boardContainer.style.position = 'relative';
boardContainer.appendChild(ctrl);
var _ou = bd.update.bind(bd);
bd.update = function(){ _ou(); lblV.textContent = fV(SV()); lblT.textContent = fT(SH()); };
bd.update();
})();
[[/ jsxgraph]]

      <div style="background:#f8f9fa;border:1px solid #dee2e6;padding:15px;border-radius:8px;margin:15px 0;">
        <strong>Conseils pour la saisie :</strong>
        <ul style="margin:8px 0 0 20px;line-height:1.7">
          <li>Entrez des nombres entiers (ex : <code>2500</code>).</li>
          <li><strong>Lier nombre et unité</strong> : reliez-les par <code>*</code> (ex : <code>1500*us</code> pour µs).</li>
          <li><strong>Unités usuelles</strong> : <code>s</code>, <code>ms</code>, <code>us</code>, <code>V</code>, <code>mV</code>.</li>
        </ul>
      </div>
      <p>1. Constante de temps \( \tau \) du circuit : [[input:ans_tau]] [[validation:ans_tau]]</p>
      <p>2. Tension initiale \( E \) du générateur : [[input:ans_E]] [[validation:ans_E]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<div style="margin-top:20px; padding:15px; background:#f8fafc; border:1px solid #cbd5e1; border-radius:8px; box-shadow:0 2px 4px rgba(0,0,0,0.05);">
        <div style="font-weight:bold; color:#0f766e; margin-bottom:10px; display:flex; align-items:center; gap:8px;">
          <span style="font-size:1.2rem;">🔑</span> Réponses attendues
        </div>
        <div style="margin-bottom:8px;font-size:.9rem;border-bottom:1px dashed #e2e8f0;padding-bottom:6px;">
          <span style="font-weight:bold;color:#0f766e;">Q1 Constante de temps :</span>
          <p>En mesurant entre \(t=0\) et le point correspondant à une chute de 63,2% de la tension initiale, on trouve \(\tau\) = {@ta_tau@}.</p>
        </div>
        <div style="margin-bottom:8px;font-size:.9rem;border-bottom:1px dashed #e2e8f0;padding-bottom:6px;">
          <span style="font-weight:bold;color:#0f766e;">Q2 Tension initiale :</span>
          <p>En mesurant le palier de tension avec les curseurs Y, on trouve \(E\) = {@ta_E@}.</p>
        </div>]]></text>
    </generalfeedback>
    <defaultgrade>2</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>ev_val: 2000 + rand(1000);
ta_val: 1000 + rand(1000);

ta_E: stackunits(ev_val, mV);
ta_tau: stackunits(ta_val, us);

tau_unit: stack_units_units(ta_tau);
E_unit: stack_units_units(ta_E);

precision_q: 0.1;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="html">
      <text>E={@ta_E@} | tau={@ta_tau@}</text>
    </questionnote>
    <questiondescription format="html">
      <text>Genere avec HeStack V2</text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>E</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans_E</name>
      <type>units</type>
      <tans>ta_E</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans_tau</name>
      <type>units</type>
      <tans>ta_tau</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>2.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>2</feedbackstyle>
      <feedbackvariables>
        <text>stud_si_tau : stack_unit_si_to_si_base(ans_tau);
teach_si_tau : stack_unit_si_to_si_base(tau_unit);
v_pure_e_tau : subst(map(lambda([u], u=1), listofvars(stud_si_tau)), stud_si_tau);
v_pure_t_tau : subst(map(lambda([u], u=1), listofvars(teach_si_tau)), teach_si_tau);
eleve_unit_tau : 2 * stud_si_tau / v_pure_e_tau;
teacher_unit_tau : 2 * teach_si_tau / v_pure_t_tau;

/* --- Traitement E --- */
stud_si_E : stack_unit_si_to_si_base(ans_E);
teach_si_E : stack_unit_si_to_si_base(E_unit);
v_pure_e_E : subst(map(lambda([u], u=1), listofvars(stud_si_E)), stud_si_E);
v_pure_t_E : subst(map(lambda([u], u=1), listofvars(teach_si_E)), teach_si_E);
eleve_unit_E : 2 * stud_si_E / v_pure_e_E;
teacher_unit_E : 2 * teach_si_E / v_pure_t_E;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Vérification de l'unité de tau</description>
        <answertest>UnitsAbsolute</answertest>
        <sans>eleve_unit_tau</sans>
        <tans>teacher_unit_tau</tans>
        <testoptions>0</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.25</truescore>
        <truepenalty></truepenalty>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt1-0-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-0-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<div style="border-left:4px solid #dc2626;padding:8px 12px;background:#fef2f2;border-radius:4px;margin-bottom:10px;">❌ <strong>Constante de temps :</strong> L'unité est incorrecte (attendue en base SI : {@teacher_unit_tau/2@}).</div>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Vérification de la valeur de tau</description>
        <answertest>UnitsRelative</answertest>
        <sans>ans_tau</sans>
        <tans>ta_tau</tans>
        <testoptions>precision_q</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.75</truescore>
        <truepenalty></truepenalty>
        <truenextnode>3</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left:4px solid #15803d;padding:8px 12px;background:#f0fdf4;border-radius:4px;margin-bottom:10px;">✅ <strong>Constante de temps :</strong> Correcte.</div>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Piège Tau - A-t-il mis E ?</description>
        <answertest>UnitsRelative</answertest>
        <sans>ans_tau</sans>
        <tans>ta_E</tans>
        <testoptions>precision_q</testoptions>
        <quiet>0</quiet>
        <truescoremode>-</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left:4px solid #ca8a04;padding:8px 12px;background:#fefce8;border-radius:4px;margin-bottom:10px;">🔶 <strong>Constante de temps :</strong> Vous avez entré la valeur de la tension \(E\) au lieu de la constante de temps \(\tau\).</div>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<div style="border-left:4px solid #dc2626;padding:8px 12px;background:#fef2f2;border-radius:4px;margin-bottom:10px;">❌ <strong>Constante de temps :</strong> La valeur numérique ou la précision est incorrecte. Mesurez précisément l'abscisse correspondant à 63,2% de la chute de tension.</div>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Vérification de l'unité de E</description>
        <answertest>UnitsAbsolute</answertest>
        <sans>eleve_unit_E</sans>
        <tans>teacher_unit_E</tans>
        <testoptions>0</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.25</truescore>
        <truepenalty></truepenalty>
        <truenextnode>4</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<div style="border-left:4px solid #dc2626;padding:8px 12px;background:#fef2f2;border-radius:4px;margin-bottom:10px;">❌ <strong>Tension :</strong> L'unité est incorrecte (attendue en base SI : {@teacher_unit_E/2@}).</div>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description>Vérification de la valeur de E</description>
        <answertest>UnitsRelative</answertest>
        <sans>ans_E</sans>
        <tans>ta_E</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.75</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left:4px solid #15803d;padding:8px 12px;background:#f0fdf4;border-radius:4px;">✅ <strong>Tension :</strong> Correcte.</div>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>5</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>5</name>
        <description>Piège E - A-t-il mis tau ?</description>
        <answertest>UnitsRelative</answertest>
        <sans>ans_E</sans>
        <tans>ta_tau</tans>
        <testoptions>precision_q</testoptions>
        <quiet>0</quiet>
        <truescoremode>-</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-5-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left:4px solid #ca8a04;padding:8px 12px;background:#fefce8;border-radius:4px;margin-bottom:10px;">🔶 <strong>Tension :</strong> Vous avez entré la valeur de la constante de temps \(\tau\) au lieu de la tension \(E\).</div>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-5-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<div style="border-left:4px solid #dc2626;padding:8px 12px;background:#fef2f2;border-radius:4px;margin-bottom:10px;">❌ <strong>Tension :</strong> La valeur numérique ou la précision est incorrecte. Mesurez l'écart entre le palier supérieur et la ligne du bas avec les curseurs Y.</div>]]></text>
        </falsefeedback>
      </node>
    </prt>
  </question>

</quiz>