<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name><text>Exercice_sans_titre_2026_07_16_18h04</text></name>
    <questiontext format="html">
      <text><![CDATA[

<!-- ENONCE-START -->tracer la courbe {@pa@}<!-- ENONCE-END -->
<div id="ggbApp_2" style="min-height:500px;"></div>
<script>
(function(){
function ggb2_sync(){
  (function(){
    var el=document.querySelector('input[name$="ggb2o1"]');
    if(!el||el.getAttribute('readonly')==='readonly')return;
    var val=api.getValue('correct');
    var sval=(val?'true':'false');
    if(el.value===sval)return;
    el.value=sval;
    el.dispatchEvent(new Event('input',{bubbles:true}));
    el.dispatchEvent(new Event('change',{bubbles:true}));
  })();
}
window.ggb2_sync=ggb2_sync;
function ggb2_start(){
  var params={
    material_id:'aassvdjt',
    width:700,height:500,
    showToolBar:false,
    showAlgebraInput:false,showMenuBar:false,enableRightClick:false,
    appletOnLoad:function(api){
      window['ggbApplet_2']=api;
      try{api.setValue('pa', {#rand(3)+1#});}catch(e){}
      try{api.setValue('pb', {#rand(5)-2#});}catch(e){}
      try{api.setValue('pc', {#rand(5)-2#});}catch(e){}
      try{api.registerObjectUpdateListener('correct', ggb2_sync);}catch(e){}
      ggb2_sync();
    }
  };
  var app=new GGBApplet(params,true);
  app.inject('ggbApp_2');
}
if(typeof GGBApplet==='undefined'){var s=document.createElement('script');s.src='https://www.geogebra.org/apps/deployggb.js';s.onload=ggb2_start;document.head.appendChild(s);}else{ggb2_start();}
})();
</script>
<div style="display:none">
[[input:ggb2o1]][[validation:ggb2o1]]
</div>
<br>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion><text></text></stackversion>
    <questionvariables>
      <text><![CDATA[]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<ol>
  <li>[[feedback:prt1]]</li>
</ol>]]></text>
    </specificfeedback>
    <questionnote format="html">
      <text>GeoGebra Q1</text>
    </questionnote>
    <questiondescription format="html">
      <text>Généré avec StackForge V2 | stackforge-types:geogebra</text>
    </questiondescription>
    <!-- stackforge::v1::eyJodG1sIjoiXG5cbjxzcGFuIGNvbnRlbnRlZGl0YWJsZT1cImZhbHNlXCIgY2xhc3M9XCJxLWNoaXAgY2hpcC1va1wiIGRhdGEtcWlkPVwiMVwiIGRhdGEtdHlwZT1cImdlb2dlYnJhXCIgcm9sZT1cImJ1dHRvblwiIHRhYmluZGV4PVwiMFwiIGFyaWEtbGFiZWw9XCJRdWVzdGlvbiAxIMK3IPCfk5AgR2VvR2VicmEuIENsaXF1ZXIgcG91ciBjb25maWd1cmVyLlwiIHN0eWxlPVwiYmFja2dyb3VuZDogcmdiYSg1NiwgMTE4LCAyOSwgMC4xMzMpOyBib3JkZXItY29sb3I6IHJnYig1NiwgMTE4LCAyOSk7XCIgZHJhZ2dhYmxlPVwidHJ1ZVwiPjxzdmcgY2xhc3M9XCJocy1pY28gY2hpcC1pY29cIiBhcmlhLWhpZGRlbj1cInRydWVcIj48dXNlIGhyZWY9XCIjaWNvLXR5cGUtZ2VvZ2VicmFcIj48L3VzZT48L3N2Zz48c3BhbiBjbGFzcz1cImNoaXAtbGFiZWxcIj5RMSDCtyBnZW9nZWJyYTwvc3Bhbj48c3BhbiBjbGFzcz1cImNoaXAtc3RhdHVzXCIgYXJpYS1sYWJlbD1cIkNvbmZpZ3Vyw6lcIj7inJM8L3NwYW4+PC9zcGFuPlxuPGJyPiIsInF1ZXN0aW9ucyI6eyIxIjp7ImlkIjoxLCJ0eXBlIjoiZ2VvZ2VicmEiLCJiYXJlbWUiOjEsInZhcnMiOiIiLCJxbm90ZSI6Ikdlb0dlYnJhIFExIiwidGV4dEZyYWciOiI8IS0tIEVOT05DRS1TVEFSVCAtLT50cmFjZXIgbGEgY291cmJlIHtAcGFAfTwhLS0gRU5PTkNFLUVORCAtLT5cbjxkaXYgaWQ9XCJnZ2JBcHBfMlwiIHN0eWxlPVwibWluLWhlaWdodDo1MDBweDtcIj48L2Rpdj5cbjxzY3JpcHQ+XG4oZnVuY3Rpb24oKXtcbmZ1bmN0aW9uIGdnYjJfc3luYygpe1xuICAoZnVuY3Rpb24oKXtcbiAgICB2YXIgZWw9ZG9jdW1lbnQucXVlcnlTZWxlY3RvcignaW5wdXRbbmFtZSQ9XCJnZ2IybzFcIl0nKTtcbiAgICBpZighZWx8fGVsLmdldEF0dHJpYnV0ZSgncmVhZG9ubHknKT09PSdyZWFkb25seScpcmV0dXJuO1xuICAgIHZhciB2YWw9YXBpLmdldFZhbHVlKCdjb3JyZWN0Jyk7XG4gICAgdmFyIHN2YWw9KHZhbD8ndHJ1ZSc6J2ZhbHNlJyk7XG4gICAgaWYoZWwudmFsdWU9PT1zdmFsKXJldHVybjtcbiAgICBlbC52YWx1ZT1zdmFsO1xuICAgIGVsLmRpc3BhdGNoRXZlbnQobmV3IEV2ZW50KCdpbnB1dCcse2J1YmJsZXM6dHJ1ZX0pKTtcbiAgICBlbC5kaXNwYXRjaEV2ZW50KG5ldyBFdmVudCgnY2hhbmdlJyx7YnViYmxlczp0cnVlfSkpO1xuICB9KSgpO1xufVxud2luZG93LmdnYjJfc3luYz1nZ2IyX3N5bmM7XG5mdW5jdGlvbiBnZ2IyX3N0YXJ0KCl7XG4gIHZhciBwYXJhbXM9e1xuICAgIG1hdGVyaWFsX2lkOidhYXNzdmRqdCcsXG4gICAgd2lkdGg6NzAwLGhlaWdodDo1MDAsXG4gICAgc2hvd1Rvb2xCYXI6ZmFsc2UsXG4gICAgc2hvd0FsZ2VicmFJbnB1dDpmYWxzZSxzaG93TWVudUJhcjpmYWxzZSxlbmFibGVSaWdodENsaWNrOmZhbHNlLFxuICAgIGFwcGxldE9uTG9hZDpmdW5jdGlvbihhcGkpe1xuICAgICAgd2luZG93WydnZ2JBcHBsZXRfMiddPWFwaTtcbiAgICAgIHRyeXthcGkuc2V0VmFs
dWUoJ3BhJywgeyNyYW5kKDMpKzEjfSk7fWNhdGNoKGUpe31cbiAgICAgIHRyeXthcGkuc2V0VmFsdWUoJ3BiJywgeyNyYW5kKDUpLTIjfSk7fWNhdGNoKGUpe31cbiAgICAgIHRyeXthcGkuc2V0VmFsdWUoJ3BjJywgeyNyYW5kKDUpLTIjfSk7fWNhdGNoKGUpe31cbiAgICAgIHRyeXthcGkucmVnaXN0ZXJPYmplY3RVcGRhdGVMaXN0ZW5lcignY29ycmVjdCcsIGdnYjJfc3luYyk7fWNhdGNoKGUpe31cbiAgICAgIGdnYjJfc3luYygpO1xuICAgIH1cbiAgfTtcbiAgdmFyIGFwcD1uZXcgR0dCQXBwbGV0KHBhcmFtcyx0cnVlKTtcbiAgYXBwLmluamVjdCgnZ2diQXBwXzInKTtcbn1cbmlmKHR5cGVvZiBHR0JBcHBsZXQ9PT0ndW5kZWZpbmVkJyl7dmFyIHM9ZG9jdW1lbnQuY3JlYXRlRWxlbWVudCgnc2NyaXB0Jyk7cy5zcmM9J2h0dHBzOi8vd3d3Lmdlb2dlYnJhLm9yZy9hcHBzL2RlcGxveWdnYi5qcyc7cy5vbmxvYWQ9Z2diMl9zdGFydDtkb2N1bWVudC5oZWFkLmFwcGVuZENoaWxkKHMpO31lbHNle2dnYjJfc3RhcnQoKTt9XG59KSgpO1xuPC9zY3JpcHQ+XG48ZGl2IHN0eWxlPVwiZGlzcGxheTpub25lXCI+XG5bW2lucHV0OmdnYjJvMV1dW1t2YWxpZGF0aW9uOmdnYjJvMV1dXG48L2Rpdj4iLCJpbnB1dFhNTCI6IiAgICA8aW5wdXQ+XG4gICAgICA8bmFtZT5nZ2IybzE8L25hbWU+XG4gICAgICA8dHlwZT5ib29sZWFuPC90eXBlPlxuICAgICAgPHRhbnM+dHJ1ZTwvdGFucz5cbiAgICAgIDxib3hzaXplPjEwPC9ib3hzaXplPlxuICAgICAgPHN0cmljdHN5bnRheD4xPC9zdHJpY3RzeW50YXg+XG4gICAgICA8aW5zZXJ0c3RhcnM+MDwvaW5zZXJ0c3RhcnM+XG4gICAgICA8c3ludGF4aGludD48L3N5bnRheGhpbnQ+XG4gICAgICA8c3ludGF4YXR0cmlidXRlPjA8L3N5bnRheGF0dHJpYnV0ZT5cbiAgICAgIDxmb3JiaWR3b3Jkcz48L2ZvcmJpZHdvcmRzPlxuICAgICAgPGFsbG93d29yZHM+PC9hbGxvd3dvcmRzPlxuICAgICAgPGZvcmJpZGZsb2F0PjA8L2ZvcmJpZGZsb2F0PlxuICAgICAgPHJlcXVpcmVsb3dlc3R0ZXJtcz4wPC9yZXF1aXJlbG93ZXN0dGVybXM+XG4gICAgICA8Y2hlY2thbnN3ZXJ0eXBlPjA8L2NoZWNrYW5zd2VydHlwZT5cbiAgICAgIDxtdXN0dmVyaWZ5PjA8L211c3R2ZXJpZnk+XG4gICAgICA8c2hvd3ZhbGlkYXRpb24+MDwvc2hvd3ZhbGlkYXRpb24+XG4gICAgICA8b3B0aW9ucz48L29wdGlvbnM+XG4gICAgPC9pbnB1dD4iLCJwcnRYTUwiOiIgICAgPHBydD5cbiAgICAgIDxuYW1lPnBydDE8L25hbWU+XG48dmFsdWU+MS4wMDAwMDAwPC92YWx1ZT5cbiAgICAgIDxhdXRvc2ltcGxpZnk+MTwvYXV0b3NpbXBsaWZ5PlxuICAgICAgPGZlZWRiYWNrc3R5bGU+MjwvZmVlZGJhY2tzdHlsZT5cbiAgICAgIDxmZWVkYmFja3ZhcmlhYmxlcz5cbiAgICAgICAgPHRleHQ+PCFbQ0RBVEFbZ2diX29rXzE6IFtpZiBnZ2IxbzEgPSB0cnVlIHRoZW4gMSBlbHNlIDBdO1xuZ2diX25fb2tfMTogYXBwbHkoXCIrXCIsIGdnYl9va18x
KTtcbmdnYl9zY18xOiBmbG9hdChnZ2Jfbl9va18xIC8gMSk7XG5nZ2JfcGN0XzE6IHJvdW5kKGdnYl9zY18xICogMTAwKTtdXT48L3RleHQ+XG4gICAgICA8L2ZlZWRiYWNrdmFyaWFibGVzPlxuICAgICAgPG5vZGU+XG48bmFtZT4wPC9uYW1lPlxuPGRlc2NyaXB0aW9uPjwvZGVzY3JpcHRpb24+XG48YW5zd2VydGVzdD5BbGdFcXVpdjwvYW5zd2VydGVzdD5cbjxzYW5zPmdnYl9zY18xPC9zYW5zPlxuPHRhbnM+MTwvdGFucz5cbjx0ZXN0b3B0aW9ucz48L3Rlc3RvcHRpb25zPlxuPHF1aWV0PjA8L3F1aWV0PlxuPHRydWVzY29yZW1vZGU+PTwvdHJ1ZXNjb3JlbW9kZT5cbjx0cnVlc2NvcmU+MTwvdHJ1ZXNjb3JlPlxuPHRydWVwZW5hbHR5PjA8L3RydWVwZW5hbHR5PlxuPHRydWVuZXh0bm9kZT4tMTwvdHJ1ZW5leHRub2RlPlxuPHRydWVhbnN3ZXJub3RlPlBSVDEtMS1UPC90cnVlYW5zd2Vybm90ZT5cbjx0cnVlZmVlZGJhY2sgZm9ybWF0PVwiaHRtbFwiPlxuICA8dGV4dD48IVtDREFUQVs8cD7inIUgPHN0cm9uZz5QYXJmYWl0ICE8L3N0cm9uZz48L3A+XV0+PC90ZXh0PlxuPC90cnVlZmVlZGJhY2s+XG48ZmFsc2VzY29yZW1vZGU+PTwvZmFsc2VzY29yZW1vZGU+XG48ZmFsc2VzY29yZT5nZ2Jfc2NfMTwvZmFsc2VzY29yZT5cbjxmYWxzZXBlbmFsdHk+MDwvZmFsc2VwZW5hbHR5PlxuPGZhbHNlbmV4dG5vZGU+LTE8L2ZhbHNlbmV4dG5vZGU+XG48ZmFsc2VhbnN3ZXJub3RlPlBSVDEtMS1GPC9mYWxzZWFuc3dlcm5vdGU+XG48ZmFsc2VmZWVkYmFjayBmb3JtYXQ9XCJodG1sXCI+XG4gIDx0ZXh0PjwhW0NEQVRBWzxwPuKdjCBTY29yZSA6IHtAZ2diX3BjdF8xQH0gJS4gU29sdXRpb24gYXR0ZW5kdWUgOjwvcD48dWw+PGxpPjxjb2RlPmNvcnJlY3Q8L2NvZGU+ID0gdHJ1ZTwvbGk+PC91bD5dXT48L3RleHQ+XG48L2ZhbHNlZmVlZGJhY2s+XG4gICAgICA8L25vZGU+XG4gICAgPC9wcnQ+IiwiZ2VuZXJhbEZlZWRiYWNrIjoiIiwiZmVlZGJhY2tSZWYiOiJbW2ZlZWRiYWNrOnBydDFdXSIsInBydCI6eyJtZXRhIjp7Im5hbWUiOiJwcnQxIiwidmFsdWUiOiIxLjAwMDAwMDAiLCJhdXRvc2ltcGxpZnkiOiIxIiwiZmVlZGJhY2tzdHlsZSI6IjIiLCJmZWVkYmFja3ZhcmlhYmxlcyI6ImdnYl9va18yOiBbaWYgZ2diMm8xID0gdHJ1ZSB0aGVuIDEgZWxzZSAwXTtcbmdnYl9uX29rXzI6IGFwcGx5KFwiK1wiLCBnZ2Jfb2tfMik7XG5nZ2Jfc2NfMjogZmxvYXQoZ2diX25fb2tfMiAvIDEpO1xuZ2diX3BjdF8yOiByb3VuZChnZ2Jfc2NfMiAqIDEwMCk7In0sIm5vZGVzIjpbeyJuYW1lIjoiMCIsImRlc2NyaXB0aW9uIjoiIiwiYW5zd2VydGVzdCI6IkFsZ0VxdWl2Iiwic2FucyI6ImdnYl9zY18xIiwidGFucyI6IjEiLCJ0ZXN0b3B0aW9ucyI6IiIsInF1aWV0IjoiMCIsInRydWVzY29yZW1vZGUiOiI9IiwidHJ1ZXNjb3JlIjoiMSIsInRydWVwZW5hbHR5IjoiMCIsInRydWVuZXh0bm9kZSI6Ii0xIiwidHJ1ZWFuc3dlcm5vdGUiOiJQUlQxLTEtVCIs
InRydWVmZWVkYmFjayI6IjxwPuKchSA8c3Ryb25nPlBhcmZhaXQgITwvc3Ryb25nPjwvcD4iLCJmYWxzZXNjb3JlbW9kZSI6Ij0iLCJmYWxzZXNjb3JlIjoiZ2diX3NjXzIiLCJmYWxzZXBlbmFsdHkiOiIwIiwiZmFsc2VuZXh0bm9kZSI6Ii0xIiwiZmFsc2VhbnN3ZXJub3RlIjoiUFJUMS0xLUYiLCJmYWxzZWZlZWRiYWNrIjoiPHA+4p2MIFNjb3JlIDoge0BnZ2JfcGN0XzFAfSAlLiBTb2x1dGlvbiBhdHRlbmR1ZSA6PC9wPjx1bD48bGk+PGNvZGU+Y29ycmVjdDwvY29kZT4gPSB0cnVlPC9saT48L3VsPiJ9XX0sInN0YXRlIjp7InR5cGUiOiJnZW9nZWJyYSIsImJhcmVtZSI6IjEiLCJ0ZXh0IjoidHJhY2VyIGxhIGNvdXJiZSB7QHBhQH0iLCJtYXRlcmlhbElkIjoiYWFzc3ZkanQiLCJ3aWR0aCI6IjcwMCIsImhlaWdodCI6IjUwMCIsInNob3dUb29sYmFyIjpmYWxzZSwiaW5wdXRzIjpbeyJnZ2JOYW1lIjoicGEiLCJleHByIjoicmFuZCgzKSsxIn0seyJnZ2JOYW1lIjoicGIiLCJleHByIjoicmFuZCg1KS0yIn0seyJnZ2JOYW1lIjoicGMiLCJleHByIjoicmFuZCg1KS0yIn1dLCJvdXRwdXRzIjpbeyJnZ2JOYW1lIjoiY29ycmVjdCIsInR5cGUiOiJib29sZWFuIiwidGFucyI6InRydWUiLCJ0b2wiOiIwLjAxIn1dLCJmYkdlbiI6IiJ9fX0sIm5leHRRaWQiOjIsInNoYXJlZFZhcnMiOltdLCJxdWl6TmFtZSI6IkV4ZXJjaWNlIHNhbnMgdGl0cmUifQ== -->
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html"><text></text></prtcorrect>
    <prtpartiallycorrect format="html"><text></text></prtpartiallycorrect>
    <prtincorrect format="html"><text></text></prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ggb2o1</name>
      <type>boolean</type>
      <tans>true</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>

    <prt>
      <name>prt1</name>
<value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>2</feedbackstyle>
      <feedbackvariables>
        <text><![CDATA[ggb_ok_1: [if ggb1o1 = true then 1 else 0];
ggb_n_ok_1: apply("+", ggb_ok_1);
ggb_sc_1: float(ggb_n_ok_1 / 1);
ggb_pct_1: round(ggb_sc_1 * 100);]]></text>
      </feedbackvariables>
      <node>
<name>0</name>
<description></description>
<answertest>AlgEquiv</answertest>
<sans>ggb_sc_1</sans>
<tans>1</tans>
<testoptions></testoptions>
<quiet>0</quiet>
<truescoremode>=</truescoremode>
<truescore>1</truescore>
<truepenalty>0</truepenalty>
<truenextnode>-1</truenextnode>
<trueanswernote>PRT1-1-T</trueanswernote>
<truefeedback format="html">
  <text><![CDATA[<p>✅ <strong>Parfait !</strong></p>]]></text>
</truefeedback>
<falsescoremode>=</falsescoremode>
<falsescore>ggb_sc_1</falsescore>
<falsepenalty>0</falsepenalty>
<falsenextnode>-1</falsenextnode>
<falseanswernote>PRT1-1-F</falseanswernote>
<falsefeedback format="html">
  <text><![CDATA[<p>❌ Score : {@ggb_pct_1@} %. Solution attendue :</p><ul><li><code>correct</code> = true</li></ul>]]></text>
</falsefeedback>
      </node>
    </prt>

  <tags>
    <tag><text>stack</text></tag>
    <tag><text>geogebra</text></tag>
  </tags>
  </question>
</quiz>