<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 396745  -->
  <question type="stack">
    <name>
      <text>G Logique — Déduire des cases manquantes</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<div style="background:#1e3a5f;border-left:5px solid #2563eb;border-radius:0 8px 8px 0;padding:10px 16px;margin-bottom:15px;display:flex;align-items:center;gap:10px;flex-wrap:wrap;">
        <strong style="font-weight:800;color:#fff;font-size:.95rem;">Logique — Cases manquantes</strong>
        <span style="background:#2563eb;color:#fff;padding:2px 9px;border-radius:20px;font-size:.78rem;font-weight:700;">/ 1 pt</span>
        <span style="background:#ffffff;color:#1e3a5f;border:1px solid #1e3a5f;padding:2px 9px;border-radius:20px;font-size:.75rem;font-weight:600;">📐 Mathématiques</span>
      </div>

      <p>Deux cases manquent dans la table de vérité de \({@expr@}\). Déduisez leur valeur :</p>

      <div style="overflow-x:auto;margin:20px 0;">
        <table style="margin:0 auto;border-collapse:collapse;border:2px solid #334155;background:#fff;font-family:monospace;font-size:1.1rem;">
          <thead>
            <tr style="background:#f1f5f9;">
              <th style="border:1px solid #94a3b8;padding:10px 30px;">P</th>
              <th style="border:1px solid #94a3b8;padding:10px 30px;">Q</th>
              <th style="border:1px solid #94a3b8;padding:10px 30px;background:#fef9c3;">{@expr@}</th>
            </tr>
          </thead>
          <tbody>
            <tr>
              <td style="border:1px solid #94a3b8;padding:10px 30px;text-align:center;">0</td>
              <td style="border:1px solid #94a3b8;padding:10px 30px;text-align:center;">0</td>
              <td style="border:1px solid #94a3b8;padding:10px 30px;text-align:center;font-weight:bold;color:#15803d;font-size:1.2rem;">{@v0@}</td>
            </tr>
            <tr style="background:#f8fafc;">
              <td style="border:1px solid #94a3b8;padding:10px 30px;text-align:center;">0</td>
              <td style="border:1px solid #94a3b8;padding:10px 30px;text-align:center;">1</td>
              <td style="border:1px solid #94a3b8;padding:10px 30px;text-align:center;background:#fefce8;">[[input:ans1]] [[validation:ans1]]</td>
            </tr>
            <tr>
              <td style="border:1px solid #94a3b8;padding:10px 30px;text-align:center;">1</td>
              <td style="border:1px solid #94a3b8;padding:10px 30px;text-align:center;">0</td>
              <td style="border:1px solid #94a3b8;padding:10px 30px;text-align:center;background:#fefce8;">[[input:ans2]] [[validation:ans2]]</td>
            </tr>
            <tr style="background:#f8fafc;">
              <td style="border:1px solid #94a3b8;padding:10px 30px;text-align:center;">1</td>
              <td style="border:1px solid #94a3b8;padding:10px 30px;text-align:center;">1</td>
              <td style="border:1px solid #94a3b8;padding:10px 30px;text-align:center;font-weight:bold;color:#15803d;font-size:1.2rem;">{@v3@}</td>
            </tr>
          </tbody>
        </table>
      </div>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<div style="border-left:4px solid #15803d;padding:10px 14px;background:#f0fdf4;border-radius:4px;margin:4px 0;">
        <p>✅ Table complète : P=0,Q=1 → <strong>{@v1@}</strong> | P=1,Q=0 → <strong>{@v2@}</strong></p>
      </div>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.2</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>expr1 : P and Q$ expr2 : P or Q$ expr3 : P implies Q$ expr4 : not(P implies Q)$ expr5 : not(P and Q)$ expr6 : not(P or Q)$ expr7 : (P implies Q) and (Q implies P)$ expr8 : (P or Q) and not(P and Q)$ expr9 : P implies not(Q)$ expr10 : not(P) implies Q$ 
liste_expr : [expr1, expr2, expr3, expr4, expr5, expr6, expr7, expr8, expr9, expr10]$ idx : rand(10) + 1$ expr : liste_expr[idx]$ 
T : truth_table(boolean_form(expr))$ Rows : args(T)$ v0 : if last(Rows[2]) then 1 else 0$ v1 : if last(Rows[3]) then 1 else 0$ v2 : if last(Rows[4]) then 1 else 0$ v3 : if last(Rows[5]) then 1 else 0$ 
/* Anti-erreurs courantes : inverser les deux cases */
ta_swap : [v2, v1]</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<ol><li>[[feedback:prt1]]</li></ol>]]></text>
    </specificfeedback>
    <questionnote format="html">
      <text>{@expr@} | P=0,Q=1→{@v1@} | P=1,Q=0→{@v2@}</text>
    </questionnote>
    <questiondescription format="html">
      <text>Généré avec HéStack V2 | hestack-types:logique-prop-cases</text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>symbol</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>v1</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>0</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>v2</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>0</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>2</feedbackstyle>
      <feedbackvariables>
        <text>student_ans : [ans1, ans2]
ta_ans : [v1, v2]</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Cases correctes</description>
        <answertest>AlgEquiv</answertest>
        <sans>student_ans</sans>
        <tans>ta_ans</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-0-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left:4px solid #15803d;padding:10px 14px;background:#f0fdf4;border-radius:4px;margin:4px 0;"><p>✅ <strong>Exact !</strong> Vous avez correctement déduit les deux valeurs manquantes.</p></div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-0-F</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Cases inversées</description>
        <answertest>AlgEquiv</answertest>
        <sans>student_ans</sans>
        <tans>ta_swap</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>1</trueanswernote>
        <truefeedback format="html">
          <text></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>1</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote></falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
  </question>

</quiz>