<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 396299  -->
  <question type="stack">
    <name>
      <text>Combinatoire : Coefficient binomial C(n,k)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<div style="background: #475569; border-left: 5px solid #334155; border-radius: 0 8px 8px 0; padding: 10px 16px; margin-bottom: 12px; display: flex; align-items: center; gap: 10px; flex-wrap: wrap;"><strong style="font-weight: 800; color: #fff; font-size: .95rem;">Combinatoire — Coefficient binomial</strong> <span style="background: #334155; color: #fff; padding: 2px 9px; border-radius: 20px; font-size: .78rem; font-weight: bold;">/ 1 pt</span></div>
<p>Calculer la valeur exacte du coefficient binomial suivant :</p>
<div style="text-align: center; font-size: 1.2rem; margin: 20px 0;">
\(C\left({@n@} \; ; \; {@k@}\right)\)
</div>
<p><em>Rappel de la formule : \(C(n, k) = \dfrac{n!}{k!(n-k)!}\)</em></p>
<div style="margin: 15px 0 10px 0;">\(C({@n@}, {@k@}) = \) [[input:ans_ck]] [[validation:ans_ck]]</div>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<div style="margin-top: 20px; padding: 15px; background: #f8fafc; border: 1px solid #cbd5e1; border-radius: 8px;">
<div style="font-weight: bold; margin-bottom: 10px;">🔑 Correction détaillée</div>
<div style="font-size: .9rem;">
On applique la formule du coefficient binomial :<br><br>
\[
C({@n@}, {@k@}) = \frac{{@n@}!}{{@k@}! \times ({@n@} - {@k@})!} = \frac{{@n@}!}{{@k@}! \times {@nmk@}!}
\]<br>
En simplifiant les factorielles (on peut simplifier le numérateur et le dénominateur par {@nmk@}!) :<br><br>
\[
C({@n@}, {@k@}) = \frac{{@n@} \times ({@n@}-1) \times \dots \times ({@nmk@}+1)}{{@k@}!} = \frac{@num_simp@}{{@den_simp@}} = {@ans_correct@}
\]
</div>
</div>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* Fonction utilitaire */
ri(a,b) := a + floor(rand((b-a+1)));

/* Génération des paramètres */
n: ri(6, 10);
k: ri(2, n-2);
nmk: n - k;

/* Réponse attendue */
ans_correct: binomial(n, k);

/* Pour l'affichage dans la correction */
num_simp: apply("*", makelist(i, i, nmk+1, n));
den_simp: k!;

/* Erreurs classiques */
/* Erreur 1 : A calculé un arrangement A(n,k) = n!/(n-k)! (a oublié le k! au dénominateur) */
err_arrang: n! / nmk!;
/* Erreur 2 : A oublié le (n-k)! au dénominateur */
err_no_nmk: n! / k!;
/* Erreur 3 : A inversé la fraction */
err_inv: (k! * nmk!) / n!;]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<div>[[feedback:prt1]]</div>]]></text>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>C({@n@},{@k@}) = {@ans_correct@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p>Généré avec HéStack V2 | hestack-types:combinatoire-coefficient-binomial</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans_ck</name>
      <type>numerical</type>
      <tans>ans_correct</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Réponse correcte</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans_ck</sans>
        <tans>ans_correct</tans>
        <testoptions>0.001</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-OK</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #15803d; padding: 10px 14px; background: #f0fdf4; border-radius: 4px;">✅ <strong>Parfait !</strong> Le coefficient binomial est correct.</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>PRT-NOK</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Erreur : A calculé un arrangement A(n,k)</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans_ck</sans>
        <tans>err_arrang</tans>
        <testoptions>0.001</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-ARRANG</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Confusion Combinaison / Arrangement !</strong> Vous avez calculé \(A({@n@}, {@k@}) = \frac{{@n@}!}{({@n@}-{@k@})!}\). Pour un coefficient binomial (une combinaison), l'ordre ne compte pas, il faut diviser en plus par {@k@}!.</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>PRT-ERR-NOT-ARRANG</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Erreur : A oublié le (n-k)! au dénominateur</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans_ck</sans>
        <tans>err_no_nmk</tans>
        <testoptions>0.001</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-NO-NMK</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Facteur manquant !</strong> Vous avez calculé \(\frac{{@n@}!}{{@k@}!}\). Il manque le facteur \(({@n@}-{@k@})!\) (soit {@nmk@}!) au dénominateur.</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>PRT-ERR-NOT-NO-NMK</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Erreur : A inversé la fraction</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans_ck</sans>
        <tans>err_inv</tans>
        <testoptions>0.001</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-INV</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Fraction inversée !</strong> La factorielle de \(n\) (le plus grand nombre) doit être au <strong>numérateur</strong>, et le produit des autres factorielles au <strong>dénominateur</strong>.</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>PRT-ERR-CALC</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #dc2626; padding: 10px 14px; background: #fff0f0; border-radius: 4px;">❌ <strong>Incorrect.</strong> Appliquez soigneusement la formule :<br>\(C(n, k) = \dfrac{n!}{k!(n-k)!}\)<br><em>(Pensez à simplifier la fraction avant de calculer pour éviter les erreurs de calculatrice avec les grands nombres).</em></div>]]></text>
        </falsefeedback>
      </node>
    </prt>
  </question>

</quiz>