<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 396308  -->
  <question type="stack">
    <name>
      <text>Probabilités : Variance d'une loi binomiale V(X) = np(1-p)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<div style="background: #475569; border-left: 5px solid #334155; border-radius: 0 8px 8px 0; padding: 10px 16px; margin-bottom: 12px; display: flex; align-items: center; gap: 10px; flex-wrap: wrap;"><strong style="font-weight: 800; color: #fff; font-size: .95rem;">Probabilités discrètes — Variance</strong> <span style="background: #334155; color: #fff; padding: 2px 9px; border-radius: 20px; font-size: .78rem; font-weight: bold;">/ 1 pt</span></div>
<p>Soit \(X\) une variable aléatoire qui suit la loi binomiale \(\mathcal{B}\left({@n@} \; ; \; {@p@}\right)\).</p>
<p><strong>1.</strong> Calculer la variance de la variable aléatoire \(X\), c'est-à-dire calculer \(V(X)\).</p>
<p><em>Le résultat devra être donné sous forme de fraction irréductible.</em></p>
<div style="margin: 15px 0 10px 0;">\(V(X) = \) [[input:ans_vx]] [[validation:ans_vx]]</div>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<div style="margin-top: 20px; padding: 15px; background: #f8fafc; border: 1px solid #cbd5e1; border-radius: 8px;">
<div style="font-weight: bold; margin-bottom: 10px;">🔑 Correction détaillée</div>
<div style="font-size: .9rem;">
Si \(X \sim \mathcal{B}(n, p)\), la variance de la loi binomiale est donnée par la formule :<br><br>
\[
V(X) = n \times p \times (1 - p)
\]<br>
<strong>Application numérique :</strong><br>
Ici, \(n = {@n@}\), \(p = {@p@}\) et donc \(1 - p = {@q@}\).<br><br>
\[
V(X) = {@n@} \times {@p@} \times {@q@} = {@ans_correct@}
\]
</div>
</div>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>/* Fonction utilitaire CORRECTE pour Maxima */
/* rand(n) dans Maxima renvoie un entier entre 1 et n */
ri(a,b) := a + rand(b - a + 1) - 1;

/* Génération des paramètres */
n: ri(10, 50);
p_pool: [1/2, 1/3, 2/3, 1/4, 3/4, 1/5, 2/5, 3/5, 4/5];
p: p_pool[rand(length(p_pool))];
q: 1 - p;

/* Réponse attendue */
ans_correct: n * p * q;

/* Erreurs classiques */
err_ex: n * p;
err_p_q: p * q;
err_np2: n * p * p;
err_nq: n * q;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<div>[[feedback:prt1]]</div>]]></text>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>B({@n@}, {@p@}), V(X) = {@ans_correct@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p>Généré avec HéStack V2 | hestack-types:probabilites-binomiale-variance</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans_vx</name>
      <type>algebraic</type>
      <tans>ans_correct</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Réponse correcte</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_vx</sans>
        <tans>ans_correct</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-OK</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #15803d; padding: 10px 14px; background: #f0fdf4; border-radius: 4px;">✅ <strong>Parfait !</strong> La variance est correcte et est bien donnée sous forme de fraction irréductible.</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>PRT-NOK</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Erreur : A calculé l'espérance E(X)</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_vx</sans>
        <tans>err_ex</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-EX</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Confusion Espérance / Variance !</strong> Vous avez calculé l'espérance \(E(X) = n \times p\). Pour la variance, il faut ajouter le facteur lié aux échecs, c'est-à-dire multiplier par \((1-p)\).</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>PRT-ERR-NOT-EX</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Erreur : A oublié le facteur n</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_vx</sans>
        <tans>err_p_q</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-NO-N</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Il manque le nombre d'épreuves !</strong> Vous avez calculé \(p(1-p)\). La variance dépend du nombre de répétitions \(n\). Il faut multiplier votre résultat par {@n@}.</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>PRT-ERR-NOT-NO-N</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Erreur : A calculé n*p^2</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_vx</sans>
        <tans>err_np2</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-NP2</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Erreur sur le deuxième facteur !</strong> Vous avez calculé \(n \times p^2\). La formule de la variance est \(n \times p \times (1-p)\), on multiplie par la probabilité d'échec \((1-p)\), pas par \(p\).</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>1</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>PRT-ERR-NO-NP2</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
    </prt>
  </question>

</quiz>