<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 396322  -->
  <question type="stack">
    <name>
      <text>Probabilités : Formule de l'union P(A∪B)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<div style="background: #475569; border-left: 5px solid #334155; border-radius: 0 8px 8px 0; padding: 10px 16px; margin-bottom: 12px; display: flex; align-items: center; gap: 10px; flex-wrap: wrap;"><strong style="font-weight: 800; color: #fff; font-size: .95rem;">Probabilités discrètes — Union d'événements</strong> <span style="background: #334155; color: #fff; padding: 2px 9px; border-radius: 20px; font-size: .78rem; font-weight: bold;">/ 1 pt</span></div>
<p>Soient \(A\) et \(B\) deux événements d'un univers muni d'une probabilité \(P\) tels que :</p>
<div style="display: flex; gap: 30px; justify-content: center; margin: 15px 0; flex-wrap: wrap;">
<div>\(P(A) = {@pA@}\)</div>
<div>\(P(B) = {@pB@}\)</div>
<div>\(P(A \cap B) = {@p_inter@}\)</div>
</div>
<p><strong>1.</strong> Calculer la probabilité de l'événement \(A \cup B\).</p>
<p><em>Le résultat devra être donné sous forme de fraction irréductible.</em></p>
<div style="margin: 15px 0 10px 0;">\(P(A \cup B) = \) [[input:ans_union]] [[validation:ans_union]]</div>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<div style="margin-top: 20px; padding: 15px; background: #f8fafc; border: 1px solid #cbd5e1; border-radius: 8px;">
<div style="font-weight: bold; margin-bottom: 10px;">🔑 Correction détaillée</div>
<div style="font-size: .9rem;">
Pour calculer la probabilité de l'union de deux événements, on utilise la formule :<br><br>
\[
P(A \cup B) = P(A) + P(B) - P(A \cap B)
\]<br>
<strong>Application numérique :</strong><br>
En remplaçant par les valeurs de l'énoncé, on obtient :<br><br>
\[
P(A \cup B) = {@pA@} + {@pB@} - {@p_inter@}
\]
\[
P(A \cup B) = {@ans_correct@}
\]
</div>
</div>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>/* Fonction utilitaire CORRECTE pour Maxima */
ri(a,b) := a + rand(b - a + 1) - 1;

/* Génération sécurisée par pool de triplets valides */
/* Format : [P(A), P(B), P(A inter B)] */
triplets: [
    [1/2, 1/2, 1/5],
    [2/3, 1/2, 1/4],
    [3/4, 1/3, 1/6],
    [4/5, 3/5, 2/5],
    [1/2, 2/3, 1/3],
    [3/5, 1/2, 1/10],
    [5/6, 1/2, 1/3],
    [3/4, 3/4, 1/2],
    [2/5, 3/4, 1/5],
    [1/3, 4/5, 1/10]
];

chosen: triplets[rand(length(triplets))];
pA: chosen[1];
pB: chosen[2];
p_inter: chosen[3];

/* Réponse attendue */
ans_correct: pA + pB - p_inter;

/* Erreurs classiques */
/* Erreur 1 : A fait une addition simple P(A) + P(B) */
err_add: pA + pB;
/* Erreur 2 : Confusion avec l'indépendance P(A) * P(B) */
err_prod: pA * pB;
/* Erreur 3 : A donné l'intersection au lieu de l'union */
err_inter: p_inter;
/* Erreur 4 : A fait une soustraction P(A) - P(B) */
err_sous: pA - pB;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<div>[[feedback:prt1]]</div>]]></text>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>P(A)={@pA@}, P(B)={@pB@}, P(A inter B)={@p_inter@}, P(A union B)={@ans_correct@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p>Généré avec HéStack V2 | hestack-types:probabilites-union</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans_union</name>
      <type>algebraic</type>
      <tans>ans_correct</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Réponse correcte</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_union</sans>
        <tans>ans_correct</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-OK</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #15803d; padding: 10px 14px; background: #f0fdf4; border-radius: 4px;">✅ <strong>Parfait !</strong> La probabilité de l'union est correcte et est bien donnée sous forme de fraction irréductible.</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>PRT-NOK</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Erreur : A fait une addition simple P(A) + P(B)</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_union</sans>
        <tans>err_add</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-ADD</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Oubli de l'intersection !</strong> Vous avez calculé \(P(A) + P(B)\). En faisant cela, vous comptez deux fois l'intersection \(A \cap B\). Il faut obligatoirement retrancher \(P(A \cap B)\).</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>PRT-ERR-NOT-ADD</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Erreur : A multiplié P(A) * P(B)</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_union</sans>
        <tans>err_prod</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-PROD</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Confusion de formule !</strong> Vous avez multiplié les probabilités \(P(A) \times P(B)\). Cette formule ne s'utilise que pour calculer une <strong>intersection</strong> dans le cas très précis d'événements <strong>indépendants</strong>. Ici on cherche l'union.</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>PRT-ERR-NOT-PROD</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Erreur : A donné P(A inter B)</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_union</sans>
        <tans>err_inter</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-INTER</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Confusion Union / Intersection !</strong> Vous avez donné la valeur de \(P(A \cap B)\) (l'intersection). On vous demande de calculer \(P(A \cup B)\) (l'union).</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>4</falsenextnode>
        <falseanswernote>PRT-ERR-NOT-INTER</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description>Erreur : A fait une soustraction P(A) - P(B)</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_union</sans>
        <tans>err_sous</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-SOUS</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Opération erronée !</strong> Vous avez soustrait \(P(B)\) à \(P(A)\). La formule de l'union nécessite une <strong>addition</strong> des probabilités de \(A\) et \(B\), suivie d'une <strong>soustraction</strong> de leur intersection.</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>PRT-ERR-CALC</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #dc2626; padding: 10px 14px; background: #fff0f0; border-radius: 4px;">❌ <strong>Incorrect.</strong> Appliquez la formule du cours :<br>\(P(A \cup B) = P(A) + P(B) - P(A \cap B)\)<br><em>(Attention à bien gérer les signes et les fractions lors du calcul).</em></div>]]></text>
        </falsefeedback>
      </node>
    </prt>
  </question>

</quiz>