<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 396304  -->
  <question type="stack">
    <name>
      <text>Probabilités : Loi binomiale P(X = k)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<div style="background: #475569; border-left: 5px solid #334155; border-radius: 0 8px 8px 0; padding: 10px 16px; margin-bottom: 12px; display: flex; align-items: center; gap: 10px; flex-wrap: wrap;"><strong style="font-weight: 800; color: #fff; font-size: .95rem;">Probabilités discrètes — Loi binomiale</strong> <span style="background: #334155; color: #fff; padding: 2px 9px; border-radius: 20px; font-size: .78rem; font-weight: bold;">/ 1 pt</span></div>
<p>Soit \(X\) une variable aléatoire qui suit la loi binomiale \(\mathcal{B}\left({@n@} \; ; \; {@p@}\right)\).</p>
<p><strong>1.</strong> Calculer la probabilité d'obtenir exactement {@k@} succès, c'est-à-dire calculer \(P(X = {@k@})\).</p>
<p><em>Le résultat devra être donné sous forme de fraction irréductible.</em></p>
<div style="margin: 15px 0 10px 0;">\(P(X = {@k@}) = \) [[input:ans_prob]] [[validation:ans_prob]]</div>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<div style="margin-top: 20px; padding: 15px; background: #f8fafc; border: 1px solid #cbd5e1; border-radius: 8px;">
<div style="font-weight: bold; margin-bottom: 10px;">🔑 Correction détaillée</div>
<div style="font-size: .9rem;">
Si \(X \sim \mathcal{B}(n, p)\), la probabilité d'obtenir exactement \(k\) succès est donnée par la formule :<br><br>
\[
P(X = k) = \binom{n}{k} \times p^k \times (1-p)^{n-k}
\]<br>
<strong>Application numérique :</strong><br>
Ici, \(n = {@n@}\), \(p = {@p@}\) et \(k = {@k@}\). On a donc \(1-p = {@q@}\).<br><br>
\[
P(X = {@k@}) = \binom{{@n@}}{{@k@}} \times \left({@p@}\right)^{{@k@}} \times \left({@q@}\right)^{{@nmk@}}
\]
\[
P(X = {@k@}) = {@ck_val@} \times \frac{@num_p@}{@den_p@} \times \frac{@num_q@}{@den_q@} = {@ans_correct@}
\]
</div>
</div>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>/* Fonction utilitaire */
ri(a,b) := a + floor(rand((b-a+1)));

/* Génération des paramètres */
n: ri(5, 8);
/* On évite k=0 ou k=n pour que les erreurs de puissances soient bien distinctes */
k: 1 + floor(rand((n-1)));
nmk: n - k;

/* On choisit une probabilité p sous forme de fraction simple */
p_pool: [1/2, 1/3, 2/3, 1/4, 3/4, 1/5, 2/5, 3/5, 4/5];
p: p_pool[1 + floor(rand(length(p_pool)))];
q: 1 - p;

/* Valeurs pour la correction */
ck_val: binomial(n, k);
num_p: p^k;
den_p: 1;
num_q: q^nmk;
den_q: 1;

/* Réponse attendue */
ans_correct: (ck_val * p^k * q^nmk);

/* Erreurs classiques */
/* Erreur 1 : A oublié le coefficient binomial C(n,k) */
err_no_ck: p^k * q^nmk;
/* Erreur 2 : A inversé les puissances (p^(n-k) et q^k) */
err_pow_swap: ck_val * p^nmk * q^k;
/* Erreur 3 : A oublié le terme (1-p)^(n-k) */
err_no_1mp: ck_val * p^k;
/* Erreur 4 : A juste écrit le coefficient binomial */
err_only_ck: ck_val;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<div>[[feedback:prt1]]</div>]]></text>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>B({@n@}, {@p@}), P(X={@k@}) = {@ans_correct@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p>Généré avec HéStack V2 | hestack-types:probabilites-binomiale-simple</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans_prob</name>
      <type>algebraic</type>
      <tans>ans_correct</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords>binomial</allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>1</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Réponse correcte</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_prob</sans>
        <tans>ans_correct</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-OK</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #15803d; padding: 10px 14px; background: #f0fdf4; border-radius: 4px;">✅ <strong>Parfait !</strong> La probabilité est correcte et est bien donnée sous forme de fraction irréductible.</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>PRT-NOK</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Erreur : A oublié le coefficient binomial C(n,k)</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_prob</sans>
        <tans>err_no_ck</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-NO-CK</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Oubli du coefficient binomial !</strong> Vous avez calculé \(p^k \times (1-p)^{n-k}\), ce qui correspond à la probabilité d'une seule issue précise (un chemin particulier). Il faut multiplier par le nombre de chemins possibles, c'est-à-dire \(\binom{n}{k}\).</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>PRT-ERR-NOT-NO-CK</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Erreur : Inversion des exposants</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_prob</sans>
        <tans>err_pow_swap</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-POW-SWAP</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Exposants inversés !</strong> Vous avez élevé \(p\) à la puissance \(n-k\) et \(1-p\) à la puissance \(k\). C'est l'inverse : le nombre de succès (\(k\)) est l'exposant de \(p\), et le nombre d'échecs (\(n-k\)) est l'exposant de \(1-p\).</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>PRT-ERR-NOT-POW-SWAP</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Erreur : A oublié le terme (1-p)^(n-k)</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_prob</sans>
        <tans>err_no_1mp</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-NO-1MP</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Il manque la probabilité d'échec !</strong> Vous avez calculé \(\binom{n}{k} \times p^k\). La formule complète impose de multiplier aussi par la probabilité d'avoir des échecs : \((1-p)^{n-k}\).</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>PRT-ERR-CALC</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #dc2626; padding: 10px 14px; background: #fff0f0; border-radius: 4px;">❌ <strong>Incorrect.</strong> Appliquez la formule du cours :<br>\(P(X = k) = \binom{n}{k} \times p^k \times (1-p)^{n-k}\)<br><em>(Attention à bien simplifier votre fraction finale).</em></div>]]></text>
        </falsefeedback>
      </node>
    </prt>
  </question>

</quiz>