<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 396245  -->
  <question type="stack">
    <name>
      <text>Matrices : Déterminant d'une matrice (2×2)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<div style="background: #475569; border-left: 5px solid #334155; border-radius: 0 8px 8px 0; padding: 10px 16px; margin-bottom: 12px; display: flex; align-items: center; gap: 10px; flex-wrap: wrap;"><strong style="font-weight: 800; color: #fff; font-size: .95rem;">Matrices — Déterminant</strong> <span style="background: #334155; color: #fff; padding: 2px 9px; border-radius: 20px; font-size: .78rem; font-weight: bold;">/ 1 pt</span></div>
<p>Soit la matrice :</p>
<div style="text-align: center; margin: 15px 0;">\(M = {@M@}\)</div>
<p><strong>Calculer</strong> le déterminant de \(M\).</p>
<div style="margin: 10px 0;">\(\det(M) = \) [[input:ans_det]] [[validation:ans_det]]</div>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<div style="margin-top: 20px; padding: 15px; background: #f8fafc; border: 1px solid #cbd5e1; border-radius: 8px;">
<div style="font-weight: bold; margin-bottom: 10px;">🔑 Correction détaillée</div>
<div style="font-size: .9rem;">
Pour une matrice \(2 \times 2\), le déterminant est :<br><br>
\[
\det\begin{pmatrix} a & b \\ c & d \end{pmatrix} = a \times d - b \times c
\]<br>
<strong>Application :</strong><br><br>
\[
\det(M) = ({@a11@}) \times ({@a22@}) - ({@a12@}) \times ({@a21@}) = {@a11*a22@} - {@a12*a21@} = {@det_correct@}
\]<br><br>
<em style="color: #64748b;">Méthode mnémotechnique : produit de la <strong>diagonale principale</strong> moins produit de la <strong>diagonale secondaire</strong>.</em>
</div>
</div>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>/* Fonction utilitaire : entier aléatoire entre a et b */
ri(a,b) := a + rand(b - a + 1)$ 
/* Génération des coefficients */
a11: ri(-4,4); a12: ri(-4,4); a21: ri(-4,4); a22: ri(-4,4);

/* Construction de la matrice */
M: matrix([a11, a12], [a21, a22]);

/* Déterminant correct */
det_correct: a11*a22 - a12*a21;

/* Erreurs classiques */
err_trace: a11 + a22;
err_plus: a11*a22 + a12*a21;
err_oppose: a12*a21 - a11*a22;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<div>[[feedback:prt1]]</div>]]></text>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>M={@M@}, det(M)={@det_correct@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p>Généré avec HéStack V2 | hestack-types:matrix-determinant-2x2</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans_det</name>
      <type>numerical</type>
      <tans>det_correct</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Réponse correcte</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_det</sans>
        <tans>det_correct</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-OK</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #15803d; padding: 10px 14px; background: #f0fdf4; border-radius: 4px;">✅ <strong>Parfait !</strong> Le déterminant est bien calculé.</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>PRT-NOK</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Erreur : A calculé la trace</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_det</sans>
        <tans>err_trace</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-TRACE</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Confusion trace / déterminant !</strong> Vous avez calculé la <strong>trace</strong> (\(a_{11} + a_{22}\)), c'est-à-dire la somme des coefficients diagonaux. Le déterminant, lui, utilise un <strong>produit</strong> : \(a_{11} \times a_{22} - a_{12} \times a_{21}\).</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>PRT-ERR-NOT-TRACE</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Erreur : Somme des deux produits au lieu de la différence</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_det</sans>
        <tans>err_plus</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-PLUS</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Erreur de signe !</strong> Vous avez additionné les deux produits (\(a_{11}a_{22} + a_{12}a_{21}\)). Il faut <strong>soustraire</strong> le produit de la diagonale secondaire : \(a_{11}a_{22} \mathbf{-} a_{12}a_{21}\).</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>PRT-ERR-NOT-PLUS</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Erreur : Signe global inversé</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_det</sans>
        <tans>err_oppose</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-OPPO</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Signe inversé !</strong> Vous avez calculé \(a_{12}a_{21} - a_{11}a_{22}\) au lieu de \(a_{11}a_{22} - a_{12}a_{21}\). N'oubliez pas : c'est la <strong>diagonale principale</strong> moins la <strong>diagonale secondaire</strong>.</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>PRT-ERR-CALC</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #dc2626; padding: 10px 14px; background: #fff0f0; border-radius: 4px;">❌ <strong>Incorrect.</strong> Appliquez la formule :<br>\(\det(M) = a_{11} \times a_{22} - a_{12} \times a_{21}\)</div>]]></text>
        </falsefeedback>
      </node>
    </prt>
  </question>

</quiz>