<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 396343  -->
  <question type="stack">
    <name>
      <text>Trigonométrie : Simplifier une expression</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<div style="background: #475569; border-left: 5px solid #334155; border-radius: 0 8px 8px 0; padding: 10px 16px; margin-bottom: 12px; display: flex; align-items: center; gap: 10px; flex-wrap: wrap;"><strong style="font-weight: 800; color: #fff; font-size: .95rem;">Trigonométrie — Simplification</strong> <span style="background: #334155; color: #fff; padding: 2px 9px; border-radius: 20px; font-size: .78rem; font-weight: bold;">/ 1 pt</span></div>
<p>Simplifier au maximum l'expression suivante à l'aide des formules trigonométriques :</p>
<div style="text-align: center; font-size: 1.2rem; margin: 25px 0; background: #f8fafc; padding: 15px; border-radius: 8px; border: 1px solid #e2e8f0;">
\(E = {@expr@}\)
</div>
<p><em>La réponse devra être donnée sous sa forme la plus simple possible (factorisée). Les valeurs approchées sont refusées.</em></p>
<div style="margin: 15px 0 10px 0;">\(E = \) [[input:ans_simp]] [[validation:ans_simp]]</div>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<div style="margin-top: 20px; padding: 15px; background: #f8fafc; border: 1px solid #cbd5e1; border-radius: 8px;">
<div style="font-weight: bold; margin-bottom: 10px;">🔑 Correction détaillée</div>
<div style="font-size: .9rem;">
On utilise les formules de linéarisation (ou formules d'addition) :</div>
<div style="font-size: .9rem; margin-top: 10px;">
\[
\cos(a+b) + \cos(a-b) = 2\cos(a)\cos(b)
\]
\[
\cos(a+b) - \cos(a-b) = -2\sin(a)\sin(b)
\]
\[
\sin(a+b) + \sin(a-b) = 2\sin(a)\cos(b)
\]
\[
\sin(a+b) - \sin(a-b) = 2\cos(a)\sin(b)
\]
</div>
<div style="font-size: .9rem; margin-top: 10px;">
En identifiant les termes dans notre expression, on obtient directement :<br><br>
\[
E = {@ans_correct@}
\]
</div>
</div>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>/* Fonction utilitaire */
ri(a,b) := a + rand(b - a + 1) - 1;

/* Choix aléatoire du modèle d'expression (sans utiliser de liste) */
r_tpl: rand(4);

/* Choix aléatoire de l'angle constant a (parmi les valeurs classiques) */
r_a: rand(3);
if r_a = 1 then a_val: pi/6
else if r_a = 2 then a_val: pi/4
else a_val: pi/3;

/* Génération de l'expression et de la réponse en fonction du modèle */
if r_tpl = 1 then (
    expr: cos(a_val + x) + cos(a_val - x),
    ans_correct: 2*cos(a_val)*cos(x),
    err_wrong_op: -2*sin(a_val)*sin(x)
) else if r_tpl = 2 then (
    expr: cos(a_val + x) - cos(a_val - x),
    ans_correct: -2*sin(a_val)*sin(x),
    err_wrong_op: 2*cos(a_val)*cos(x)
) else if r_tpl = 3 then (
    expr: sin(a_val + x) + sin(a_val - x),
    ans_correct: 2*sin(a_val)*cos(x),
    err_wrong_op: 2*cos(a_val)*sin(x)
) else (
    expr: sin(a_val + x) - sin(a_val - x),
    ans_correct: 2*cos(a_val)*sin(x),
    err_wrong_op: 2*sin(a_val)*cos(x)
);

/* Erreurs classiques */
/* Erreur 1 : A oublié le facteur 2 */
err_no_2: ans_correct / 2;

/* Erreur 2 : Confusion des fonctions dans le résultat (sin au lieu de cos) */
if r_tpl = 1 then err_swap: 2*cos(a_val)*sin(x)
else if r_tpl = 2 then err_swap: -2*cos(a_val)*cos(x)
else if r_tpl = 3 then err_swap: 2*cos(a_val)*cos(x)
else err_swap: 2*sin(a_val)*cos(x);

/* Erreur 3 : A inversé le signe de tout */
err_sign: -ans_correct;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<div>[[feedback:prt1]]</div>]]></text>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>{@expr@} = {@ans_correct@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p>Généré avec HéStack V2 | hestack-types:trigonometrie-simplification</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans_simp</name>
      <type>algebraic</type>
      <tans>ans_correct</tans>
      <boxsize>20</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint></syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords></allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>2</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Réponse correcte</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_simp</sans>
        <tans>ans_correct</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-OK</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #15803d; padding: 10px 14px; background: #f0fdf4; border-radius: 4px;">✅ <strong>Parfait !</strong> L'expression est parfaitement simplifiée.</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>PRT-NOK</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Erreur : Mauvaise formule appliquée</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_simp</sans>
        <tans>err_wrong_op</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-OP</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Mauvaise formule !</strong> Vous avez utilisé une formule d'addition qui ne correspond pas aux fonctions présentes (par exemple une formule avec du sinus alors qu'il y a du cosinus, ou un problème de signe avec la soustraction).</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>PRT-ERR-NOT-OP</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Erreur : Oubli du facteur 2</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_simp</sans>
        <tans>err_no_2</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-NO2</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Oubli du coefficient !</strong> Toutes les formules de linéarisation ont un facteur <strong>2</strong> au numérateur. Votre résultat est exactement la moitié de la réponse attendue.</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>PRT-ERR-NOT-NO2</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Erreur : Confusion sin/cos dans le résultat</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_simp</sans>
        <tans>err_swap</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-SWAP</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Confusion Sinus / Cosinus !</strong> La structure de votre réponse est bonne, mais vous avez interverti les fonctions sinus et cosinus dans le produit final.</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>4</falsenextnode>
        <falseanswernote>PRT-ERR-NOT-SWAP</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description>Erreur : Signe global inversé</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_simp</sans>
        <tans>err_sign</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-SIGN</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Erreur de signe !</strong> L'expression est mathématiquement l'opposée de la bonne réponse. Vérifiez le signe de la formule (surtout pour les différences du type \(\cos(A) - \cos(B)\)).</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>PRT-ERR-CALC</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #dc2626; padding: 10px 14px; background: #fff0f0; border-radius: 4px;">❌ <strong>Incorrect.</strong> Utilisez les formules d'addition ou de linéarisation :<br>\(\cos(a+b) \pm \cos(a-b)\)<br>\(\sin(a+b) \pm \sin(a-b)\)<br><em>(N'oubliez pas le facteur 2 et vérifiez les signes).</em></div>]]></text>
        </falsefeedback>
      </node>
    </prt>
  </question>

</quiz>