<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<!-- question: 396387  -->
  <question type="stack">
    <name>
      <text>Limites : Limite à gauche en un point (Racine carrée)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<div style="background: #475569; border-left: 5px solid #334155; border-radius: 0 8px 8px 0; padding: 10px 16px; margin-bottom: 12px; display: flex; align-items: center; gap: 10px; flex-wrap: wrap;"><strong style="font-weight: 800; color: #fff; font-size: .95rem;">Limites de fonctions — À gauche en \(a\)</strong> <span style="background: #334155; color: #fff; padding: 2px 9px; border-radius: 20px; font-size: .78rem; font-weight: bold;">/ 1 pt</span></div>
<p>Soit \(f\) la fonction définie pour tout réel \(x \leq {@a@}\), \(x \neq {@a@}\), par :</p>
<div style="text-align: center; font-size: 1.2rem; margin: 25px 0; background: #f8fafc; padding: 15px; border-radius: 8px; border: 1px solid #e2e8f0;">
\(f(x) = \frac{{@num@}}{{@den@}}\)
</div>
<p> Calculer la limite de la fonction \(f\) à gauche en {@a@}.</p>
<p><em>Si la limite est infinie, on saisira <strong>inf</strong> ou <strong>-inf</strong>.</em></p>
<div style="margin: 15px 0 10px 0;">\(\lim\limits_{x \to {@a@}^-} f(x) = \) [[input:ans_lim]] [[validation:ans_lim]]</div>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<div style="margin-top: 20px; padding: 15px; background: #f8fafc; border: 1px solid #cbd5e1; border-radius: 8px;">
<div style="font-weight: bold; margin-bottom: 10px;">🔑 Correction détaillée</div>
<div style="font-size: .9rem;">
En remplaçant \(x\) par {@a@}, on obtient \(\frac{\sqrt{0}}{0}\), c'est une <strong>forme indéterminée</strong>.<br><br>
<strong>Le point clé ici :</strong><br>
Puisque l'on calcule la limite <strong>à gauche</strong> (\(x \to {@a@}^-\)), on sait que \(x < {@a@}\), donc :<br>
- \(x - {@a@} < 0\)<br>
- \(a - x > 0\) (donc \(\sqrt{a - {@a@}}\) est bien défini et positif).<br><br>
<strong>Levée de l'indétermination :</strong><br>
On fait apparaître le terme positif \((a - x)\) au dénominateur en factorisant par \(-1\) :<br>
\[
x - {@a@} = -(a - {@a@})
\]
Pour tout \(x < {@a@}\) :
\[
f(x) = \frac{\sqrt{a - {@a@}}}{-(a - {@a@})} = -\frac{\sqrt{a - {@a@}}}{a - {@a@}}
\]
Puisque \(a - x > 0\), on peut simplifier \(\frac{\sqrt{a - x}}{a - x} = \frac{1}{\sqrt{a - x}}\).<br>
On obtient donc : \(f(x) = -\frac{1}{\sqrt{a - {@a@}}}\).<br><br>
<strong>Conclusion :</strong><br>
Lorsque \(x \to {@a@}^-\), le terme \(a - {@a@}\) tend vers \(0^+\) (zéro par valeurs positives).<br>
Ainsi, \(\sqrt{a - {@a@}} \to 0^+\), et son inverse tend vers \(+\infty\). Avec le signe "moins" devant, on obtient :
\[
\lim\limits_{x \to {@a@}^-} f(x) = -\infty
\]
</div>
</div>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* 1. INITIALISATIONS PAR DÉFAUT */
a: 2$ num: sqrt(a-x)$ den: x-a$ func: num/den$ ans_correct: -inf$ txt_lim: "-\infty"$ err_0: 0$ err_sign: inf$ err_a: 999$ 
/* 2. GÉNÉRATION ALÉATOIRE */
r_a: rand(5)$ if r_a = 1 then a: 2$ if r_a = 2 then a: 3$ if r_a = 3 then a: -1$ if r_a = 4 then a: -2$ if r_a = 5 then a: 4$ 
num: sqrt(a - x)$ den: x - a$ func: num/den$ ans_correct: -inf$ txt_lim: "-\infty"$ 
/* 3. ERREURS CLASSIQUES */
err_0: 0$ err_sign: inf$ err_a: a$]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<div>[[feedback:prt1]]</div>]]></text>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>a={@a@}, Limite à gauche = -inf</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p>Généré avec HéStack V2 | hestack-types:limites-point-fini-a-gauche</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed></variantsselectionseed>
    <input>
      <name>ans_lim</name>
      <type>algebraic</type>
      <tans>ans_correct</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>inf ou -inf</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords></forbidwords>
      <allowwords>inf</allowwords>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options></options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Réponse correcte</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_lim</sans>
        <tans>ans_correct</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-OK</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #15803d; padding: 10px 14px; background: #f0fdf4; border-radius: 4px;">✅ <strong>Parfait !</strong> Vous avez bien justifié le signe "moins" grâce au fait que \(x - {@a@} < 0\).</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>PRT-NOK</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Erreur : A répondu +inf (oubli du signe)</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_lim</sans>
        <tans>err_sign</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-SIGN</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Oubli du signe "moins" !</strong> Vous avez probablement simplifié directement \(\frac{\sqrt{a-x}}{x-a}\) en \(\frac{1}{\sqrt{a-x}}\). C'est faux car pour \(x \to {@a@}^-\), on a \(x - {@a@} < 0\). Il faut écrire \(x-a = -(a-x)\) avant de simplifier, ce qui fait apparaître un signe \(-\) devant la fraction.</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>PRT-ERR-NOT-SIGN</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Erreur : A répondu 0</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_lim</sans>
        <tans>err_0</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-0</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Piège de simplification !</strong> Vous avez probablement simplifié \(\frac{\sqrt{a-x}}{x-a}\) par \(\sqrt{x-a}\) au lieu de gérer le signe. Rappelez-vous que la racine carrée est toujours positive, il est impossible que le résultat fasse 0 ici. Il y a une forme indéterminée qu'il faut lever.</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>PRT-ERR-NOT-0</falseanswernote>
        <falsefeedback format="html">
          <text></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Erreur : A répondu la valeur de a</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_lim</sans>
        <tans>err_a</tans>
        <testoptions></testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty></truepenalty>
        <truenextnode>-1</truenextnode>
        <trueanswernote>PRT-ERR-A</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #f97316; padding: 10px 14px; background: #fff7ed; border-radius: 4px;">🚨 <strong>Valeur interdite !</strong> Vous avez donné la valeur {@a@}. Cela correspondrait à un remplacement direct dans la fraction non simplifiée, ce qui est mathématiquement interdit (\(\frac{0}{0}\) n'existe pas). Il faut impérativement transformer l'expression avant de passer à la limite.</div>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty></falsepenalty>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>PRT-ERR-CALC</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<div style="border-left: 4px solid #dc2626; padding: 10px 14px; background: #fff0f0; border-radius: 4px;">❌ <strong>Incorrect.</strong> Méthode pour la limite à gauche :<br>1. Justifiez que \(x < {@a@}\), donc \(x - {@a@} < 0\) et \(a - x > 0\).<br>2. Remplacez \(x - a\) par \(-(a - x)\) au dénominateur.<br>3. Simplifiez \(\frac{\sqrt{a-x}}{-(a-x)}\) pour obtenir \(-\frac{1}{\sqrt{a-x}}\).<br>4. Concluez en utilisant le fait que \(a - {@a@} \to 0^+\).</div>]]></text>
        </falsefeedback>
      </node>
    </prt>
  </question>

</quiz>